<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>10.1. Коришћење и дефинисање функција &#8212; Petlja - Информатика и рачунарство за шести разред</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="10.2. Дефинисање функција" href="10_2_cas.html" />
    <link rel="prev" title="10. лекција" href="index.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Информатика и рачунарство за шести разред';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Информатика и рачунарство за шести разред</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">1. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_cas.html">1.1. Робот Карел и линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_cas.html">1.2. Робот Карел и бројачка петља <cite>for</cite></a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_3_cas.html">1.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">2. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_cas.html">2.1. Робот Карел - условне петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_cas.html">2.2. Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_3_cas.html">2.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">3. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3_1_cas.html">3.1. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">4. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_1_test.html">4.1. Квиз - робот Карел (линијска структура и понављање)</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_2_test.html">4.2. Квиз - робот Карел (условна петља и гранање)</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">5. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_1_cas.html">5.1. Корњача графика - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_2_cas.html">5.2. Корњача графика - понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_3_cas.html">5.3. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../6_lekcija/index.html">6. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_1_cas.html">6.1. Корњача графика - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_2_cas.html">6.2. Корњача графика - угнежђене петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_3_cas.html">6.3. Корњача графика - процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_4_cas.html">6.4. Корњача графика - торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_5_cas.html">6.5. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">7. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_1.html">Тест - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../8_lekcija/index.html">8. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_1_cas.html">8.1. Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_2_cas.html">8.2. Променљиве - имена међурезултата</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_3_cas.html">8.3. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_4_cas.html">8.4. Вежбање - изрази</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../9_lekcija/index.html">9. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_1_cas.html">9.1. Цели и реални бројеви</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_2_cas.html">9.2. Реално и целобројно дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_3_cas.html">9.3. Вежбање - цели и реални бројеви</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">10. лекција</a><ul class="current">
<li class="toctree-l2 current"><a class="current reference internal" href="#">10.1. Коришћење и дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="10_2_cas.html">10.2. Дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="10_3_cas.html">10.3. Вежбање - функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../11_lekcija/index.html">11. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_1.html">11.1. Квиз - Аритметичке операције</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_2.html">11.2.Квиз - целобројно и реално дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_3.html">11.3. Квиз - Функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">12. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_1_cas.html">12.1. Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_2_cas.html">12.2. Комбинација више логичких услова</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../13_lekcija/index.html">13. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_1_cas.html">13.1 Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_2_cas.html">13.2. Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_3_cas.html">13.3. Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../14_lekcija/index.html">14. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_1_test.html">Тест - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_2_test.html">Тест - понављање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">15. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_1_cas.html">15.1. Структуре података: Торке, листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_2_cas.html">15.2. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_3_cas.html">15.3. Скупови</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../16_lekcija/index.html">16. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../16_lekcija/16_1.html">16.1. Квиз - структуре података</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>10.1. Коришћење и дефинисање функција<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<div class="section" id="id2">
<h2>Уграђене функције<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Хајде прво да погледамо наредни видео који је у вези са уграђеним функцијама, а затим детаљно прочитај све што следи и уради све питалице и задатке.</p>

    <div  id="CJLpHTm1OLw" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/CJLpHTm1OLw/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-CJLpHTm1OLw" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-CJLpHTm1OLw">
        </div>
    </div>
    
    
    <div class="section" id="id3">
<h3>Појам функције<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>
<p><strong>Функције</strong> су један од основних појмова математике и
програмирања. Оне представљају пресликавања одређених улазних
вредности (каже се параметара, аргумената) у једну или више излазних
вредности тј. резултата. На пример, функција може да као улазни
параметар добије дужину странице једнакостраничног троугла <span class="math notranslate nohighlight">\(а\)</span> и
да као резултат врати дужину обима тог троугла <span class="math notranslate nohighlight">\(3\cdot a\)</span>. Таква
зависност би се у математици описала као <span class="math notranslate nohighlight">\(f(a) = 3\cdot a\)</span> што
значи да функција коју смо назвали <cite>f</cite> (што је најчешће име за
функцију) на основу датог параметра <span class="math notranslate nohighlight">\(а\)</span> одређује вредност
<span class="math notranslate nohighlight">\(3\cdot a\)</span>. Слично, функција која израчунава обим правоугаоника
на основу дужине његових страница би се у математици описала као
<span class="math notranslate nohighlight">\(g(a, b) = 2\cdot a + 2 \cdot b\)</span>. У овом случају је функцији дат
назив <span class="math notranslate nohighlight">\(g\)</span>, она има два улазна параметра (<span class="math notranslate nohighlight">\(a\)</span> и <span class="math notranslate nohighlight">\(b\)</span>)
и враћа резултат који је одређен формулом <span class="math notranslate nohighlight">\(2\cdot a + 2 \cdot
b\)</span>.</p>
<p>Касније ћемо видети како можемо дефинисати функције у језику Python, а
за почетак ћемо се само позабавити коришћењем неких основних уграђених
функција (оне су већ дефинисане језиком и можемо их слободно користити
тј. позивати у нашим програмима). У наставку ћемо детаљније описати
следеће функције:</p>
<ul class="simple">
<li><p><code class="docutils literal notranslate"><span class="pre">min</span></code>, <code class="docutils literal notranslate"><span class="pre">max</span></code> - минимум и максимум</p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">abs</span></code> - апсолутна вредност</p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">math.pow</span></code>, <code class="docutils literal notranslate"><span class="pre">math.sqrt</span></code> -  степен, квадратни корен</p></li>
<li><p><code class="docutils literal notranslate"><span class="pre">round</span></code>, <code class="docutils literal notranslate"><span class="pre">math.floor</span></code>, <code class="docutils literal notranslate"><span class="pre">math.ceil</span></code> - заокругљивање реалних
бројева</p></li>
</ul>
</div>
<div class="section" id="id4">
<h3>Минимум и максимум<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>У многим задацима потребно је одредити мањи или већи од два дата
броја. Пошто је тај задатак веома чест, програмски језик Python нуди
решење у виду функција <code class="docutils literal notranslate"><span class="pre">min</span></code> и <code class="docutils literal notranslate"><span class="pre">max</span></code>. На пример, вредност израза
<code class="docutils literal notranslate"><span class="pre">min(2,</span> <span class="pre">5)</span></code> је мањи од бројева 2 и 5 тј. број 2 (тај израз
представља позив функције <code class="docutils literal notranslate"><span class="pre">min</span></code> са аргументима <code class="docutils literal notranslate"><span class="pre">2</span></code> и <code class="docutils literal notranslate"><span class="pre">5</span></code>), док
је вредност израза <code class="docutils literal notranslate"><span class="pre">max(2,</span> <span class="pre">5)</span></code> већи од бројева 2 и 5 тј. број 5 (тај
израз представља позив функције <code class="docutils literal notranslate"><span class="pre">max</span></code> са аргументима <code class="docutils literal notranslate"><span class="pre">2</span></code> и <code class="docutils literal notranslate"><span class="pre">5</span></code>).</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_минимум">
            <p>Вредност израза <code class="docutils literal notranslate"><span class="pre">min(24,</span> <span class="pre">17)</span></code> је <input type="text"></p>

        <script type="application/json">
            [[{"number": [17, 17], "feedback": "<p>\u0422\u0430\u0447\u043d\u043e!</p>\n"}, {"regex": "^ *x *$", "regexFlags": "", "feedback": "<p>\u041d\u0435\u0442\u0430\u0447\u043d\u043e. min \u043e\u0437\u043d\u0430\u0447\u0430\u0432\u0430 \u043c\u0430\u045a\u0438 \u0431\u0440\u043e\u0458</p>\n"}]]
        </script>

        </div>
        </div>
            
            <div class="course-box course-box-question course-content">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="mchoice_минимум">
            <p>Q-1: Која је вредност израза <code class="docutils literal notranslate"><span class="pre">min(min(5,</span> <span class="pre">2),</span> <span class="pre">3)</span></code>?</p>

            <li data-component="answer" data-correct id="mchoice_минимум_opt_a">2</li><li data-component="feedback" id="mchoice_минимум_opt_a">Тачно!</li>
            
            <li data-component="answer"  id="mchoice_минимум_opt_b">3</li><li data-component="feedback" id="mchoice_минимум_opt_b">min означава мањи од два броја.</li>
            
            <li data-component="answer"  id="mchoice_минимум_opt_c">5</li><li data-component="feedback" id="mchoice_минимум_opt_c">min означава мањи од два броја.</li>
            

            </ul>
            </div>
            <p>Минимум и максимум неколико бројева могуће је израчунати тако што се
сви наведу као параметри функција <cite>min</cite> тј. <cite>max</cite>. На пример, вредност
израза <code class="docutils literal notranslate"><span class="pre">max(3,</span> <span class="pre">2,</span> <span class="pre">5,</span> <span class="pre">4)</span></code> је <code class="docutils literal notranslate"><span class="pre">5</span></code>.</p>
<p>Покажимо сада неколико задатака у чијем нам решењу ове функције могу
помоћи.</p>
<div class="section" id="id5">
<h4>Већа оцена<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h4>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Љубица је једног дана добила оцену из биологије и географије. Која
је већа оцена коју је Љубица добила тог дана?</p>

    </p></div></div>

<div data-childcomponent="већа_оцена" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=већа_оцена data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
biologija  = int(input("biologija: "))
geografija = int(input("geografija: "))
print("veća ocena: ", 0)  # ispravi ovu liniju
</textarea>
</div>
</div>
<div class="section" id="id7">
<h4>Највећи број поена у игрици<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h4>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Асмир, Снежа и Мица су играли игрицу на рачунару. Колики је највећи
број поена (high-score) који је освоји неко од њих.</p>

    </p></div></div>

<div data-childcomponent="највећи_број_поена" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=највећи_број_поена data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="asmir sneza mica najveci"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
asmir = int(input("Koliko je poena osvojio Asmir: "))
sneza = int(input("Koliko je poena osvojila Sneža: "))
mica  = int(input("Koliko je poena osvojila Mica: "))
# -*- acsection: main -*-
najveci = 0 # ispravi ovu liniju
# -*- acsection: after-main -*-
print("Највећи број поена је: ", najveci)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for asmir, sneza, mica in [(3542, 2388, 4123), (3562, 4221, 1002), (7428, 2842, 9413)]:
          self.assertEqual(acMainSection(asmir = asmir, sneza = sneza, mica = mica)["najveci"],max(asmir, sneza, mica),"Ако су редом освајали %s, %s и %s поена, тада је највећи број поена %s." % (asmir, sneza, mica, max(asmir, sneza, mica)))
myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="id9">
<h3>Апсолутна вредност<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h3>
<p>Још једна веома корисна функција коју си упознао/упознала у математици
је апсолутна вредност. Апсолутном вредношћу се одређује одступање
броја од нуле.  На пример, број 4 одступа од нуле за 4, док број -5
одступа од нуле за 5 јединица. Дакле апсолутна вредност броја
<span class="math notranslate nohighlight">\(x\)</span>, која се, подсетимо се, обележава са <span class="math notranslate nohighlight">\(|x|\)</span>, једнака је
самом броју <span class="math notranslate nohighlight">\(x\)</span> ако је <span class="math notranslate nohighlight">\(x \geq 0\)</span> тј.  броју <span class="math notranslate nohighlight">\(-x\)</span>,
ако је <span class="math notranslate nohighlight">\(x &lt; 0\)</span>.</p>
<p>У језику Python апсолутну вредност можемо израчунати помоћу функције
<code class="docutils literal notranslate"><span class="pre">abs</span></code>. Тако је вредност израза <code class="docutils literal notranslate"><span class="pre">abs(5)</span></code> једнака <code class="docutils literal notranslate"><span class="pre">5</span></code>, док је
вредност израза <code class="docutils literal notranslate"><span class="pre">abs(-3)</span></code> једнака 3.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_апсолутна_вредност">
            <p>Вредност израза <code class="docutils literal notranslate"><span class="pre">abs(-11.2)</span></code> је <input type="text"></p>

        <script type="application/json">
            [[{"number": [11.2, 11.2], "feedback": "<p>\u0422\u0430\u0447\u043d\u043e!</p>\n"}, {"regex": "^ *x *$", "regexFlags": "", "feedback": "<p>abs \u043e\u0437\u043d\u0430\u0447\u0430\u0432\u0430 \u0430\u043f\u0441\u043e\u043b\u0443\u0442\u043d\u0443 \u0432\u0440\u0435\u0434\u043d\u043e\u0441\u0442 \u0431\u0440\u043e\u0458\u0430 \u0438 \u043e\u043d\u0430 \u043d\u0435 \u043c\u043e\u0436\u0435 \u0431\u0438\u0442\u0438 \u043d\u0435\u0433\u0430\u0442\u0438\u0432\u043d\u0430</p>\n"}]]
        </script>

        </div>
        </div>
            <p>Једна од најзначајнијих примена апсолутне вредности је да се израчуна
удаљеност бројева, без обзира на њихов међусобни однос. Наиме,
растојање између бројева <span class="math notranslate nohighlight">\(x_1\)</span> и <span class="math notranslate nohighlight">\(x_2\)</span> је једнако
апсолутној вредности њихове разлике тј. вредности <span class="math notranslate nohighlight">\(|x_1 - x_2|\)</span>,
без обзира да ли је <span class="math notranslate nohighlight">\(x_1 \geq x_2\)</span> или је <span class="math notranslate nohighlight">\(x_1 &lt;
x_2\)</span>. Размотримо наредни задатак.</p>
<div class="section" id="id10">
<h4>Удаљеност спратова<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h4>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Cпратови у једној згради су обележени бројевима од -2 до 10
(бројеви -2 и -1 означавају два нивоа испод земље, 0 означава
приземље, док остали бројеви означавају спратове изнад земље). Ако
се знају спратови на којима се налазе два другара, израчунај колико
су спратова удаљени.</p>

    </p></div></div>
<p>Пошто не знамо да ли лифт иде на горе или на доле тј. да ли је полазни
спрат мањи или већи од долазног, растојање између спратова можемо
израчунати тако што израчунамо апсолутну вредност разлике између тих
спратова. На пример, ако би први другар био на трећем спрату, а други
у првом подруму, исправи програм који рачуна то растојање.</p>

<div data-childcomponent="лифт_апсолутна_вредност" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=лифт_апсолутна_вредност data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
sprat1 = 3
sprat2 = -1
rastojanje = abs(sprat1 - sprat2)
print(rastojanje)
</textarea>
</div>
<p>Наравно, уместо бројева 3 и -1 могуће је стављати и друге, а могуће је
ове бројеве учитати приликом покретања програма. Пробај са неколико
својих тест-примера, а пробај и са тест-примерима које смо ти ми
припремили.</p>

<div data-childcomponent="лифт_апсолутна_вредност_input" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=лифт_апсолутна_вредност_input data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="sprat1 sprat2 rastojanje"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
sprat1 = int(input("На ком се спрату налази први другар?"))
sprat2 = int(input("На ком се спрату налази други другар?"))
# -*- acsection: main -*-
rastojanje = 0 # ispravi svoj red
# -*- acsection: after-main -*-
print(rastojanje)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for sprat1, sprat2 in [(8, 4), (3, 6), (-1, 8), (4, -2)]:
          self.assertEqual(acMainSection(sprat1 = sprat1, sprat2 = sprat2)["rastojanje"],abs(sprat1-sprat2),"Растојање између спратова %s и %s је %s." % (sprat1, sprat2, abs(sprat1-sprat2)))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id12">
<h4>Менхетн растојање<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h4>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Менхетн, део града Њујорка је организован у авеније у правцу
север-југ и улице у правцу isток-запад. Размак између две улице је
80m, а између две авеније је 275m. Ако се Том налази на углу улице
<span class="math notranslate nohighlight">\(u_1\)</span> и авеније <span class="math notranslate nohighlight">\(a_1\)</span> и жели да стигне на угао улице
<span class="math notranslate nohighlight">\(u_2\)</span> и авеније <span class="math notranslate nohighlight">\(a_2\)</span>, колико ће метара морати да
пређе.</p>

    </p></div></div>
<a class="reference internal image-reference" href="../_images/manhattan_distance.png"><img alt="../_images/manhattan_distance.png" class="align-center" src="../_images/manhattan_distance.png" style="width: 500px;" /></a>
<p>Том има више начина да стигне са једног на друго место (може да иде
цик-цак, на разне начине), међутим, пређено растојање је исто као када
би прво ишао улицом <span class="math notranslate nohighlight">\(u_1\)</span> све док не дође до угла са авенијом
<span class="math notranslate nohighlight">\(a_2\)</span>, а затим да се креће авенијом <span class="math notranslate nohighlight">\(а_2\)</span> све док не дође
до угла са улицом <span class="math notranslate nohighlight">\(u_2\)</span>. Дакле, потребно је израчунати растојање
између авенија <span class="math notranslate nohighlight">\(a_1\)</span> и <span class="math notranslate nohighlight">\(a_2\)</span> (да би се оно добило у
метрима потребно је помножити апсолутну разлику између њихових редних
бројева размаком између суседних авенија) и на то додати растојање
између улица <span class="math notranslate nohighlight">\(u_1\)</span> и <span class="math notranslate nohighlight">\(u_2\)</span> (да би се оно добило у метрима
потребно је помножити апсолутну разлику између њихових редних бројева
размаком између суседних улица).</p>
<p>Исправи наредни код тако да коректно израчуна пређени пут (наравно, програм
треба да ради и када се улазни подаци промене или учитају са улаза).</p>

<div data-childcomponent="менхетн" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=менхетн data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="ulica1 avenija1 ulica2 avenija2 rastojanje razmak_ulica razmak_avenija"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
razmak_ulica = 80
razmak_avenija = 275
ulica1 = 51
avenija1 = 6
ulica2 = 58
avenija2 = 3
# -*- acsection: main -*-
rastojanje = abs(avenija1 - avenija2) * 0 + \
             0 * razmak_ulica
# -*- acsection: after-main -*-
print(rastojanje)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
       for (ulica1, avenija1, ulica2, avenija2, razmak_ulica, razmak_avenija, rastojanje) in [(3, 5, 8, 4, 80, 275, 675), (1, 7, 2, 4, 80, 275, 905), (9, 4, 11, 2, 80, 275, 710), (4, 8, 1, 5, 80, 275, 1065)]:
          self.assertEqual((acMainSection(ulica1 = ulica1, avenija1 = avenija1, ulica2 = ulica2, avenija2 = avenija2, razmak_ulica = razmak_ulica, razmak_avenija = razmak_avenija)["rastojanje"]),  rastojanje , "Растојање између тачака (%s, %s) и (%s, %s) је %s." % (ulica1, avenija1, ulica2, avenija2, rastojanje))
myTests().main()
</textarea>
</div>
<p>Приметимо да је формула у претходном примеру била веома дугачка и
проценили смо да је прегледније да је одштампамо кроз више редова. Да
бисмо нагласили да се нека наредба наставља и у следећој линији на
крај линије стављамо симбол <code class="docutils literal notranslate"><span class="pre">\</span></code>.</p>

    <div data-component="reveal" id="менхетн_решење1"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="менхетн_решење2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=менхетн_решење2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
ulica1 = 51
avenija1 = 6
ulica2 = 58
avenija2 = 3
razmak_ulica = 80
razmak_avenija = 275
rastojanje = abs(avenija1 - avenija2) * razmak_avenija + \
             abs(ulica1 - ulica2) * razmak_ulica
print(rastojanje)
</textarea>
</div>

    </div>
    </div>
</div>
</div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - 10. лекција' data-toggle="tooltip"><a href="index.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - 10.2. Дефинисање функција' data-toggle="tooltip"><a href="10_2_cas.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - 10. лекција' data-toggle="tooltip"><a href="index.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - 10.2. Дефинисање функција' data-toggle="tooltip"><a href="10_2_cas.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>