<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>10.1. Израчунавања - задаци за вежбање</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/qchoice.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/associations.css" />
    <link rel="stylesheet" type="text/css" href="../_static/editor.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelChat.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/associations.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/editor.js"></script>
    <script type="text/javascript" src="../_static/jszip.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Информатика и рачунарство за шести разред';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
    <div id="usefulLinksToggle" class="useful-links-toggle">
        <a class="dropdown-toggle dropdown-toggle-burger" role="button" href="#" id="usefulLinksDropdown" data-toggle="dropdown" aria-haspopup="true" aria-expanded="false">
            <i class="fas fa-bars fa-2x"></i>
        </a>
        <div class="dropdown-menu dropdown-nav-links shadow-lg" aria-labelledby="usefulLinksDropdown" x-placement="bottom-start" style="position: absolute; transform: translate3d(-55px, 32px, 0px); top: 0px; left: 0px; will-change: transform;">
          <div class="links-title"> Корисни линкови</div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/prirucnik-python-gim" target="_blank">Петљин приручник за Пајтон</a></div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/pygame-prirucnik-gim" target="_blank">Петљин приручник за библиотеку Пајгејм</a></div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/python-zbirka-pitanja-gim" target="_blank">Збирка кратких питања из програмирања за први разред</a></div>
          
        </div>
    </div>

    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Програмирање уз језик Python</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-1_lekcija>
                    <h5>1. Робот Карел - Линијски програми и бројачка петља</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../1_lekcija/1_1_cas.html><div class="studioLecture" id=activity-1_1_cas><i class="fas fa-file-alt activity-icon"></i>1.1. Линијски програми</div></a>
                <a href=../1_lekcija/1_2_cas.html><div class="studioLecture" id=activity-1_2_cas><i class="fas fa-file-alt activity-icon"></i>1.2 Вежбање</div></a>
                <a href=../1_lekcija/1_3_cas.html><div class="studioLecture" id=activity-1_3_cas><i class="fas fa-file-alt activity-icon"></i>1.3. Бројачка петља for</div></a>
                <a href=../1_lekcija/1_4_cas.html><div class="studioLecture" id=activity-1_4_cas><i class="far fa-question-circle activity-icon"></i>1.4. Квиз</div></a>
                <a href=../1_lekcija/1_5_cas.html><div class="studioLecture" id=activity-1_5_cas><i class="fas fa-file-alt activity-icon"></i>1.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-2_lekcija>
                    <h5>2. Робот Карел - Условна петља и гранање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../2_lekcija/2_1_cas.html><div class="studioLecture" id=activity-2_1_cas><i class="fas fa-file-alt activity-icon"></i>2.1. Условне петље</div></a>
                <a href=../2_lekcija/2_2_cas.html><div class="studioLecture" id=activity-2_2_cas><i class="fas fa-file-alt activity-icon"></i>2.2. Вежбање</div></a>
                <a href=../2_lekcija/2_3_cas.html><div class="studioLecture" id=activity-2_3_cas><i class="fas fa-file-alt activity-icon"></i>2.3. Гранање</div></a>
                <a href=../2_lekcija/2_4_cas.html><div class="studioLecture" id=activity-2_4_cas><i class="far fa-question-circle activity-icon"></i>2.4. Квиз</div></a>
                <a href=../2_lekcija/2_5_cas.html><div class="studioLecture" id=activity-2_5_cas><i class="fas fa-file-alt activity-icon"></i>2.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-3_lekcija>
                    <h5>3. Робот Карел - Задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../3_lekcija/3_1_cas.html><div class="studioLecture" id=activity-3_1_cas><i class="fas fa-file-alt activity-icon"></i>3.1. Робот Карел - задаци за вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-4_lekcija>
                    <h5>4. Корњача графика - Линијски програми и петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../4_lekcija/4_1_cas.html><div class="studioLecture" id=activity-4_1_cas><i class="fas fa-file-alt activity-icon"></i>4.1. Наредбе корњаче</div></a>
                <a href=../4_lekcija/4_2_cas.html><div class="studioLecture" id=activity-4_2_cas><i class="fas fa-file-alt activity-icon"></i>4.2. Линијски програми</div></a>
                <a href=../4_lekcija/4_3_cas.html><div class="studioLecture" id=activity-4_3_cas><i class="fas fa-file-alt activity-icon"></i>4.3. Петље</div></a>
                <a href=../4_lekcija/4_4_cas.html><div class="studioLecture" id=activity-4_4_cas><i class="far fa-question-circle activity-icon"></i>4.4. Квиз</div></a>
                <a href=../4_lekcija/4_5_cas.html><div class="studioLecture" id=activity-4_5_cas><i class="fas fa-file-alt activity-icon"></i>4.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-5_lekcija>
                    <h5>5. Корњача графика - Гранање, угнежђене петље, функције, торке и листе</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../5_lekcija/5_1_cas.html><div class="studioLecture" id=activity-5_1_cas><i class="fas fa-file-alt activity-icon"></i>5.1. Гранање</div></a>
                <a href=../5_lekcija/5_2_cas.html><div class="studioLecture" id=activity-5_2_cas><i class="fas fa-file-alt activity-icon"></i>5.2. Угнежђене петље</div></a>
                <a href=../5_lekcija/5_3_cas.html><div class="studioLecture" id=activity-5_3_cas><i class="fas fa-file-alt activity-icon"></i>5.3. Функције</div></a>
                <a href=../5_lekcija/5_4_cas.html><div class="studioLecture" id=activity-5_4_cas><i class="fas fa-file-alt activity-icon"></i>5.4. Торке и листе</div></a>
                <a href=../5_lekcija/5_5_cas.html><div class="studioLecture" id=activity-5_5_cas><i class="far fa-question-circle activity-icon"></i>5.5. Квиз</div></a>
                <a href=../5_lekcija/5_6_cas.html><div class="studioLecture" id=activity-5_6_cas><i class="fas fa-file-alt activity-icon"></i>5.6. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-6_lekcija>
                    <h5>6. Корњача графика - Задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../6_lekcija/6_1_cas.html><div class="studioLecture" id=activity-6_1_cas><i class="fas fa-file-alt activity-icon"></i>6.1. Корњача графика - задаци вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-7_lekcija>
                    <h5>7. Израчунавања - Изрази, оператори, променљиве</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../7_lekcija/7_1_cas.html><div class="studioLecture" id=activity-7_1_cas><i class="fas fa-file-alt activity-icon"></i>7.1. Основне аритметичке операције и примена</div></a>
                <a href=../7_lekcija/7_2_cas.html><div class="studioLecture" id=activity-7_2_cas><i class="fas fa-file-alt activity-icon"></i>7.2. Променљиве - имена придружена подацима</div></a>
                <a href=../7_lekcija/7_3_cas.html><div class="studioLecture" id=activity-7_3_cas><i class="fas fa-file-alt activity-icon"></i>7.3. Тестирање и дебаговање програма</div></a>
                <a href=../7_lekcija/7_4_cas.html><div class="studioLecture" id=activity-7_4_cas><i class="far fa-question-circle activity-icon"></i>7.4. Квиз</div></a>
                <a href=../7_lekcija/7_5_cas.html><div class="studioLecture" id=activity-7_5_cas><i class="fas fa-file-alt activity-icon"></i>7.5. Вежбање</div></a>
                <a href=../7_lekcija/7.6.%20Израчунавања%20-%20додатни%20задаци.html><div class="studioLecture" id=activity-7.6.%20Израчунавања%20-%20додатни%20задаци><i class="far fa-question-circle activity-icon"></i>7.6. Израчунавања - додатни задаци</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-8_lekcija>
                    <h5>8. Израчунавања - Цели и реални бројеви, дељење</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../8_lekcija/8_1_cas.html><div class="studioLecture" id=activity-8_1_cas><i class="fas fa-file-alt activity-icon"></i>8.1. Цели и реални бројеви</div></a>
                <a href=../8_lekcija/8_2_cas.html><div class="studioLecture" id=activity-8_2_cas><i class="fas fa-file-alt activity-icon"></i>8.2. Реално и целобројно дељење</div></a>
                <a href=../8_lekcija/8_3_cas.html><div class="studioLecture" id=activity-8_3_cas><i class="fas fa-file-alt activity-icon"></i>8.3. Позициони запис бројева</div></a>
                <a href=../8_lekcija/8_4_cas.html><div class="studioLecture" id=activity-8_4_cas><i class="far fa-question-circle activity-icon"></i>8.4. Квиз</div></a>
                <a href=../8_lekcija/8_5_cas.html><div class="studioLecture" id=activity-8_5_cas><i class="fas fa-file-alt activity-icon"></i>8.5. Вежбање</div></a>
                <a href=../8_lekcija/8.6.%20Цели%20и%20реални%20бројеви,%20дељење%20-%20додатни%20задци.html><div class="studioLecture" id=activity-8.6.%20Цели%20и%20реални%20бројеви,%20дељење%20-%20додатни%20задци><i class="far fa-question-circle activity-icon"></i>8.6. Цели и реални бројеви, дељење - додатни задци</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-9_lekcija>
                    <h5>9. Израчунавања - Коришћење и дефинисање функција</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../9_lekcija/9_1_cas.html><div class="studioLecture" id=activity-9_1_cas><i class="fas fa-file-alt activity-icon"></i>9.1. Коришћење функција</div></a>
                <a href=../9_lekcija/9_2_cas.html><div class="studioLecture" id=activity-9_2_cas><i class="fas fa-file-alt activity-icon"></i>9.2. Дефинисање функција</div></a>
                <a href=../9_lekcija/9_3_cas.html><div class="studioLecture" id=activity-9_3_cas><i class="far fa-question-circle activity-icon"></i>9.3. Квиз</div></a>
                <a href=../9_lekcija/9_4_cas.html><div class="studioLecture" id=activity-9_4_cas><i class="fas fa-file-alt activity-icon"></i>9.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-10_lekcija>
                    <h5>10. Израчунавања - Задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../10_lekcija/10_1_cas.html><div class="studioLecture" id=activity-10_1_cas><i class="fas fa-file-alt activity-icon"></i>10.1. Израчунавања - задаци за вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-11_lekcija>
                    <h5>11. Гранање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../11_lekcija/11_1_cas.html><div class="studioLecture" id=activity-11_1_cas><i class="fas fa-file-alt activity-icon"></i>11.1. Гранање</div></a>
                <a href=../11_lekcija/11_2_cas.html><div class="studioLecture" id=activity-11_2_cas><i class="fas fa-file-alt activity-icon"></i>11.2. Комбинација више логичких услова</div></a>
                <a href=../11_lekcija/11_3_cas.html><div class="studioLecture" id=activity-11_3_cas><i class="fas fa-file-alt activity-icon"></i>11.3. Конструкција if-elif-else</div></a>
                <a href=../11_lekcija/11_4_cas.html><div class="studioLecture" id=activity-11_4_cas><i class="far fa-question-circle activity-icon"></i>11.4. Квиз</div></a>
                <a href=../11_lekcija/11_5_cas.html><div class="studioLecture" id=activity-11_5_cas><i class="fas fa-file-alt activity-icon"></i>11.5. Вежбање</div></a>
                <a href=../11_lekcija/11.6.%20Гранање%20-%20додатни%20задтак.html><div class="studioLecture" id=activity-11.6.%20Гранање%20-%20додатни%20задтак><i class="far fa-question-circle activity-icon"></i>11.6. Гранање - додатни задтак</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-12_lekcija>
                    <h5>12. Понављање - петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../12_lekcija/12_1_cas.html><div class="studioLecture" id=activity-12_1_cas><i class="fas fa-file-alt activity-icon"></i>12.1. Понављање - петља for</div></a>
                <a href=../12_lekcija/12_2_cas.html><div class="studioLecture" id=activity-12_2_cas><i class="fas fa-file-alt activity-icon"></i>12.2. Понављање - условна петља while</div></a>
                <a href=../12_lekcija/12_3_cas.html><div class="studioLecture" id=activity-12_3_cas><i class="fas fa-file-alt activity-icon"></i>12.3. Угњежђене петље</div></a>
                <a href=../12_lekcija/12_4_cas.html><div class="studioLecture" id=activity-12_4_cas><i class="far fa-question-circle activity-icon"></i>12.4. Квиз</div></a>
                <a href=../12_lekcija/12_5_cas.html><div class="studioLecture" id=activity-12_5_cas><i class="fas fa-file-alt activity-icon"></i>12.5. Вежбање</div></a>
                <a href=../12_lekcija/12.6.%20Петље%20-%20додатни%20задтак.html><div class="studioLecture" id=activity-12.6.%20Петље%20-%20додатни%20задтак><i class="far fa-question-circle activity-icon"></i>12.6. Петље - додатни задтак</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-13_lekcija>
                    <h5>13. Контрола тока - Задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../13_lekcija/13_1_cas.html><div class="studioLecture" id=activity-13_1_cas><i class="fas fa-file-alt activity-icon"></i>13.1. Контрола тока - задаци вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-14_lekcija>
                    <h5>14. Структуре података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../14_lekcija/14_1_cas.html><div class="studioLecture" id=activity-14_1_cas><i class="fas fa-file-alt activity-icon"></i>14.1. Торке и листе</div></a>
                <a href=../14_lekcija/14_2_cas.html><div class="studioLecture" id=activity-14_2_cas><i class="fas fa-file-alt activity-icon"></i>14.2. Функције за рад са торкама и листама</div></a>
                <a href=../14_lekcija/14_3_cas.html><div class="studioLecture" id=activity-14_3_cas><i class="fas fa-file-alt activity-icon"></i>14.3. Ниске</div></a>
                <a href=../14_lekcija/14_4_cas.html><div class="studioLecture" id=activity-14_4_cas><i class="fas fa-file-alt activity-icon"></i>14.4. Скупови и речници</div></a>
                <a href=../14_lekcija/14_5_cas.html><div class="studioLecture" id=activity-14_5_cas><i class="far fa-question-circle activity-icon"></i>14.5. Квиз</div></a>
                <a href=../14_lekcija/14_6_cas.html><div class="studioLecture" id=activity-14_6_cas><i class="fas fa-file-alt activity-icon"></i>14.6. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-15_lekcija>
                    <h5>15. Примена петље у структурама података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../15_lekcija/15_1_cas.html><div class="studioLecture" id=activity-15_1_cas><i class="fas fa-file-alt activity-icon"></i>15.1. Примена петље у торци, листи и ниски</div></a>
                <a href=../15_lekcija/15_2_cas.html><div class="studioLecture" id=activity-15_2_cas><i class="fas fa-file-alt activity-icon"></i>15.2. Учитавање и обрада листе</div></a>
                <a href=../15_lekcija/15_3_cas.html><div class="studioLecture" id=activity-15_3_cas><i class="fas fa-file-alt activity-icon"></i>15.3. Примена петље у речнику</div></a>
                <a href=../15_lekcija/15_4_cas.html><div class="studioLecture" id=activity-15_4_cas><i class="far fa-question-circle activity-icon"></i>15.4. Квиз</div></a>
                <a href=../15_lekcija/15_5_cas.html><div class="studioLecture" id=activity-15_5_cas><i class="fas fa-file-alt activity-icon"></i>15.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-16_lekcija>
                    <h5>16. Откривање и исправљање грешака и коришћење дебагера</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../16_lekcija/16_1_cas.html><div class="studioLecture" id=activity-16_1_cas><i class="fas fa-file-alt activity-icon"></i>16.1. Дебаговање на рачунару у програму IDLE</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>10.1. Израчунавања - задаци за вежбање<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<div class="section" id="id2">
<h2>Обућа<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>У једној фабрици обуће прошле године је произведено 356.000 пари
ципела и 247.923 пара чизама. Напиши програм који израчунава колико
је пари ципела произведено више него пари чизама, колико је укупно
произведено пари обуће, и колико је укупно произведено комада
обуће. Напиши програм тако да исправно ради и када су бројеви дати у
задатку другачији.</p>

    </div></div>

<div data-childcomponent="обућа" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=обућа data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
pari_cipela = 356000
pari_cizama = 247923
# dopuni rešenje
</textarea>
</div>
<p>Провери да ли је програм дошао до правог решења.</p>

        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_обућа1">
            <p>Колико има више пари ципела него пари чизама?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^108077$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041e\u0434\u0443\u0437\u043c\u0438 \u0431\u0440\u043e\u0458 \u043f\u0430\u0440\u0438 \u0447\u0438\u0437\u0430\u043c\u0430 \u043e\u0434 \u0431\u0440\u043e\u0458\u0430 \u0446\u0438\u043f\u0435\u043b\u0430</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_обућа2">
            <p>Колико укупно има пари ципела и чизама?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^603923$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u0421\u0430\u0431\u0435\u0440\u0438 \u0431\u0440\u043e\u0458 \u043f\u0430\u0440\u0438 \u0446\u0438\u043f\u0435\u043b\u0430 \u0438 \u0447\u0438\u0437\u0430\u043c\u0430</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_обућа3">
            <p>Колико укупно има комада обуће?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^1207846$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043c\u043d\u043e\u0436\u0438 \u0431\u0440\u043e\u0458 \u043f\u0430\u0440\u0438 \u0446\u0438\u043f\u0435\u043b\u0430 \u0438 \u0447\u0438\u0437\u0430\u043c\u0430 \u0441\u0430 \u0434\u0432\u0430</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div data-component="reveal" id="обућа_решење"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="обућа_решење_код" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=обућа_решење_код data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
pari_cipela = 356000
pari_cizama = 247923
vise_pari_cipela = pari_cipela - pari_cizama
print(vise_pari_cipela)
ukupno_pari_obuce = pari_cipela + pari_cizama
print(ukupno_pari_obuce)
ukupno_komada_obuce = ukupno_pari_obuce * 2
print(ukupno_komada_obuce)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id5">
<h2>Карте у возу<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Кондуктер је једног дана прегледао карте са бројевима од 276.898 до
277.236. Напиши програм који израчунава колико путника се тог дана
возило возом?</p>

    </div></div>

<div data-childcomponent="број_путника" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=број_путника data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
od_karte = 276898
do_karte = 277236
broj_karata = 0     # popravi ovaj red
print(broj_karata)
</textarea>
</div>
<p>Провери да ли је програм дошао до правог решења.</p>

        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_број_путника">
            <p>Колико се путника возило возом?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^339$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440!</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u0420\u0430\u0437\u043c\u0438\u0441\u043b\u0438 \u043a\u0430\u043a\u043e \u0431\u0438 \u0441\u0435 \u0440\u0435\u0437\u0443\u043b\u0442\u0430\u0442 \u0438\u0437\u0440\u0430\u0447\u0443\u043d\u0430\u043e \u0434\u0430 \u0441\u0443 \u043f\u0440\u043e\u0434\u0430\u0442\u0435 \u043a\u0430\u0440\u0442\u0435 \u043e\u0434 \u0431\u0440\u043e\u0458\u0430 3 \u0434\u043e \u0431\u0440\u043e\u0458\u0430 7. \u041f\u0430\u0437\u0438 \u0434\u0430 \u043d\u0435 \u043f\u043e\u0433\u0440\u0435\u0448\u0438\u0448 \u0437\u0430 \u0458\u0435\u0434\u0430\u043d.</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div data-component="reveal" id="број_путника_решење"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="број_путника_решење_код" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=број_путника_решење_код data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
od_karte = 277236
do_karte = 276898
broj_karata = od_karte - do_karte + 1
print(broj_karata)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id8">
<h2>Аутомобили<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Продавница половних аутомобила има 317 возила. Од тог броја 95 је
потпуно исправно. Колико је новца потребно за поправку преосталих аутомобила,
ако поправка једног аута кошта 9756 динара.</p>

    </div></div>

<div data-childcomponent="аутомобили" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=аутомобили data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="automobili ispravni popravka_jednog popravka_svih"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
automobili = 317
ispravni = 95
popravka_jednog = 9756
# -*- acsection: main -*-
# dopuni ovde kod
# -*- acsection: after-main -*-
print(popravka_svih)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for automobili, ispravni, popravka_jednog, popravka_svih in [(442, 118, 11324, 3668976), (92, 83, 955, 8595)]:
          self.assertEqual(acMainSection(automobili = automobili, ispravni = ispravni, popravka_jednog = popravka_jednog)["popravka_svih"],popravka_svih,"Ако је у продавници било %s аутомобила, од којих је %s било исправно, а ако поправка једног кошта %s динара, поправка укупно кошта %s динара." % (automobili, ispravni, popravka_jednog, popravka_svih))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="аутомобили_решење_reveal"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="аутомобили_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=аутомобили_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
automobili = 317
ispravni = 95
neispravni = automobili - ispravni
popravka_jednog = 9756
popravka_svih   = neispravni * popravka_jednog
print(popravka_svih)
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id11">
<h2>Сличице<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Урош и Растко имају заједно 317 сличица. Растко и Павле имају
заједно 295 сличица. Павле и Урош имају заједно 212 сличица. Напиши
програм који израчунава и исписује колико сличица има свако од њих.</p>

    </div></div>

<div data-childcomponent="сличице_у_паровима" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=сличице_у_паровима data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
uros_i_rastko = 317
rastko_i_pavle = 295
pavle_i_uros = 212
# dopuni resenje
print(pavle, uros, rastko)
</textarea>
</div>
<p>Провери да ли је програм дошао до правог решења.</p>

        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_сличице_1">
            <p>Колико сличица има Урош?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^117$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_сличице_2">
            <p>Колико сличица има Растко?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^200$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            
        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_сличице_3">
            <p>Колико сличица има Павле?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^95$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            
    <div data-component="reveal" id="сличице_помоћ"   data-showtitle="Прикажи малу помоћ" data-hidetitle="Сакриј малу помоћ">
    <p>Израчунај прво колико сви заједно имају сличица.</p>

    </div>
    
    <div data-component="reveal" id="сличице_решење"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="сличице_у_паровима_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=сличице_у_паровима_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
uros_i_rastko = 317
rastko_i_pavle = 295
pavle_i_uros = 212
zajedno = (uros_i_rastko + rastko_i_pavle + pavle_i_uros) // 2
pavle = zajedno - uros_i_rastko
uros = zajedno - rastko_i_pavle
rastko = zajedno - pavle_i_uros
print(pavle, uros, rastko)
</textarea>
</div>

    </div>
    <p>Провери да ли твој програм и даље ради када се измене вредности
улазних података. Израчунај колико Павле има сличица.</p>

<div data-childcomponent="slicice_test" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=slicice_test data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="rastko_i_pavle uros_i_rastko pavle_i_uros pavle"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
uros_i_rastko = 317
rastko_i_pavle = 295
pavle_i_uros = 212
# -*- acsection: main -*-
# dopuni ovde kod
# -*- acsection: after-main -*-
print(pavle)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for rastko_i_pavle, uros_i_rastko, pavle_i_uros, pavle in [(73, 94, 81, 30), (103, 75, 94, 61)]:
          self.assertEqual(acMainSection(rastko_i_pavle = rastko_i_pavle, uros_i_rastko = uros_i_rastko, pavle_i_uros = pavle_i_uros)["pavle"],pavle,"Ако Растко и Павле имају %s сличица, Урош и Растко %s сличица, а Павле и Урош %s сличица, тада Павле има %s сличица." % (rastko_i_pavle, uros_i_rastko, pavle_i_uros, pavle))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id14">
<h2>Израз<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Број 345 увећај 76 пута, па добијени резултат увећај за количник
бројева 141126 и 258 (први број јесте дељив са другим). Који се
резултат добија? Задатак реши једним изразом (немој да рачунаш
пешке).</p>

    </div></div>

<div data-childcomponent="операције" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=операције data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print()   # у заграде упиши израз
</textarea>
</div>
<p>Провери да ли је програм дошао до правог решења.</p>

        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_израз3">
            <p>Које је решење?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^26767$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="id16">
<h2>Оловке<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Сташа је 16 оловака платила 3408 динара. Колико ће платити 21 оловку?</p>

    </div></div>

<div data-childcomponent="оловке_израз" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=оловке_израз data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print() # у заграде упиши израз
</textarea>
</div>
<p>Провери да ли је програм дошао до правог решења.</p>

        <div class="course-box course-box-question course-content petlja-problem-box fitb-question">
        <div class="image-background"></div>
        <div class="petlja-problem-box-icon-holder"> </div>
        <img src="../_static/img/fitb-img.svg" class="petlja-problem-image  fitb-image" /> 
        <img src="../_static/img/qchoice-img.svg" class="petlja-problem-image-quiz  qchoice-image" /> 
        
        <div data-component="fillintheblank" id="fill_оловке">
            <p>Колико кошта 21 оловка?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^4473$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440!</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="id18">
<h2>Куповина<a class="headerlink" href="#id18" title="Permalink to this headline">¶</a></h2>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Ана има 8460 динара. Петину новца је потрошила у књижари, а трећину
укупне суме у самопослузи. Колико новца је остало Ани?</p>

    </div></div>

<div data-childcomponent="књижара_самопослуга" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=књижара_самопослуга data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="novac ostalo"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
novac  = 8460
# -*- acsection: main -*-
# dopuni ovde kod
# -*- acsection: after-main -*-
print(ostalo)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for novac, ostalo in [(1110, 518), (1380, 644)]:
          self.assertEqual(acMainSection(novac = novac)["ostalo"],ostalo,"Ако је Ана имала %s динара, остало јој је %s динара." % (novac, ostalo))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="књижара_самопослуга_решење_reveal"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="књижара_самопослуга_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=књижара_самопослуга_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
novac = 8460
knjizara = novac/5
samoposluga = novac/3
ostalo = novac - (knjizara + samoposluga)
print(ostalo)
</textarea>
</div>
<p>Решење смо могли добити и помоћу једног израза.</p>

<div data-childcomponent="књижара_самопослуга_израз" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=књижара_самопослуга_израз data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(8460 - (8460 / 5 + 8460 / 3))
</textarea>
</div>

    </div>
    </div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходна лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>

<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>


<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>