<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>Ниске &#8212; Petlja - Информатика и рачунарство за шести разред</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Вежбање" href="7_4_cas.html" />
    <link rel="prev" title="Променљиве - имена међурезултата" href="7_2_cas.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Информатика и рачунарство за шести разред';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Информатика и рачунарство за шести разред</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">Прва лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_cas.html">Робот Карел и линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_cas.html">Понављање - бројачка петља <cite>for</cite></a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_3_cas.html">Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_test.html">Тест - линијска структура</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_test.html">Тест - понављање, бројачка петља for</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">Друга лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_cas.html">Робот Карел - условне петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_cas.html">Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_3_cas.html">Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_test.html">Тест - условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_test.html">Tест - гранање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">Трећа лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3_1_cas.html">Робот Карел - вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">Четврта лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_1_cas.html">Корњача графика - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_2_cas.html">Корњача графика - понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">Пета лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_1_cas.html">Корњача графика - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_2_cas.html">Корњача графика - угнежђене петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_3_cas.html">Корњача графика - процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_4_cas.html">Корњача графика - торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_5_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">Седма лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="7_1_cas.html">Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="7_2_cas.html">Променљиве - имена међурезултата</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="7_4_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../8_lekcija/index.html">Осма лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_1_cas.html">Цели и реални бројеви</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_2_cas.html">Реално и целобројно дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../9_lekcija/index.html">Девета лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_1_cas.html">Коришћење и дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_2_cas.html">Дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../11_lekcija/index.html">Једанаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_1_cas.html">Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_2_cas.html">Конструкција elif</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">Дванаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_1_cas.html">Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_2_cas.html">Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_3_cas.html">Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../14_lekcija/index.html">Четрнаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_1_cas.html">Структуре података: Торке, листе, ниске, скупови, речници</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_2_cas.html">Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_2_cas.html#id11">Врсте реченица</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_3_cas.html">Скупови</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_4_cas.html">Речници</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">Петнаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_1_cas.html">Обрада структура података коришћењем петљи</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>Ниске<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Приметили смо да смо приликом учитавања (наредбе <code class="docutils literal notranslate"><span class="pre">input</span></code>) или исписа
(наредбе <code class="docutils literal notranslate"><span class="pre">print</span></code>) текст који се дословно приказује кориснику
наводили између наводника. Такав текст назива се <strong>ниска</strong> или
<strong>стринг</strong> (од енглеске речи string која значи ниска). На пример,
ниске су <code class="docutils literal notranslate"><span class="pre">&quot;Zdravo&quot;</span></code> или <code class="docutils literal notranslate"><span class="pre">&quot;Programski</span> <span class="pre">jezik</span> <span class="pre">Python.&quot;</span></code>.  Уместо
двоструких равноправно се могу користити и једноструки наводници (на
пример, <code class="docutils literal notranslate"><span class="pre">'Zdravo'</span></code>), међутим, да бисмо вас мање збуњивали, ми ћемо
увек користити двоструке наводнике.</p>
<p>Текст (па и ниске) се састоји од слова (малих и великих), цифара,
размака, интерпункцијских знакова (на пример тачака, зареза, упитника,
узвичника) и слично. Све те знакове једним именом називамо
<strong>карактери</strong>. Неки програмски језици подржавају само веома узак скуп
карактера (од слова је могуће користити само слова енглеске абецеде),
међутим, програмски језик Python3 користи широк скуп карактера који
обухвата и све карактере потребне за писање на већини језика света,
укључујући и слова ћириличног и латиничног писма која се користе у
српском језику.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Поменути основни скуп карактера довољан само за запис текста на
енглеском језику назива се ASCII, док се овај шири скуп карактера
назива Unicode.</p>

    </p></div></div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>У језику Python 3 могуће је и имена променљивих написати ћирилицом,
међутим, то некада може довести до проблема (ако се, на пример,
едитор текста који се користи да се програм откуца не подеси
адекватно), тако да ћемо за сваки случај имена променљивих увек
писати латиницом, без коришћења српских слова (š, ž, č, …).</p>

    </p></div></div>

    <div class="rst-level rst-level-2">

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Ако текст садржи знаке наводника или неке друге специјалне
карактере потребно је те карактере обележити косим цртама. На
пример, <code class="docutils literal notranslate"><span class="pre">&quot;Rekao</span> <span class="pre">je:</span> <span class="pre">\&quot;Zdravo,</span> <span class="pre">svima\&quot;.</span></code> Ове косе црте се не
исписују приликом извршавања програма и штампања ниски.</p>

<div data-childcomponent="escape_sequence" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=escape_sequence data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print("Rekao je: \"Zdravo, svima\".")
</textarea>
</div>

    </p></div></div>

    </div>
<p>Текст може бити сачуван и у променљивама. На пример,</p>

<div data-childcomponent="поздрави_перу" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=поздрави_перу data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
ime = "Pera Perić"
print("Zdravo, ti se zoveš", ime)
</textarea>
</div>
<p>Текст можемо учитати и од корисника, након покретања програма. За то
можемо употребити функцију <code class="docutils literal notranslate"><span class="pre">input</span></code>. На пример, наредни програм пита
корисника како се зове, а онда га поздрави.</p>

<div data-childcomponent="поздрави_корисника" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=поздрави_корисника data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
ime = input("Unesi svoje ime: ")
print("Zdravo, ti se zoveš", ime)
</textarea>
</div>
<p>Приметимо да се приликом учитавања бројева заправо учитава текст
(помоћу <code class="docutils literal notranslate"><span class="pre">input</span></code>), а онда се из учитаног текста одређују бројеви који
су тим текстом представљени. Таква претварања (конверзије) текста у
број врше се помоћу <code class="docutils literal notranslate"><span class="pre">int</span></code> и <code class="docutils literal notranslate"><span class="pre">float</span></code>.</p>
<p>Радом са текстом и нискама ћемо се много детаљније бавити током наредних
часова.</p>
<div class="section" id="id4">
<h2>Аутоматско тестирање програма<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h2>
<p>У програмирању је веома битно да програми које напишеш увек дају тачне
резултате. Најчешћи начин да се провери коректност програма је да се
он покрене на већем броју унапред припремљених тест-примера (у којима
су дате различите вредности улазних параметара) и да се провери да ли
на њима даје очекиване резултате. Ако програм на њима исправно ради,
веома вероватно је да ће исправно радити и на свим другим
примерима. Трудили смо се да у наставку омогућимо аутоматско тестирање
на што већем броју задатака. Један такав је и наредни.</p>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Марко је прочитао књигу за три дана. Првог дана је прочитао 100
страна. Другог дана је прочитао 17 страна више него првог, а трећег
два пута више него другог. Колико та књига има страна? Напиши
програм који то одређује, а ради исправно и када је број страна
прочитаних првог дана другачији од 100.</p>

    </p></div></div>
<p>Програм је у тренутном облику неисправан (увек исписује
нулу). Притисни дугме <strong>Тестирај програм</strong> и видећеш да резултати који
се добију нису једнаки оним очекиваним и приказаће ти се поља црвене
боје, што указује на грешке.</p>

<div data-childcomponent="читање" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=читање data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="prvi_dan ukupno"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
prvi_dan  = 100
# -*- acsection: main -*-
drugi_dan = 0      # popravi ovaj red
treci_dan = 0      # popravi ovaj red
ukupno = 0         # popravi ovaj red
# -*- acsection: after-main -*-
print(ukupno)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for prvi_dan, ukupno in [(84, 387), (122, 539), (153, 663)]:
          self.assertEqual(acMainSection(prvi_dan = prvi_dan)["ukupno"],ukupno,"Ако је први дан прочитао %s стране, број страна књиге је %s." % (prvi_dan, ukupno))
myTests().main()
</textarea>
</div>
<p>Поправи претходни програм тако што ћеш изменити само оне редове који
су обележени - друге редове не дирај. Када завршиш поново притисни
дугме <strong>Тестирај програм</strong>. Ако је све како треба, приказаће ти се три
поља зелене боје. Наиме, када притиснеш то дугме програм се тестира на
основу неколико унапред припремљених тестова. Програм се неколико пута
пушта из почетка (за сваки тест по једном), али тако да се уместо
вредности <code class="docutils literal notranslate"><span class="pre">100</span></code> променљивој <code class="docutils literal notranslate"><span class="pre">prvi_dan</span></code> додељују неке друге
вредности (у овом конкретном примеру покушава се са вредностима
<code class="docutils literal notranslate"><span class="pre">84</span></code>, <code class="docutils literal notranslate"><span class="pre">122</span></code> и <code class="docutils literal notranslate"><span class="pre">153</span></code>). Након тога програм који си
написао/написала израчунава укупан број страна и резултат који је твој
програм израчунао се пореди са стварним бројем страна књиге за тај
тест (за улаз <code class="docutils literal notranslate"><span class="pre">84</span></code> стварни број страна је <code class="docutils literal notranslate"><span class="pre">387</span></code>, за улаз <code class="docutils literal notranslate"><span class="pre">122</span></code>
стварни број страна је <code class="docutils literal notranslate"><span class="pre">539</span></code>, а за улаз <code class="docutils literal notranslate"><span class="pre">153</span></code> стварни број страна
је <code class="docutils literal notranslate"><span class="pre">663</span></code>). Ако се та два броја поклапају, на том случају за ту почетну
вредност се приказује зелена боја, а ако не, приказује се црвена боја.</p>
</div>
<div class="section" id="id6">
<h2>Извршавање корак по корак<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h2>
<p>Једна корисна техника која помаже да се боље разуме начин рада неког
програма, али и да се уоче и исправе евентуалне грешке у програму је
то да се програм извршава корак по корак, тј. да се извршава једна по
једна његова наредба, при чему се након сваке извршене наредбе прати
вредност променљивих у програму.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Извршавање програма корак по корак уз праћење вредности свих
међурезултата назива се <strong>дебаговање</strong> (требљење од бубица) и јако
је важна техника за откривање грешака у програмима. У старим
рачунарима који су били велики као читава соба, дешавало се да
мољци и сличне бубе уђу у рачунар и проузрокују неки квар. Од тада
се све грешке у програмима називају <strong>багови</strong> тј. <strong>бубице</strong>.</p>

    </p></div></div>
<p>Илуструјмо ову могућност на сајту Петље кроз наредни, једноставни,
пример.</p>
<div class="section" id="id7">
<h3>Воћњак са јабукама<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Пера је засадио 380 стабала јабуке. Ђура је засадио 142 стабла јабука
више од Пере, а Мика је засадио два пута више од Пере. Колико су
стабала засадили заједно?</p>

    </p></div></div>

<div data-childcomponent="јабуке" class="course-box course-box-problem course-content">
<p>Исправи следећи програм тако да исправно израчунава колико су
стабала засадили заједно (програм треба исправно да ради и ако се
број стабала које је засадио Пера промени).</p>

<textarea data-component="activecode" id=јабуке data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="pera zajedno"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
pera = 380
# -*- acsection: main -*-
djura = ???
mika = ???
zajedno = pera + djura + mika
# -*- acsection: after-main -*-
print(zajedno)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for pera, zajedno in [(100, 542), (200, 942)]:
          self.assertEqual(acMainSection(pera = pera)["zajedno"],zajedno,"Ако је Пера засадио %s стабала, заједно су засадили %s стабала." % (pera, zajedno))
myTests().main()
</textarea>
</div>
<p>Притисни сада дугме <strong>Корак по корак</strong>. Оно ти пружа могућност да
програм извршаваш корак по корак. Дугметом <em>Forward</em> извршаваш
наредну наредбу (ону обележену црвеном стрелицом). Наредба која је
претходно извршена обележена је светло-плавом стрелицом. У делу
<em>Frames</em> можеш видети вредности свих до сада израчунатих резултата,
док у прозору тога можеш видети излаз програма (резултате одштампане
наредбом <code class="docutils literal notranslate"><span class="pre">print</span></code>).</p>
<p>И у овом задатку је омогућено аутоматско тестирање, па дугметом
<strong>Тестирај програм</strong> провери да ли је твоје решење тачно.</p>
</div>
</div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - Променљиве - имена међурезултата' data-toggle="tooltip"><a href="7_2_cas.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - Вежбање' data-toggle="tooltip"><a href="7_4_cas.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - Променљиве - имена међурезултата' data-toggle="tooltip"><a href="7_2_cas.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - Вежбање' data-toggle="tooltip"><a href="7_4_cas.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>