<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>8.4. Вежбање - изрази &#8212; Petlja - Информатика и рачунарство за шести разред</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="9. лекција" href="../9_lekcija/index.html" />
    <link rel="prev" title="8.3. Ниске" href="8_3_cas.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Информатика и рачунарство за шести разред';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Информатика и рачунарство за шести разред</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">1. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_cas.html">1.1. Робот Карел и линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_cas.html">1.2. Робот Карел и бројачка петља <cite>for</cite></a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_3_cas.html">1.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">2. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_cas.html">2.1. Робот Карел - условне петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_cas.html">2.2. Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_3_cas.html">2.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">3. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3_1_cas.html">3.1. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">4. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_1_test.html">4.1. Квиз - робот Карел (линијска структура и понављање)</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_2_test.html">4.2. Квиз - робот Карел (условна петља и гранање)</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">5. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_1_cas.html">5.1. Корњача графика - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_2_cas.html">5.2. Корњача графика - понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_3_cas.html">5.3. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../6_lekcija/index.html">6. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_1_cas.html">6.1. Корњача графика - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_2_cas.html">6.2. Корњача графика - угнежђене петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_3_cas.html">6.3. Корњача графика - процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_4_cas.html">6.4. Корњача графика - торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_5_cas.html">6.5. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">7. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_1.html">Тест - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">8. лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="8_1_cas.html">8.1. Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="8_2_cas.html">8.2. Променљиве - имена међурезултата</a></li>
<li class="toctree-l2"><a class="reference internal" href="8_3_cas.html">8.3. Ниске</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">8.4. Вежбање - изрази</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../9_lekcija/index.html">9. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_1_cas.html">9.1. Цели и реални бројеви</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_2_cas.html">9.2. Реално и целобројно дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_3_cas.html">9.3. Вежбање - цели и реални бројеви</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../10_lekcija/index.html">10. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_1_cas.html">10.1. Коришћење и дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_2_cas.html">10.2. Дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_3_cas.html">10.3. Вежбање - функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../11_lekcija/index.html">11. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_1.html">11.1. Квиз - Аритметичке операције</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_2.html">11.2.Квиз - целобројно и реално дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_3.html">11.3. Квиз - Функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">12. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_1_cas.html">12.1. Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_2_cas.html">12.2. Комбинација више логичких услова</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../13_lekcija/index.html">13. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_1_cas.html">13.1 Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_2_cas.html">13.2. Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_3_cas.html">13.3. Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../14_lekcija/index.html">14. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_1_test.html">Тест - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_2_test.html">Тест - понављање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">15. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_1_cas.html">15.1. Структуре података: Торке, листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_2_cas.html">15.2. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_3_cas.html">15.3. Скупови</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../16_lekcija/index.html">16. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../16_lekcija/16_1.html">16.1. Квиз - структуре података</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>8.4. Вежбање - изрази<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Ова лекција није много тешка, али је прилично обимна, па је сасвим
могуће да неки од приказаних задатака нисте стигли да проанализирате
током часа. Ако има таквих задатака, уради их сада, у склопу домаћег
задатка. Након тога уради и наредне задатке. Додатне задаткеза вежбу
можеш пронаћи <a class="reference external" href="IzracunavanjeZadaci.html">овде</a>.</p>
<div class="section" id="id3">
<h2>Израз<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Збир бројева 23765 и 7825 умањи 45 пута, па добијени број повећај
за 1609. Колики је резултат?  Задатак реши једним изразом (немој да
рачунаш пешке).</p>

    </p></div></div>

<div data-childcomponent="израз_2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=израз_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print() # у заграде упиши израз
</textarea>
</div>
<p>Провери да ли је твој програм израчунао тачно решење.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_израз2">
            <p>Колико је решење?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^2311$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041f\u043e\u043a\u0443\u0448\u0430\u0458 \u043f\u043e\u043d\u043e\u0432\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            </div>
<div class="section" id="id5">
<h2>Тркачи<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Васа је прешао 2347 метара. Воја 987 метара више од Васе, а Милош два
пута више од Воје. Колико су метара укупно прешли?</p>

    </p></div></div>

<div data-childcomponent="три_тркача" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=три_тркача data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="vasa ukupno"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
vasa  = 2347
# -*- acsection: main -*-
# dopuni ovde kod
# -*- acsection: after-main -*-
print(ukupno)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for vasa, ukupno in [(2462, 12809), (773, 6053)]:
          self.assertEqual(acMainSection(vasa = vasa)["ukupno"],ukupno,"Ако је Васа претрчао %s метара, укупно су претрчали %s метара." % (vasa, ukupno))
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="тркачи_решење_reveal"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="три_тркача_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=три_тркача_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
vasa  = 2347
voja  = vasa + 987
milos = voja * 2
ukupno = vasa + voja + milos
print(ukupno)
</textarea>
</div>

    </div>
    <p>Технике које смо до сада научили довољне су нам да бисмо решили велики број
математичких задатака. Размотримо неколико.</p>
</div>
<div class="section" id="id8">
<h2>Године маме и тате<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Милица има 4 године, њена мама има 7 пута више година него она, а њен
тата има 8 пута више година него она. Колико је година Миличин тата старији
од њене маме?</p>

    </p></div></div>

        <div class="course-box course-box-problem course-content" style="max-width: none;">
        <pre data-component="parsons" id="godine"      >
        <span data-question> Q-85: Поређај делове кода тако да представљају исправно решење овог задатка.</span>milica = 4
---
mama = 7 * milica
tata = 8 * milica
---
razlika = tata - mama
---
print(razlika)
        </pre>
        </div>
    </div>
<div class="section" id="id9">
<h2>Река Морава<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Велика Морава је дугачка 185km и настаје од Јужне Мораве, која је
90km дужа, и Западне Мораве, која је 123km дужа од ње. Колика је
укупна дужина ове три реке?</p>

    </p></div></div>

<div data-childcomponent="морава" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=морава data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
velika_morava = 185
juzna_morava = velika_morava + 90
zapadna_morava = velika_morava + 123
ukupno = ??? # ispravi ovaj red
print(ukupno)
</textarea>
</div>
</div>
<div class="section" id="id11">
<h2>Немањићи<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Стефан Немањић је постао краљ Србије 1217 и владао је 11
година. После њега је Радослав владао до 1234. године, па Владислав,
који је владао 9 година и предао престо брату Урошу Првом, који је
владао до 1276. У којим временским периодима су владали ови српски
краљеви?</p>

    </p></div></div>

<div data-childcomponent="немањићи" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=немањићи data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
Stefan_pocetak = 1217
Stefan_kraj = 1217 + 11
Radoslav_pocetak = Stefan_kraj
Radoslav_kraj = 1234
Vladislav_pocetak = 0
Vladislav_kraj = 0
Uros_pocetak = 0
Uros_kraj = 0
print("Стефан:", Stefan_pocetak, "-", Stefan_kraj)
print("Радослав:", Radoslav_pocetak, "-", Radoslav_kraj)
print("Владислав:", Vladislav_pocetak, "-", Vladislav_kraj)
print("Урош:", Uros_pocetak, "-", Uros_kraj)
</textarea>
</div>
<p>Исправи претходни програм тако да исправно израчуна периоде у којима
су владали краљеви. Ако све урадиш како треба добићеш следеће резултате:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="n">Стефан</span><span class="p">:</span> <span class="mi">1217</span> <span class="o">-</span> <span class="mi">1228</span>
<span class="n">Радослав</span><span class="p">:</span> <span class="mi">1228</span> <span class="o">-</span> <span class="mi">1234</span>
<span class="n">Владислав</span><span class="p">:</span> <span class="mi">1234</span> <span class="o">-</span> <span class="mi">1243</span>
<span class="n">Урош</span><span class="p">:</span> <span class="mi">1243</span> <span class="o">-</span> <span class="mi">1276</span>
</pre></div>
</div>
</div>
<div class="section" id="id13">
<h2>Једначина<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h2>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Напиши програм који израчунава који број треба додати броју 123780
да се добије број 321732.</p>

    </p></div></div>
<p>Нажалост, Python не може директно да решава једначине. Ти мораш да
напишеш израз којим се непозната вредност израчунава на основу
познатих, а онда ти он може помоћи у рачунању.</p>

<div data-childcomponent="непознати_сабирак" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=непознати_сабирак data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
prvi_sabirak = 123780
zbir = 321732
drugi_sabirak = 0    # popravi resenje
print(drugi_sabirak)
</textarea>
</div>
<p>Провери да ли је твој програм израчунао тачно решење.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_једначина">
            <p>Колико је решење?</p>
<input type="text">
        <script type="application/json">
            [[{"regex": "^ *^197952$ *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u0430\u043d \u043e\u0434\u0433\u043e\u0432\u043e\u0440</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041e\u0434 \u0437\u0431\u0438\u0440\u0430 \u043e\u0434\u0443\u0437\u043c\u0438 \u043f\u043e\u0437\u043d\u0430\u0442\u0438 \u0441\u0430\u0431\u0438\u0440\u0430\u043a\u201d</p>\n"}]]
        </script>

        </div>
        </div>
            <p>Ако у решењу нису коришћене вредности, већ само називи променљивих,
програм би требало да исправно решава задатке и за друге
бројеве. Тестирај га на тест-примерима које смо припремили.</p>

<div data-childcomponent="непознати_сабирак_тест" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=непознати_сабирак_тест data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="prvi_sabirak zbir drugi_sabirak"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
prvi_sabirak = 123780
zbir = 321732
# -*- acsection: main -*-
drugi_sabirak = 0    # popravi resenje
# -*- acsection: after-main -*-
print(drugi_sabirak)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for prvi_sabirak, zbir, drugi_sabirak in [(100, 230, 130), (200, 942, 742)]:
          self.assertEqual(acMainSection(prvi_sabirak = prvi_sabirak, zbir = zbir)["drugi_sabirak"],drugi_sabirak,"Ако је једначина %s + x = %s, тада је x = %s." % (prvi_sabirak, zbir, drugi_sabirak))
myTests().main()
</textarea>
</div>
</div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - 8.3. Ниске' data-toggle="tooltip"><a href="8_3_cas.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - 9. лекција' data-toggle="tooltip"><a href="../9_lekcija/index.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - 8.3. Ниске' data-toggle="tooltip"><a href="8_3_cas.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - 9. лекција' data-toggle="tooltip"><a href="../9_lekcija/index.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>