10.1. Коришћење и дефинисање функција
#####################################

Уграђене функције
=================

Хајде прво да погледамо наредни видео који је у вези са уграђеним функцијама, а затим детаљно прочитај све што следи и уради све питалице и задатке.

.. ytpopup:: CJLpHTm1OLw
    :width: 735
    :height: 415
    :align: center

Појам функције
--------------
	   
**Функције** су један од основних појмова математике и
програмирања. Оне представљају пресликавања одређених улазних
вредности (каже се параметара, аргумената) у једну или више излазних
вредности тј. резултата. На пример, функција може да као улазни
параметар добије дужину странице једнакостраничног троугла :math:`а` и
да као резултат врати дужину обима тог троугла :math:`3\cdot a`. Таква
зависност би се у математици описала као :math:`f(a) = 3\cdot a` што
значи да функција коју смо назвали `f` (што је најчешће име за
функцију) на основу датог параметра :math:`а` одређује вредност
:math:`3\cdot a`. Слично, функција која израчунава обим правоугаоника
на основу дужине његових страница би се у математици описала као
:math:`g(a, b) = 2\cdot a + 2 \cdot b`. У овом случају је функцији дат
назив :math:`g`, она има два улазна параметра (:math:`a` и :math:`b`)
и враћа резултат који је одређен формулом :math:`2\cdot a + 2 \cdot
b`.

Касније ћемо видети како можемо дефинисати функције у језику Python, а
за почетак ћемо се само позабавити коришћењем неких основних уграђених
функција (оне су већ дефинисане језиком и можемо их слободно користити
тј. позивати у нашим програмима). У наставку ћемо детаљније описати
следеће функције:

- ``min``, ``max`` - минимум и максимум
- ``abs`` - апсолутна вредност
- ``math.pow``, ``math.sqrt`` -  степен, квадратни корен
- ``round``, ``math.floor``, ``math.ceil`` - заокругљивање реалних
  бројева
  
Минимум и максимум
------------------

У многим задацима потребно је одредити мањи или већи од два дата
броја. Пошто је тај задатак веома чест, програмски језик Python нуди
решење у виду функција ``min`` и ``max``. На пример, вредност израза
``min(2, 5)`` је мањи од бројева 2 и 5 тј. број 2 (тај израз
представља позив функције ``min`` са аргументима ``2`` и ``5``), док
је вредност израза ``max(2, 5)`` већи од бројева 2 и 5 тј. број 5 (тај
израз представља позив функције ``max`` са аргументима ``2`` и ``5``).

.. fillintheblank:: fill_минимум
		    
      Вредност израза ``min(24, 17)`` је |blank|

      -     :17: Тачно!
            :x: Нетачно. min означава мањи број

.. mchoice:: mchoice_минимум
   :answer_a: 2
   :answer_b: 3
   :answer_c: 5
   :correct: a
   :feedback_a: Тачно!
   :feedback_b: min означава мањи од два броја.
   :feedback_c: min означава мањи од два броја. 
		
   Која је вредност израза ``min(min(5, 2), 3)``?

Минимум и максимум неколико бројева могуће је израчунати тако што се
сви наведу као параметри функција `min` тј. `max`. На пример, вредност
израза ``max(3, 2, 5, 4)`` је ``5``.

Покажимо сада неколико задатака у чијем нам решењу ове функције могу
помоћи.

Већа оцена
''''''''''
.. level:: 1
  
.. questionnote::

   Љубица је једног дана добила оцену из биологије и географије. Која
   је већа оцена коју је Љубица добила тог дана?

.. activecode:: већа_оцена

   biologija  = int(input("biologija: "))
   geografija = int(input("geografija: "))
   print("veća ocena: ", 0)  # ispravi ovu liniju

Највећи број поена у игрици
'''''''''''''''''''''''''''
.. level:: 1
   
.. questionnote::

   Асмир, Снежа и Мица су играли игрицу на рачунару. Колики је највећи
   број поена (high-score) који је освоји неко од њих.

.. activecode:: највећи_број_поена
   :runortest: asmir, sneza, mica, najveci
   :enablecopy:
      
   # -*- acsection: general-init -*-
   # -*- acsection: var-init -*-
   asmir = int(input("Koliko je poena osvojio Asmir: "))
   sneza = int(input("Koliko je poena osvojila Sneža: "))
   mica  = int(input("Koliko je poena osvojila Mica: "))
   # -*- acsection: main -*-
   najveci = 0 # ispravi ovu liniju
   # -*- acsection: after-main -*-
   print("Највећи број поена је: ", najveci)
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):
          for asmir, sneza, mica in [(3542, 2388, 4123), (3562, 4221, 1002), (7428, 2842, 9413)]:
             self.assertEqual(acMainSection(asmir = asmir, sneza = sneza, mica = mica)["najveci"],max(asmir, sneza, mica),"Ако су редом освајали %s, %s и %s поена, тада је највећи број поена %s." % (asmir, sneza, mica, max(asmir, sneza, mica)))
   myTests().main()
   



   
Апсолутна вредност
------------------

Још једна веома корисна функција коју си упознао/упознала у математици
је апсолутна вредност. Апсолутном вредношћу се одређује одступање
броја од нуле.  На пример, број 4 одступа од нуле за 4, док број -5
одступа од нуле за 5 јединица. Дакле апсолутна вредност броја
:math:`x`, која се, подсетимо се, обележава са :math:`|x|`, једнака је
самом броју :math:`x` ако је :math:`x \geq 0` тј.  броју :math:`-x`,
ако је :math:`x < 0`.

У језику Python апсолутну вредност можемо израчунати помоћу функције
``abs``. Тако је вредност израза ``abs(5)`` једнака ``5``, док је
вредност израза ``abs(-3)`` једнака 3.

.. fillintheblank:: fill_апсолутна_вредност
		    
      Вредност израза ``abs(-11.2)`` је |blank|

      -     :11.2: Тачно!
            :x: abs означава апсолутну вредност броја и она не може бити негативна


Једна од најзначајнијих примена апсолутне вредности је да се израчуна
удаљеност бројева, без обзира на њихов међусобни однос. Наиме,
растојање између бројева :math:`x_1` и :math:`x_2` је једнако
апсолутној вредности њихове разлике тј. вредности :math:`|x_1 - x_2|`,
без обзира да ли је :math:`x_1 \geq x_2` или је :math:`x_1 <
x_2`. Размотримо наредни задатак.

Удаљеност спратова
''''''''''''''''''
.. level:: 1

.. questionnote::

   Cпратови у једној згради су обележени бројевима од -2 до 10
   (бројеви -2 и -1 означавају два нивоа испод земље, 0 означава
   приземље, док остали бројеви означавају спратове изнад земље). Ако
   се знају спратови на којима се налазе два другара, израчунај колико
   су спратова удаљени.

Пошто не знамо да ли лифт иде на горе или на доле тј. да ли је полазни
спрат мањи или већи од долазног, растојање између спратова можемо
израчунати тако што израчунамо апсолутну вредност разлике између тих
спратова. На пример, ако би први другар био на трећем спрату, а други
у првом подруму, исправи програм који рачуна то растојање.
   
.. activecode:: лифт_апсолутна_вредност
		
   sprat1 = 3
   sprat2 = -1
   rastojanje = abs(sprat1 - sprat2)
   print(rastojanje)

Наравно, уместо бројева 3 и -1 могуће је стављати и друге, а могуће је
ове бројеве учитати приликом покретања програма. Пробај са неколико
својих тест-примера, а пробај и са тест-примерима које смо ти ми
припремили.

.. activecode:: лифт_апсолутна_вредност_input
   :runortest: sprat1, sprat2, rastojanje
		
   # -*- acsection: general-init -*-
   # -*- acsection: var-init -*-
   sprat1 = int(input("На ком се спрату налази први другар?"))
   sprat2 = int(input("На ком се спрату налази други другар?"))
   # -*- acsection: main -*-
   rastojanje = 0 # ispravi svoj red
   # -*- acsection: after-main -*-
   print(rastojanje)
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):
       def testOne(self):
          for sprat1, sprat2 in [(8, 4), (3, 6), (-1, 8), (4, -2)]:
             self.assertEqual(acMainSection(sprat1 = sprat1, sprat2 = sprat2)["rastojanje"],abs(sprat1-sprat2),"Растојање између спратова %s и %s је %s." % (sprat1, sprat2, abs(sprat1-sprat2)))
   myTests().main()
   

Менхетн растојање
'''''''''''''''''
.. level:: 1
      
.. questionnote::

   Менхетн, део града Њујорка је организован у авеније у правцу
   север-југ и улице у правцу isток-запад. Размак између две улице је
   80m, а између две авеније је 275m. Ако се Том налази на углу улице
   :math:`u_1` и авеније :math:`a_1` и жели да стигне на угао улице
   :math:`u_2` и авеније :math:`a_2`, колико ће метара морати да
   пређе.

.. image:: ../../_images/manhattan_distance.png
   :width: 500px   
   :align: center

Том има више начина да стигне са једног на друго место (може да иде
цик-цак, на разне начине), међутим, пређено растојање је исто као када
би прво ишао улицом :math:`u_1` све док не дође до угла са авенијом
:math:`a_2`, а затим да се креће авенијом :math:`а_2` све док не дође
до угла са улицом :math:`u_2`. Дакле, потребно је израчунати растојање
између авенија :math:`a_1` и :math:`a_2` (да би се оно добило у
метрима потребно је помножити апсолутну разлику између њихових редних
бројева размаком између суседних авенија) и на то додати растојање
између улица :math:`u_1` и :math:`u_2` (да би се оно добило у метрима
потребно је помножити апсолутну разлику између њихових редних бројева
размаком између суседних улица).
	   
Исправи наредни код тако да коректно израчуна пређени пут (наравно, програм
треба да ради и када се улазни подаци промене или учитају са улаза).
	   
.. activecode:: менхетн
   :runortest: ulica1, avenija1, ulica2, avenija2, rastojanje, razmak_ulica, razmak_avenija

   # -*- acsection: general-init -*-
   # -*- acsection: var-init -*-
   razmak_ulica = 80	
   razmak_avenija = 275
   ulica1 = 51
   avenija1 = 6
   ulica2 = 58
   avenija2 = 3
   # -*- acsection: main -*-
   rastojanje = abs(avenija1 - avenija2) * 0 + \
                0 * razmak_ulica
   # -*- acsection: after-main -*-
   print(rastojanje)
   ====
   from unittest.gui import TestCaseGui
   class myTests(TestCaseGui):

       def testOne(self):
          for (ulica1, avenija1, ulica2, avenija2, razmak_ulica, razmak_avenija, rastojanje) in [(3, 5, 8, 4, 80, 275, 675), (1, 7, 2, 4, 80, 275, 905), (9, 4, 11, 2, 80, 275, 710), (4, 8, 1, 5, 80, 275, 1065)]:
             self.assertEqual((acMainSection(ulica1 = ulica1, avenija1 = avenija1, ulica2 = ulica2, avenija2 = avenija2, razmak_ulica = razmak_ulica, razmak_avenija = razmak_avenija)["rastojanje"]),  rastojanje , "Растојање између тачака (%s, %s) и (%s, %s) је %s." % (ulica1, avenija1, ulica2, avenija2, rastojanje))
   myTests().main()
   

Приметимо да је формула у претходном примеру била веома дугачка и
проценили смо да је прегледније да је одштампамо кроз више редова. Да
бисмо нагласили да се нека наредба наставља и у следећој линији на
крај линије стављамо симбол ``\``.

.. reveal:: менхетн_решење1
   :showtitle: Прикажи решење
   :hidetitle: Сакриј решење

   .. activecode:: менхетн_решење2

      ulica1 = 51
      avenija1 = 6
      ulica2 = 58
      avenija2 = 3
      razmak_ulica = 80		
      razmak_avenija = 275
      rastojanje = abs(avenija1 - avenija2) * razmak_avenija + \
                   abs(ulica1 - ulica2) * razmak_ulica
      print(rastojanje)


