<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>11.5. Вежбање</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/simanim.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pycode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/p5js.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dbDirective.css" />
    <link rel="stylesheet" type="text/css" href="../_static/associations.css" />
    <link rel="stylesheet" type="text/css" href="../_static/regex-check.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/blockly/blockly_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/blocks_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/python_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/msg-sr.js"></script>
    <script type="text/javascript" src="../_static/blockpy/utilities.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_errors.js"></script>
    <script type="text/javascript" src="../_static/blockpy/ast_node_visitor.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter.js"></script>
    <script type="text/javascript" src="../_static/blockpy/pytifa.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_to_blockly.js"></script>
    <script type="text/javascript" src="../_static/blockpy/imported.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/class.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comment.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/dict.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/if.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/io.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/lists.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/sets.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/loops.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/parking.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/tuple.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/turtles.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/text.js"></script>
    <script type="text/javascript" src="../_static/blockpy-modal.js"></script>
    <script type="text/javascript" src="../_static/simanim.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/pyodide/v0.16.1/full/pyodide.js"></script>
    <script type="text/javascript" src="../_static/pycode.js"></script>
    <script type="text/javascript" src="//toolness.github.io/p5.js-widget/p5-widget.js"></script>
    <script type="text/javascript" src="../_static/p5js.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/dbDirective.js"></script>
    <script type="text/javascript" src="../_static/sql.js"></script>
    <script type="text/javascript" src="../_static/associations.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/regex-check.js"></script>
    <script type="text/javascript" src="../_static/regex-check-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/regex-check-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Информатика и рачунарство за шести разред';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
    <div id="usefulLinksToggle" class="useful-links-toggle">
        <a class="dropdown-toggle dropdown-toggle-burger" role="button" href="#" id="usefulLinksDropdown" data-toggle="dropdown" aria-haspopup="true" aria-expanded="false">
            <i class="fas fa-bars fa-2x"></i>
        </a>
        <div class="dropdown-menu dropdown-nav-links shadow-lg" aria-labelledby="usefulLinksDropdown" x-placement="bottom-start" style="position: absolute; transform: translate3d(-55px, 32px, 0px); top: 0px; left: 0px; will-change: transform;">
          <div class="links-title"> Корисни линкови</div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/prirucnik-python-gim" target="_blank">Петљин приручник за Пајтон</a></div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/pygame-prirucnik-gim" target="_blank">Петљин приручник за библиотеку Пајгејм</a></div>
          <div class="useful-link">
            <a href="https://petlja.org/biblioteka/r/kursevi/python-zbirka-pitanja-gim" target="_blank">Збирка кратких питања из програмирања за први разред</a></div>
          
        </div>
    </div>

    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Програмирање у Пајтону</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-1_lekcija>
                    <h5>1. Робот Карел - линијски програми и бројачка петља</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../1_lekcija/1_1_cas.html><div class="studioLecture" id=activity-1_1_cas><i class="fas fa-file-alt activity-icon"></i>1.1. Линијски програми</div></a>
                <a href=../1_lekcija/1_2_cas.html><div class="studioLecture" id=activity-1_2_cas><i class="fas fa-file-alt activity-icon"></i>1.2. Бројачка петља for</div></a>
                <a href=../1_lekcija/1_3_cas.html><div class="studioLecture" id=activity-1_3_cas><i class="far fa-question-circle activity-icon"></i>1.3. Квиз</div></a>
                <a href=../1_lekcija/1_4_cas.html><div class="studioLecture" id=activity-1_4_cas><i class="fas fa-file-alt activity-icon"></i>1.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-2_lekcija>
                    <h5>2. Робот Карел - условна петља и гранање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../2_lekcija/2_1_cas.html><div class="studioLecture" id=activity-2_1_cas><i class="fas fa-file-alt activity-icon"></i>2.1. Условне петље</div></a>
                <a href=../2_lekcija/2_2_cas.html><div class="studioLecture" id=activity-2_2_cas><i class="fas fa-file-alt activity-icon"></i>2.2. Гранање</div></a>
                <a href=../2_lekcija/2_3_cas.html><div class="studioLecture" id=activity-2_3_cas><i class="far fa-question-circle activity-icon"></i>2.3. Квиз</div></a>
                <a href=../2_lekcija/2_4_cas.html><div class="studioLecture" id=activity-2_4_cas><i class="fas fa-file-alt activity-icon"></i>2.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-3_lekcija>
                    <h5>3. Робот Карел - задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../3_lekcija/3_1_cas.html><div class="studioLecture" id=activity-3_1_cas><i class="fas fa-file-alt activity-icon"></i>3.1. Робот Карел - задаци за вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-4_lekcija>
                    <h5>4. Корњача графика - линијски програми и петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../4_lekcija/4_1_cas.html><div class="studioLecture" id=activity-4_1_cas><i class="fas fa-file-alt activity-icon"></i>4.1. Наредбе корњаче</div></a>
                <a href=../4_lekcija/4_2_cas.html><div class="studioLecture" id=activity-4_2_cas><i class="fas fa-file-alt activity-icon"></i>4.2. Линијски програми</div></a>
                <a href=../4_lekcija/4_3_cas.html><div class="studioLecture" id=activity-4_3_cas><i class="fas fa-file-alt activity-icon"></i>4.3. Петље</div></a>
                <a href=../4_lekcija/4_4_cas.html><div class="studioLecture" id=activity-4_4_cas><i class="far fa-question-circle activity-icon"></i>4.4. Квиз</div></a>
                <a href=../4_lekcija/4_5_cas.html><div class="studioLecture" id=activity-4_5_cas><i class="fas fa-file-alt activity-icon"></i>4.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-5_lekcija>
                    <h5>5. Корњача графика - гранање, процедуре, торке/листе</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../5_lekcija/5_1_cas.html><div class="studioLecture" id=activity-5_1_cas><i class="fas fa-file-alt activity-icon"></i>5.1. Гранање</div></a>
                <a href=../5_lekcija/5_2_cas.html><div class="studioLecture" id=activity-5_2_cas><i class="fas fa-file-alt activity-icon"></i>5.2. Угнежђене петље</div></a>
                <a href=../5_lekcija/5_3_cas.html><div class="studioLecture" id=activity-5_3_cas><i class="fas fa-file-alt activity-icon"></i>5.3. Процедуре</div></a>
                <a href=../5_lekcija/5_4_cas.html><div class="studioLecture" id=activity-5_4_cas><i class="fas fa-file-alt activity-icon"></i>5.4. Торке/листе</div></a>
                <a href=../5_lekcija/5_5_cas.html><div class="studioLecture" id=activity-5_5_cas><i class="far fa-question-circle activity-icon"></i>5.5. Квиз</div></a>
                <a href=../5_lekcija/5_6_cas.html><div class="studioLecture" id=activity-5_6_cas><i class="fas fa-file-alt activity-icon"></i>5.6. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-6_lekcija>
                    <h5>6. Корњача графика - задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../6_lekcija/6_1_cas.html><div class="studioLecture" id=activity-6_1_cas><i class="fas fa-file-alt activity-icon"></i>6.1. Корњача графика - задаци вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-7_lekcija>
                    <h5>7. Израчунавања - изрази, оператори, променљиве</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../7_lekcija/7_1_cas.html><div class="studioLecture" id=activity-7_1_cas><i class="fas fa-file-alt activity-icon"></i>7.1. Основне аритметичке операције и примена</div></a>
                <a href=../7_lekcija/7_2_cas.html><div class="studioLecture" id=activity-7_2_cas><i class="fas fa-file-alt activity-icon"></i>7.2. Променљиве - имена међурезултата</div></a>
                <a href=../7_lekcija/7_3_cas.html><div class="studioLecture" id=activity-7_3_cas><i class="fas fa-file-alt activity-icon"></i>7.3. Тестирање и дебаговање програма</div></a>
                <a href=../7_lekcija/7_4_cas.html><div class="studioLecture" id=activity-7_4_cas><i class="far fa-question-circle activity-icon"></i>7.4. Квиз</div></a>
                <a href=../7_lekcija/7_5_cas.html><div class="studioLecture" id=activity-7_5_cas><i class="fas fa-file-alt activity-icon"></i>7.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-8_lekcija>
                    <h5>8. Израчунавања - цели и реални бројеви, дељење</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../8_lekcija/8_1_cas.html><div class="studioLecture" id=activity-8_1_cas><i class="fas fa-file-alt activity-icon"></i>8.1. Цели и реални бројеви</div></a>
                <a href=../8_lekcija/8_2_cas.html><div class="studioLecture" id=activity-8_2_cas><i class="fas fa-file-alt activity-icon"></i>8.2. Реално и целобројно дељење</div></a>
                <a href=../8_lekcija/8_3_cas.html><div class="studioLecture" id=activity-8_3_cas><i class="fas fa-file-alt activity-icon"></i>8.3. Позициони запис бројева</div></a>
                <a href=../8_lekcija/8_4_cas.html><div class="studioLecture" id=activity-8_4_cas><i class="far fa-question-circle activity-icon"></i>8.4. Квиз</div></a>
                <a href=../8_lekcija/8_5_cas.html><div class="studioLecture" id=activity-8_5_cas><i class="fas fa-file-alt activity-icon"></i>8.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-9_lekcija>
                    <h5>9. Израчунавања - коришћење и дефинисање функција</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../9_lekcija/9_1_cas.html><div class="studioLecture" id=activity-9_1_cas><i class="fas fa-file-alt activity-icon"></i>9.1. Коришћење функција</div></a>
                <a href=../9_lekcija/9_2_cas.html><div class="studioLecture" id=activity-9_2_cas><i class="fas fa-file-alt activity-icon"></i>9.2. Дефинисање функција</div></a>
                <a href=../9_lekcija/9_3_cas.html><div class="studioLecture" id=activity-9_3_cas><i class="far fa-question-circle activity-icon"></i>9.3. Квиз</div></a>
                <a href=../9_lekcija/9_4_cas.html><div class="studioLecture" id=activity-9_4_cas><i class="fas fa-file-alt activity-icon"></i>9.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-10_lekcija>
                    <h5>10. Израчунавања - задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../10_lekcija/10_1_cas.html><div class="studioLecture" id=activity-10_1_cas><i class="fas fa-file-alt activity-icon"></i>10.1. Израчунавања - задаци за вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-11_lekcija>
                    <h5>11. Гранање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../11_lekcija/11_1_cas.html><div class="studioLecture" id=activity-11_1_cas><i class="fas fa-file-alt activity-icon"></i>11.1. Гранање</div></a>
                <a href=../11_lekcija/11_2_cas.html><div class="studioLecture" id=activity-11_2_cas><i class="fas fa-file-alt activity-icon"></i>11.2. Комбинација више логичких услова</div></a>
                <a href=../11_lekcija/11_3_cas.html><div class="studioLecture" id=activity-11_3_cas><i class="fas fa-file-alt activity-icon"></i>11.3. Конструкција elif</div></a>
                <a href=../11_lekcija/11_4_cas.html><div class="studioLecture" id=activity-11_4_cas><i class="far fa-question-circle activity-icon"></i>11.4. Квиз</div></a>
                <a href=../11_lekcija/11_5_cas.html><div class="studioLecture" id=activity-11_5_cas><i class="fas fa-file-alt activity-icon"></i>11.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-12_lekcija>
                    <h5>12. Понављање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../12_lekcija/12_1_cas.html><div class="studioLecture" id=activity-12_1_cas><i class="fas fa-file-alt activity-icon"></i>12.1. Понављање</div></a>
                <a href=../12_lekcija/12_2_cas.html><div class="studioLecture" id=activity-12_2_cas><i class="fas fa-file-alt activity-icon"></i>12.2. Условна петља</div></a>
                <a href=../12_lekcija/12_3_cas.html><div class="studioLecture" id=activity-12_3_cas><i class="fas fa-file-alt activity-icon"></i>12.3. Угнежђене петље</div></a>
                <a href=../12_lekcija/12_4_cas.html><div class="studioLecture" id=activity-12_4_cas><i class="far fa-question-circle activity-icon"></i>12.4. Квиз</div></a>
                <a href=../12_lekcija/12_5_cas.html><div class="studioLecture" id=activity-12_5_cas><i class="fas fa-file-alt activity-icon"></i>12.5. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-13_lekcija>
                    <h5>13. Контрола тока - задаци за вежбање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../13_lekcija/13_1_cas.html><div class="studioLecture" id=activity-13_1_cas><i class="fas fa-file-alt activity-icon"></i>13.1. Контрола тока - задаци вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-14_lekcija>
                    <h5>14. Структуре података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../14_lekcija/14_1_cas.html><div class="studioLecture" id=activity-14_1_cas><i class="fas fa-file-alt activity-icon"></i>14.1. Торке, листе и ниске</div></a>
                <a href=../14_lekcija/14_2_cas.html><div class="studioLecture" id=activity-14_2_cas><i class="fas fa-file-alt activity-icon"></i>14.2. Скупови и речници</div></a>
                <a href=../14_lekcija/14_3_cas.html><div class="studioLecture" id=activity-14_3_cas><i class="far fa-question-circle activity-icon"></i>14.3. Квиз</div></a>
                <a href=../14_lekcija/14_4_cas.html><div class="studioLecture" id=activity-14_4_cas><i class="fas fa-file-alt activity-icon"></i>14.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-15_lekcija>
                    <h5>15. Структуре података и петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../15_lekcija/15_1_cas.html><div class="studioLecture" id=activity-15_1_cas><i class="fas fa-file-alt activity-icon"></i>15.1. Торке, листе, ниске и петље</div></a>
                <a href=../15_lekcija/15_2_cas.html><div class="studioLecture" id=activity-15_2_cas><i class="fas fa-file-alt activity-icon"></i>15.2. Речници и петље</div></a>
                <a href=../15_lekcija/15_3_cas.html><div class="studioLecture" id=activity-15_3_cas><i class="far fa-question-circle activity-icon"></i>15.3. Квиз</div></a>
                <a href=../15_lekcija/15_4_cas.html><div class="studioLecture" id=activity-15_4_cas><i class="fas fa-file-alt activity-icon"></i>15.4. Вежбање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-16_lekcija>
                    <h5>16. Откривање и исправљање грешака</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../16_lekcija/16_1_cas.html><div class="studioLecture" id=activity-16_1_cas><i class="fas fa-file-alt activity-icon"></i>16.1. Откривање и исправљање грешака и коришћење дебагера</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>11.5. Вежбање<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<div class="section" id="id2">
<h2>Индекс телесне масе<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Индекс телесне масе човека (енгл. body mass index, bmi) се дефинише
као количник његове масе у килограмима и квадрата његове висине у
метрима. Нормалним се сматра индекс телесне масе из (затвореног)
интервала од <span class="math notranslate nohighlight">\(18,5\frac{kg}{m^2}\)</span> до 25
<span class="math notranslate nohighlight">\(25\frac{kg}{m^2}\)</span>. Да ли човек који је висок 180 центиметара
и тежак 79 килограма нормалне дебљине?</p>

    </p></div></div>

<div data-childcomponent="telesna_masa" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=telesna_masa data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
masa = 79
visina = 180 / 100
bmi = masa / (visina * visina)
if 18.5 <= bmi and bmi <= 25:
    print("indeks telesne mase je u preporučenim granicama")
else:
    print("Masa treba da bude od", 18.5*visina*visina, "do", 25*visina*visina, "kg")
</textarea>
</div>
</div>
<div class="section" id="id3">
<h2>Кућни ред<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Кућни ред забрањује прављење буке пре 6 часова и између 13 и 17
часова, и након 22 часа. Напиши програм који радницима говори да ли
у неком датом тренутку могу да изводе бучније радове.</p>

    </p></div></div>
<p>Сматраћемо да су ови интервали полуотворени тј. да је дозвољено
правити буку у интервалу <span class="math notranslate nohighlight">\([6, 13)\)</span> и <span class="math notranslate nohighlight">\([17, 22)\)</span> и у
програму ћемо проверити да ли време припада неком од тих интервала.</p>

<div data-childcomponent="кућни_ред_1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=кућни_ред_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
sati = 15
minuta = 23
if (sati >= 6 and sati < 13) or (sati >= 17 and sati < 22):
    print("Možete da izvodite bučnije radove")
else:
    print("Ne bi trebalo da izvodite bučnije radove")
</textarea>
</div>
<p>или</p>

<div data-childcomponent="кућни_ред_2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=кућни_ред_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
if not (sati < 6 or (13 <= sati and sati < 17) or sati >= 22):
    print("Možete da izvodite bučnije radove")
else:
    print("Ne bi trebalo da izvodite bučnije radove")
</textarea>
</div>
</div>
<div class="section" id="id6">
<h2>Преступна година<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Година је преступна ако је дељива са 4 и није дељива са 100 или је
дељива са 400. Напиши функцију која проверава да ли је година
преступна.</p>

    </p></div></div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Овако уведен систем преступних година уводи тачно 97 преступних
година на сваких 400 година. Тиме је време једне године тачно (97 ×
366 + 303 × 365) / 400 = 365,2425 дана, што је прилично добра
процена вредности 365,242374, коју астрономи обично узимају за
трајање године.</p>

    </p></div></div>
<p>Подсетимо се, проверу да ли је број дељив другим бројем можеш извршити
тако што израчунаш остатак при дељењу (коришћењем оператора <code class="docutils literal notranslate"><span class="pre">%</span></code>) и
провериш да ли је он једнак нули. На пример, проверу да ли је година
дељива са 4 можеш извршити изразом <code class="docutils literal notranslate"><span class="pre">godina</span> <span class="pre">%</span> <span class="pre">4</span> <span class="pre">==</span> <span class="pre">0</span></code>.</p>

<div data-childcomponent="преступна_година" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=преступна_година data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
def prestupna(godina):
    return True

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(prestupna(2000), True, "2000 јесте преступна")
       self.assertEqual(prestupna(1900), False, "1900 није преступна")
       self.assertEqual(prestupna(2017), False, "2017 није преступна")
       self.assertEqual(prestupna(2020), True, "2020 јесте преступна")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id8">
<h2>Предност у превозу<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Предност у превозу имају труднице, деца млађа од 7 година и старији
од 65 година. Допуни програм који испитује да ли особа има предност
(параметар <code class="docutils literal notranslate"><span class="pre">godine</span></code> садржи број година, а параметар <code class="docutils literal notranslate"><span class="pre">trudnica</span></code>
има вредност <code class="docutils literal notranslate"><span class="pre">true</span></code> ако је особа трудница, односно <code class="docutils literal notranslate"><span class="pre">false</span></code> у
супротном).</p>

    </p></div></div>

<div data-childcomponent="предност_у_превозу" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=предност_у_превозу data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
def prednost(godine, trudnica)
    return True

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(prednost(34, True), True, "Трудница од 34 године има предност")
       self.assertEqual(prednost(6, False), True, "Дете од 6 година има предност")
       self.assertEqual(prednost(8, False), False, "Дете од 8 година нема предност")
       self.assertEqual(prednost(42, False), False, "Мушкарац од 42 године нема предност")
       self.assertEqual(prednost(67, False), True, "Пензионерка од 67 година има предност")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id10">
<h2>Број потеза краља<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Позиције на шаховској табли се обележавају, на пример, са <code class="docutils literal notranslate"><span class="pre">a3</span></code>,
<code class="docutils literal notranslate"><span class="pre">b5</span></code>, <code class="docutils literal notranslate"><span class="pre">h1</span></code> и слично. Прво се наводи словна ознака колоне (од
<code class="docutils literal notranslate"><span class="pre">a</span></code> до <code class="docutils literal notranslate"><span class="pre">h</span></code>), а затим бројевна ознака врсте (од 1 до 8). Ако је
дата позиција краља на шаховској табли на којој нема других фигура
осим тог краља, напиши програм који одређује број потеза које тај
краљ може да направи (краљ се у шаху може померити на било које
њему суседно поље).</p>

    </p></div></div>
<p>Постоје три различите врсте поља. Поља у угловима табле (поља <code class="docutils literal notranslate"><span class="pre">a1</span></code>,
<code class="docutils literal notranslate"><span class="pre">a8</span></code>, <code class="docutils literal notranslate"><span class="pre">h1</span></code> и <code class="docutils literal notranslate"><span class="pre">h8</span></code>) су таква да краљ може да се помери на само
три околне позиције. Поља која су на ивицама (у колонама <code class="docutils literal notranslate"><span class="pre">a</span></code> или
<code class="docutils literal notranslate"><span class="pre">h</span></code> тј. у врстама <code class="docutils literal notranslate"><span class="pre">1</span></code> или <code class="docutils literal notranslate"><span class="pre">8</span></code>), али нису у угловима су таква да
краљ може да се помери на пет околних позиција. Са осталих поља краљ
може да се помери на осам суседних позиција.</p>

<div data-childcomponent="краљ_број_потеза" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=краљ_број_потеза data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
kralj = input("Unesi poziciju kralja:")
kolona = kralj[0]
vrsta  = kralj[1]
kolona_na_ivici = kolona == "a" or kolona == "h"
vrsta_na_ivici  = vrsta == "1" or vrsta == "8"
if True:  # ispravi ovaj red
    broj_poteza = 3
elif True:  # ispravi ovaj red
    broj_poteza = 5
else:
    broj_poteza = 0   # ispravi ovaj red
print(broj_poteza)
</textarea>
</div>
</div>
<div class="section" id="id12">
<h2>Килобајт<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Напиши програм који проверава да ли је корисник научио колико
бајтова постоји у једном килобајту.</p>

    </p></div></div>

<div data-childcomponent="бајтова_у_килобајту" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=бајтова_у_килобајту data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
odgovor = int(input("Koliko bajtova ima u jednom kilobajtu"))
if True:  # popravi ovaj uslov
    print("Bravo")
else:
    print("Odgovor nije tačan")
</textarea>
</div>

            <div class="course-box course-box-question course-content">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="бајтова_у_килобајту_2">
            <p>Q-8: Који од наредних услова треба употребити у претходном програму?</p>

            <li data-component="answer"  id="бајтова_у_килобајту_2_opt_a">odgovor == 1000</li><li data-component="feedback" id="бајтова_у_килобајту_2_opt_a">Килобајт садржи 1024 бајта.</li>
            
            <li data-component="answer"  id="бајтова_у_килобајту_2_opt_b">odgovor = 1000</li><li data-component="feedback" id="бајтова_у_килобајту_2_opt_b">Килобајт садржи 1024 бајта, а поређење једнакости се записује
коришћењем ==.</li>
            
            <li data-component="answer" data-correct id="бајтова_у_килобајту_2_opt_c">odgovor == 1024</li><li data-component="feedback" id="бајтова_у_килобајту_2_opt_c">Браво!</li>
            
            <li data-component="answer"  id="бајтова_у_килобајту_2_opt_d">odgovor = 1024</li><li data-component="feedback" id="бајтова_у_килобајту_2_opt_d">Поређење једнакости се записује коришћењем ==.</li>
            

            </ul>
            </div>
            </div>
<div class="section" id="id14">
<h2>Жупан или краљ<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Стефан Немањић је био Велики жупан од 1196 до 1217, а краљ од 1217
до 1228. Да ли је дуже владао као краљ или Велики жупан?</p>

    </p></div></div>

<div data-childcomponent="жупан_краљ" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=жупан_краљ data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
zupan = 0   # ispravi ovaj red
kralj = 0   # ispravi ovaj red
if zupan < kralj:
print("Дуже је био краљ.")     # ispravi ovaj red
else:
print("Дуже је био жупан.")    # ispravi ovaj red
</textarea>
</div>
</div>
<div class="section" id="id16">
<h2>Оцена на факултету<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>На факултету се оцена одређује на основу броја поена на следећи
начин. За 50 поена и мање добија се оцена 5, за поене од 51 до 60
добија се оцена 6, од 61 до 70 оцена 7, од 71 до 80 оцена 8, од 81
до 90 оцена 9 и за поене од 91 до 100 добија се оцена 10. Напиши
програм који за дати број поена одређује оцену.</p>

    </p></div></div>

<div data-childcomponent="оцена_на_основу_поена" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=оцена_на_основу_поена data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
poeni = int(input("Unesi broj poena:"))

if poeni < 51:
    ocena = 5
elif True:      # ispravi ovaj red
    ocena = 6
elif True:      # ispravi ovaj red
    ocena = 7
elif poeni < 81:
    ocena = 0   # ispravi ovaj red
elif True:      # ispravi ovaj red
    ocena = 9
else:
    ocena = 0    # ispravi ovaj red

print(ocena)
</textarea>
</div>
<p>Рецимо и да смо овај задатак могли решити уз помоћ мало математике,
без гранања.</p>

<div data-childcomponent="оцена_на_основу_поена_без_гранања" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=оцена_на_основу_поена_без_гранања data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
poeni = int(input("Unesi broj poena:"))
ocena = (poeni - 1) // 10 + 1
print(ocena)
</textarea>
</div>
</div>
<div class="section" id="id19">
<h2>Потези фигура у шаху<a class="headerlink" href="#id19" title="Permalink to this headline">¶</a></h2>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Краљ у шаху може да се помера само једно поље (на било које од
могућих 8 суседних поља). Топ у шаху може да се помера вертикално
или хоризонтално, било који број поља. Ловац у шаху може да се
помера дијагонално, било који број поља. Краљица може да се помера
хоризонтално, вертикално или дијагонално, било који број поља. Коњ
се помера тако што иде два поља вертикално и једно поље
хоризонтално или два поља хоризонтално и једно поље вертикално. Са
сваку од описаних шаховских фигура дефиниши функцију која за дата
два поља на шаховској табли (одређена својим координатама) одређује
да ли фигура на празној табли може да стигне са првог на друго
поље.</p>

    </p></div></div>
<p>Први услов је да полазно поље мора бити различито од долазног. Ако су
дата поља са координатама <code class="docutils literal notranslate"><span class="pre">(x1,</span> <span class="pre">y1)</span></code> и <code class="docutils literal notranslate"><span class="pre">(x2,</span> <span class="pre">y2)</span></code>, овај се услов
просто може проверити помоћу <code class="docutils literal notranslate"><span class="pre">(x1,</span> <span class="pre">y1)</span> <span class="pre">!=</span> <span class="pre">(x2,</span> <span class="pre">y2)</span></code>. Други начин не
користи парове и захтева да је бар једна од две координате различита
тј. да важи <code class="docutils literal notranslate"><span class="pre">x1</span> <span class="pre">!=</span> <span class="pre">x2</span> <span class="pre">or</span> <span class="pre">y1</span> <span class="pre">!=</span> <span class="pre">y2</span></code>.</p>
<p>Већина провера се може засновати на анализи хоризонталног и
вертикалног растојања између два поља. Када смо разматрали апсолутну
вредност рекли смо да се растојање између два броја може једноставно
израчунати као апсолутна вредност њихове разлике. Тако, , хоризонтално
растојање можемо одредити помоћу <code class="docutils literal notranslate"><span class="pre">abs(x1</span> <span class="pre">-</span> <span class="pre">x2)</span></code>, док вертикално
растојање можемо одредити помоћу <code class="docutils literal notranslate"><span class="pre">abs(y1</span> <span class="pre">-</span> <span class="pre">y2)</span></code>.</p>
<ul class="simple">
<li><p>Краљ се може померити ако је веће од ова два растојања једнако тачно
1 (тада је растојање по једној координати 1, а по другој 0 или 1,
што је тачно услов померања краља).</p></li>
<li><p>Топ се може померити ако полазно и долазно поље имају исту
координату <code class="docutils literal notranslate"><span class="pre">x</span></code> или исту координату <code class="docutils literal notranslate"><span class="pre">y</span></code> (обе координате не смеју
бити једнаке јер би тада полазно и долазно поље било исто).</p></li>
<li><p>Ловац се може померити ако два поља леже на истој дијагонали. Овај
се критеријум може проверити тако што се утврди да је хоризонтално
растојање између два поља једнако њиховом вертикалном растојању (у
сваком дијагоналном кораку се та растојања умањују за 1, све док се
не стигне до долазног поља када оба та растојања истовремено постају
нула, што значи да су у почетку морала бити једнака).</p></li>
<li><p>Проверу за краљицу можеш извршити тако што ћеш искомбиновати две већ
направљене провере (ону за топа и ловца).</p></li>
<li><p>На крају, проверу за коња можеш извршити тако што ћеш проверити да
ли је хоризонтално растојање једнако 2, а вертикално 1 или је
хоризонтално растојање једнако 1, а вертикално 2.</p></li>
</ul>
<p>Имајући претходну анализу у виду, покушај да дефинишеш функције које
проверавају доступност поља. Када покренеш програм видећеш све потезе
које свака од фигура може да направи са датог поља, на основу чега ћеш
видети да ли је твоје решење исправно.</p>

<div data-childcomponent="шаховске_фигуре" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=шаховске_фигуре data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
def kralj(x1, y1, x2, y2):
    return False

def top(x1, y1, x2, y2):
    return False

def lovac(x1, y1, x2, y2):
    return False

def kraljica(x1, y1, x2, y2):
    return False

def konj(x1, y1, x2, y2):
    return False

polje = (3, 5)

====
def ispisi_dostupna_polja(polje, naziv_figure, figura_provera):
    for x in range(1, 8+1):
       for y in range(1, 8+1):
          if (figura_provera(polje[0], polje[1], x, y)):
             print(naziv_figure, x, y)

ispisi_dostupna_polja(polje, 'kralj:', kralj)
ispisi_dostupna_polja(polje, 'top:', top)
ispisi_dostupna_polja(polje, 'lovac:', lovac)
ispisi_dostupna_polja(polje, 'kraljica:', kraljica)
ispisi_dostupna_polja(polje, 'konj:', konj)
</textarea>
</div>

    <div data-component="reveal" id="терен_тест_решење11"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="шаховске_фигуре_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=шаховске_фигуре_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
def kralj(x1, y1, x2, y2):
    return max(abs(x1-x2), abs(y1-y2)) == 1

def top(x1, y1, x2, y2):
    return (x1 == x2 or y1 == y2) and (x1, y1) != (x2, y2)

def lovac(x1, y1, x2, y2):
    return abs(x1 - x2) == abs(y1 - y2) and (x1, y1) != (x2, y2)

def kraljica(x1, y1, x2, y2):
    return top(x1, y1, x2, y2) or lovac(x1, y1, x2, y2)

def konj(x1, y1, x2, y2):
    return ((abs(x1 - x2) == 2 and abs(y1 - y2) == 1) or \
            (abs(x1 - x2) == 1 and abs(y1 - y2) == 2)) and \
           (x1, y1) != (x2, y2)

polje = (3, 5)

====
def ispisi_dostupna_polja(polje, naziv_figure, figura_provera):
  for x in range(1, 8+1):
    for y in range(1, 8+1):
       if (figura_provera(polje[0], polje[1], x, y)):
          print(naziv_figure, x, y)

ispisi_dostupna_polja(polje, 'kralj:', kralj)
ispisi_dostupna_polja(polje, 'top:', top)
ispisi_dostupna_polja(polje, 'lovac:', lovac)
ispisi_dostupna_polja(polje, 'kraljica:', kraljica)
ispisi_dostupna_polja(polje, 'konj:', konj)
</textarea>
</div>

    </div>
    </div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="11_4_cas.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходна лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>

<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>


<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>