<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>Речници &#8212; Petlja - Информатика и рачунарство за шести разред</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="Петнаеста лекција" href="../15_lekcija/index.html" />
    <link rel="prev" title="Скупови" href="14_3_cas.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Информатика и рачунарство за шести разред';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Информатика и рачунарство за шести разред</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">Прва лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_cas.html">Робот Карел и линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_cas.html">Понављање - бројачка петља <cite>for</cite></a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_3_cas.html">Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_test.html">Тест - линијска структура</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_test.html">Тест - понављање, бројачка петља for</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">Друга лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_cas.html">Робот Карел - условне петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_cas.html">Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_3_cas.html">Вежбање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_test.html">Тест - условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_test.html">Tест - гранање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">Трећа лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3_1_cas.html">Робот Карел - вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">Четврта лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_1_cas.html">Корњача графика - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_2_cas.html">Корњача графика - понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">Пета лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_1_cas.html">Корњача графика - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_2_cas.html">Корњача графика - угнежђене петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_3_cas.html">Корњача графика - процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_4_cas.html">Корњача графика - торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_5_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">Седма лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_1_cas.html">Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_2_cas.html">Променљиве - имена међурезултата</a></li>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_3_cas.html">Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_4_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../8_lekcija/index.html">Осма лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_1_cas.html">Цели и реални бројеви</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_2_cas.html">Реално и целобројно дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../9_lekcija/index.html">Девета лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_1_cas.html">Коришћење и дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_2_cas.html">Дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../9_lekcija/9_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../11_lekcija/index.html">Једанаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_1_cas.html">Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_2_cas.html">Конструкција elif</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_3_cas.html">Вежбање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">Дванаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_1_cas.html">Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_2_cas.html">Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_3_cas.html">Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">Четрнаеста лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="14_1_cas.html">Структуре података: Торке, листе, ниске, скупови, речници</a></li>
<li class="toctree-l2"><a class="reference internal" href="14_2_cas.html">Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="14_2_cas.html#id11">Врсте реченица</a></li>
<li class="toctree-l2"><a class="reference internal" href="14_3_cas.html">Скупови</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">Речници</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">Петнаеста лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_1_cas.html">Обрада структура података коришћењем петљи</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>Речници<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Листе нам дају могућност да организујемо податке за које је познат
неки редослед и у којима сваки податак има свој редни број (на пример,
листу смо употребили да организујемо имена путника у авиону или
податке о месецима у години) и приступ елементима листе врши се на
основу њиховог индекса (редног броја, тј. позиције). Међутим, често
подаци представљају неко пресликавање у којем на основу датог <em>кључа</em>
желимо да приступимо <em>вредности</em> која је придружена том кључу. На
пример, на основу шифре производа приступамо његовој цени, на основу
ЈМБГ грађана приступамо његовом имену, на основу броја банковног
рачуна приступамо салду (износу новца) на том рачуну и слично. У тим
ситуацијама уместо листе много је боље користити <strong>речнике</strong> (каже се
и мапе или асоцијативне низове).</p>
<div class="section" id="id2">
<h2>Дефинисање и коришћење речника<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id3">
<h3>Цене аутомобила<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Аутомобили у каталогу имају придружене цене и ми желимо да у нашем
програму можемо да одредимо цену аутомобила на основу његовог
модела. Напиши програм који на основу унетог модела аутомобила
(ниска) одређује његову цену (цео број).</p>

    </p></div></div>
<p>Кључна одлука за ефикасно решење овог задатка је да каталог цена
аутомобила представимо помоћу речника.</p>

<div data-childcomponent="цене_аутомобила" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=цене_аутомобила data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
cene_automobila = {"fiat 500l": 11990,
                   "renault clio": 9650,
                   "toyota corolla": 13990}
automobil = input("Unesi model automobila:")
print(cene_automobila[automobil])
</textarea>
</div>
<p>У претходном примеру модели аутомобила су кључеви који се пресликавају
у цене аутомобила које су вредности. Кључеви су биле ниске, а
вредности цели бројеви, што је чест случај.</p>
</div>
<div class="section" id="id5">
<h3>Географске координате градова<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>

    <div data-level="2" style="display:none">

    </div>
<p>У наредном примеру вредности у речнику су уређени парови.</p>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Познате су географске координате неколико главних европских
градова. За дато име града одреди њене географске
координате. Одреди посебно географску дужину и посебно географску
ширину.</p>

    </p></div></div>

<div data-childcomponent="географске_координате" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=географске_координате data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
gradovi = {"Beograd": (44.7866, 20.4489),
           "Budimpešta": (47.4979, 19.0402),
           "Beč": (48.2082, 16.3738),
           "Bratislava": (48.1486, 17.1077)}
grad = input("Unesi ime grada: ")
# ispravi naredni red tako da se iz rečnika pročitaju koordinate grada
koordinate = ???
print(koordinate)
# dopuni naredni kod tako da se ispišu geografska širina i dužina
print("Geografska širina: ")
print("Geografska dužina: ")
</textarea>
</div>
</div>
<div class="section" id="id7">
<h3>Цена производа<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h3>
<p>Слично као и код осталих структура података и провера да ли је неком
кључу у речнику придружена вредност се може проверити оператором
<code class="docutils literal notranslate"><span class="pre">in</span></code>.</p>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У речнику се чувају цене различитих производа у продавници. Пронађи
и испиши цену датог производа или пријави да се тај производ не
продаје.</p>

    </p></div></div>

<div data-childcomponent="припада_речнику" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=припада_речнику data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
cene = {"hleb": 39.90, "mleko": 89, "jaje": 14.90}
proizvod = "hleb"
if proizvod in cene:
    print(proizvod, "košta", cene[proizvod], "dinara")
else:
    print(proizvod, "se ne prodaje u ovoj prodavnici")
</textarea>
</div>
<div class="section" id="id9">
<h4>Домаћи задатак<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h4>
<p>Ова лекција је била прилично обимна, па је могуће да нисте током часа
успели да је целу пређете. За домаћи пређи све задатке које су на часу
остали неурађени, а затим уради и наредних неколико задатака.</p>
</div>
</div>
</div>
<div class="section" id="id10">
<h2>Задаци<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id11">
<h3>Да ли си награђен?<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Позната је листа награђених ученика. Напиши програм који проверава
да ли је корисник чије се име учитава на почетку програма освојио
награду.</p>

    </p></div></div>

<div data-childcomponent="члан_листе_награђених" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=члан_листе_награђених data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
nagradjeni = ["Lenka", "Sofija", "Dimitrije"]
ime = ""    # ispravi ovaj red tako da se unosi ime ucenika
if ???:     # dodaj uslov koji proverava da li se uneto ime nalazi u listi
    print("Bravo, nagrada je tvoja!")
else:
    print("Nažalost, nisi među nagrađenima")
</textarea>
</div>
</div>
<div class="section" id="id13">
<h3>Путници на датим седиштима у авиону<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Познат је списак имена путника у авиону. Седишта су нумерисана од 1
па надаље. Ако стјуардеса унесе број седишта, напиши програм који
одређује име путника на том седишту. Након тога испиши имена путника
која седе на прва четири, као и на последња два седишта, као и број
седишта на којем седи Мика Микић.</p>

    </p></div></div>
<p>Да бисмо олакшали бројање, на првом месту у листи је постављен
вештачки елемент (празна ниска).</p>

<div data-childcomponent="путници_у_авиону" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=путници_у_авиону data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
putnici = ["", "Петар Петровић", "Зорана Зорановић", "Мика Микић", \
           "Јелена Јеленковић", "Ана Анић", "Милан Милановић"]
# ispravi naredne redove tako da se ispise ime putnika na unetom broju sedista
sediste = int(input("Unesi broj sedista:"))
print("Na sedistu", sediste, "sedi", putnici[0])
# ispravi naredni red tako da se ispisu imena prva 4 putnika
print(putnici[0])
# ispravi naredni red tako da se ispisu imena poslednja 2 putnika
print(putnici[0])
# ispravi naredni red tako da se ispise broj sedista na kojem sedi Мика Микић
print(putnici[0])
</textarea>
</div>
</div>
<div class="section" id="id15">
<h3>Шифра<a class="headerlink" href="#id15" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Лидијина шифра за gmail се састоји од 6 слова. Одлучила је да је
закомпликује, али на правилан начин, тако да касније може да је се
сети. Шифру ће закомпликовати тако што ће заменити редослед прва
два и последња два слова, а средња два слова поновити два пута. На
пример, ако јој је шифра била lidija закомпликована шифра јој је
jadidili. Напиши програм који исписује закомпликовану шифру.</p>

    </p></div></div>

<div data-childcomponent="закомпликована_шифра" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=закомпликована_шифра data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
sifra = "lidija"
zakomplikovana_sifra = sifra[4:6] + ""  # popravi ovaj red
print(zakomplikovana_sifra)
</textarea>
</div>
</div>
<div class="section" id="id17">
<h3>Учитавање разломка<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h3>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Корисник уноси текст који описује разломак (на пример,
<cite>3/4</cite>). Израчунај децималну вредност тог разломка.</p>

    </p></div></div>

<div data-childcomponent="разломак" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=разломак data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
razlomak = input("Unesi razlomak: ")
pozicija_crte = razlomak.find("/")
brojilac = int(0)    # popravi ovaj red
imenilac = int(razlomak[pozicija_crte+1:])
print(0)   # popravi ovaj red
</textarea>
</div>
</div>
<div class="section" id="id19">
<h3>Сугласници<a class="headerlink" href="#id19" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Сугласници се по месту изговора могу бити предњонепчани (Ј, Љ, Њ, Ћ,
Ђ, Ш, Ж, Ч, Џ) и задњонепчани К, Г, Х. По звучности, сугласници се
деле на звучне (Б, Г, Д, Ђ, Ж, З, Џ), безвучне (П, К, Т, Ћ, Ш, С, Ч,
Ф, Х, Ц), и сонанте(М, В, Р, Л, Н, Љ, Њ, Ј).</p>
<ul class="simple">
<li><p>Који сугласници су истовремено и звучни и предњонепчани?</p></li>
<li><p>Који безвучни сугласници нису задњонепчани?</p></li>
<li><p>Који сугласници су сонанти или звучни?</p></li>
</ul>

    </p></div></div>
<p>Поправи наредни код тако да исправно решава овај задатак. Програм
након покретања треба да испише наредне резултате.</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="p">{</span><span class="s1">&#39;Ђ&#39;</span><span class="p">,</span> <span class="s1">&#39;Џ&#39;</span><span class="p">,</span> <span class="s1">&#39;Ж&#39;</span><span class="p">}</span>
<span class="p">{</span><span class="s1">&#39;Ч&#39;</span><span class="p">,</span> <span class="s1">&#39;Ћ&#39;</span><span class="p">,</span> <span class="s1">&#39;Ш&#39;</span><span class="p">,</span> <span class="s1">&#39;Т&#39;</span><span class="p">,</span> <span class="s1">&#39;П&#39;</span><span class="p">,</span> <span class="s1">&#39;Ф&#39;</span><span class="p">,</span> <span class="s1">&#39;С&#39;</span><span class="p">,</span> <span class="s1">&#39;Ц&#39;</span><span class="p">}</span>
<span class="p">{</span><span class="s1">&#39;Џ&#39;</span><span class="p">,</span> <span class="s1">&#39;З&#39;</span><span class="p">,</span> <span class="s1">&#39;Д&#39;</span><span class="p">,</span> <span class="s1">&#39;Ђ&#39;</span><span class="p">,</span> <span class="s1">&#39;Ј&#39;</span><span class="p">,</span> <span class="s1">&#39;Л&#39;</span><span class="p">,</span> <span class="s1">&#39;Н&#39;</span><span class="p">,</span> <span class="s1">&#39;М&#39;</span><span class="p">,</span> <span class="s1">&#39;Г&#39;</span><span class="p">,</span> <span class="s1">&#39;Р&#39;</span><span class="p">,</span> <span class="s1">&#39;Б&#39;</span><span class="p">,</span> <span class="s1">&#39;Љ&#39;</span><span class="p">,</span> <span class="s1">&#39;Њ&#39;</span><span class="p">,</span> <span class="s1">&#39;Ж&#39;</span><span class="p">,</span> <span class="s1">&#39;В&#39;</span><span class="p">}</span>
</pre></div>
</div>

<div data-childcomponent="скупови_сугласника" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=скупови_сугласника data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
zadnjonepcani  = {"К", "Г", "Х"}
prednjonepcani = {"Ј", "Љ", "Њ", "Ћ", "Ђ", "Ш", "Ж", "Ч", "Џ"}
zvucni   = {"Б", "Г", "Д", "Ђ", "Ж", "З", "Џ"}
bezvucni = {"П", "К", "Т", "Ћ", "Ш", "С", "Ч", "Ф", "Х", "Ц"}
sonanti  = {"М", "В", "Р", "Л", "Н", "Љ", "Њ", "Ј"}

zvucni_prednjonepcani = {}           # popravi ovaj red
bezvucni_ne_zadnjonepcani = {}       # popravi ovaj red
zvucni_ili_sonanti = {}              # popravi ovaj red
print(zvucni_prednjonepcani)
print(bezvucni_ne_zadnjonepcani)
print(zvucni_ili_sonanti)
</textarea>
</div>
</div>
<div class="section" id="id21">
<h3>Звучни и безвучни сугласници<a class="headerlink" href="#id21" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Звучни и безвучни сугласници јављају се у паровима.  Звучни
сугласници су БГДЂЖЗЏ (можеш их лакше запамтити помоћу првих слова
реченице Баба грди деду: „Ђаволе живи зашто џандрљаш”?) Њихови
безвучни парови су редом ПКТЋШСЧ. Напиши програм који за дати
звучни сугласник одређује његов безвучни пар.</p>

    </p></div></div>
<p>Поправи наредни програм тако што ћеш попунити речник одговарајућим
паровима сугласника. Ако све урадиш како треба, требало би да добијеш
7 зелених поља (сваки сугласник мора да прође тест). Користи велика
слова ћирилице.</p>

<div data-childcomponent="безвучни_звучни_парови" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=безвучни_звучни_парови data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
def bezvucni(zvucni):
    bezvucni_par = {}    # поправи овај ред
    return bezvucni_par[zvucni]

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
       self.assertEqual(bezvucni("Б"), "П", "Безвучни пар гласа Б је глас П.")
       self.assertEqual(bezvucni("Г"), "К", "Безвучни пар гласа Г је глас К.")
       self.assertEqual(bezvucni("Д"), "Т", "Безвучни пар гласа Д је глас Т.")
       self.assertEqual(bezvucni("Ђ"), "Ћ", "Безвучни пар гласа Ђ је глас Ћ.")
       self.assertEqual(bezvucni("Ж"), "Ш", "Безвучни пар гласа Ж је глас Ш.")
       self.assertEqual(bezvucni("З"), "С", "Безвучни пар гласа З је глас С.")
       self.assertEqual(bezvucni("Џ"), "Ч", "Безвучни пар гласа Џ је глас Ч.")

myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id23">
<h3>Главни градови<a class="headerlink" href="#id23" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У речнику се чувају главни градови неколико држава. Напиши програм
који учитава име државе и исписује главни град или каже да одговор
на то питање не зна.</p>

    </p></div></div>

<div data-childcomponent="главни_град" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=главни_град data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
glavni_grad = {"Srbija": "Beograd",
               "Mađarska": "Budimpešta",
               "Rumunija": "Bukurešt",
               "Bugarska": "Sofija",
               "Makedonija": "Skoplje",
               "Albanija": "Tirana",
               "Crna Gora": "Podgorica",
               "Bosna i Hercegovina": "Sarajevo",
               "Hrvatska": "Zagreb"}
# završi ovaj program
</textarea>
</div>
</div>
</div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - Скупови' data-toggle="tooltip"><a href="14_3_cas.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - Петнаеста лекција' data-toggle="tooltip"><a href="../15_lekcija/index.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - Скупови' data-toggle="tooltip"><a href="14_3_cas.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - Петнаеста лекција' data-toggle="tooltip"><a href="../15_lekcija/index.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>