<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <title>9.2. Реално и целобројно дељење &#8212; Petlja - Информатика и рачунарство за шести разред</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script src="../_static/runestonebase.js"></script>
    <script src="../_static/jquery.js"></script>
    <script src="../_static/underscore.js"></script>
    <script src="../_static/doctools.js"></script>
    <script src="../_static/language_data.js"></script>
    <script src="../_static/matrixeq.js"></script>
    <script src="../_static/mchoice.js"></script>
    <script src="../_static/timedmc.js"></script>
    <script src="../_static/timed.js"></script>
    <script src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script src="../_static/mchoice-i18n.en.js"></script>
    <script src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script src="../_static/jquery.highlight.js"></script>
    <script src="../_static/bookfuncs.js"></script>
    <script src="../_static/codemirror.js"></script>
    <script src="../_static/xml.js"></script>
    <script src="../_static/css.js"></script>
    <script src="../_static/python.js"></script>
    <script src="../_static/htmlmixed.js"></script>
    <script src="../_static/javascript.js"></script>
    <script src="../_static/activecode-i18n.en.js"></script>
    <script src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script src="../_static/skulpt.min.js"></script>
    <script src="../_static/skulpt-stdlib.js"></script>
    <script src="../_static/activecode.js"></script>
    <script src="../_static/clike.js"></script>
    <script src="../_static/timed_activecode.js"></script>
    <script src="../_static/tabbedstuff.js"></script>
    <script src="../_static/runestonevideo.js"></script>
    <script src="../_static/shortanswer.js"></script>
    <script src="../_static/timed_shortanswer.js"></script>
    <script src="../_static/poll.js"></script>
    <script src="../_static/showEval.js"></script>
    <script src="../_static/d3.v2.min.js"></script>
    <script src="../_static/jquery.ba-bbq.min.js"></script>
    <script src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script src="../_static/pytutor.js"></script>
    <script src="../_static/codelens.js"></script>
    <script src="../_static/clickable.js"></script>
    <script src="../_static/timedclickable.js"></script>
    <script src="../_static/dragndrop.js"></script>
    <script src="../_static/timeddnd.js"></script>
    <script src="../_static/dragndrop-i18n.en.js"></script>
    <script src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script src="../_static/fitb.js"></script>
    <script src="../_static/timedfitb.js"></script>
    <script src="../_static/fitb-i18n.en.js"></script>
    <script src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script src="../_static/webglinteractive.js"></script>
    <script src="../_static/FileSaver.min.js"></script>
    <script src="../_static/Blob.js"></script>
    <script src="../_static/reveal.js"></script>
    <script src="../_static/animationbase.js"></script>
    <script src="../_static/datafile.js"></script>
    <script src="../_static/lib/prettify.js"></script>
    <script src="../_static/lib/hammer.min.js"></script>
    <script src="../_static/parsons.js"></script>
    <script src="../_static/parsons-i18n.en.js"></script>
    <script src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script src="../_static/timedparsons.js"></script>
    <script src="../_static/karelCorner.js"></script>
    <script src="../_static/karelRobot.js"></script>
    <script src="../_static/karelWorld.js"></script>
    <script src="../_static/karelRobotDrawer.js"></script>
    <script src="../_static/karelUI.js"></script>
    <script src="../_static/karel.js"></script>
    <script src="../_static/karel-i18n.en.js"></script>
    <script src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script src="../_static/notes.js"></script>
    <script src="../_static/pygamelib-init.js"></script>
    <script src="../_static/blockly/blockly_compressed.js"></script>
    <script src="../_static/blockly/blocks_compressed.js"></script>
    <script src="../_static/blockly/python_compressed.js"></script>
    <script src="../_static/blockly/msg-sr.js"></script>
    <script src="../_static/blockpy/utilities.js"></script>
    <script src="../_static/blockpy/python_errors.js"></script>
    <script src="../_static/blockpy/ast_node_visitor.js"></script>
    <script src="../_static/blockpy/abstract_interpreter.js"></script>
    <script src="../_static/blockpy/pytifa.js"></script>
    <script src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script src="../_static/blockpy/python_to_blockly.js"></script>
    <script src="../_static/blockpy/imported.js"></script>
    <script src="../_static/blockpy/blocks/class.js"></script>
    <script src="../_static/blockpy/blocks/comment.js"></script>
    <script src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script src="../_static/blockpy/blocks/dict.js"></script>
    <script src="../_static/blockpy/blocks/if.js"></script>
    <script src="../_static/blockpy/blocks/io.js"></script>
    <script src="../_static/blockpy/blocks/lists.js"></script>
    <script src="../_static/blockpy/blocks/sets.js"></script>
    <script src="../_static/blockpy/blocks/loops.js"></script>
    <script src="../_static/blockpy/blocks/parking.js"></script>
    <script src="../_static/blockpy/blocks/tuple.js"></script>
    <script src="../_static/blockpy/blocks/turtles.js"></script>
    <script src="../_static/blockpy/blocks/text.js"></script>
    <script src="../_static/blockpy-modal.js"></script>
    <script async="async" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script src="../_static/jquery-fix.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script src="../_static/bootstrap-sphinx.js"></script>
    <script src="../_static/waypoints.min.js"></script>
    <script src="../_static/rangy-core.js"></script>
    <script src="../_static/rangy-textrange.js"></script>
    <script src="../_static/rangy-cssclassapplier.js"></script>
    <script src="../_static/user-highlights.js"></script>
    <script src="../_static/jquery.idle-timer.js"></script>
    <script src="../_static/processing-1.4.1.min.js"></script>
    <script src="../_static/jquery.hotkey.js"></script>
    <script src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="9.3. Вежбање - цели и реални бројеви" href="9_3_cas.html" />
    <link rel="prev" title="9.1. Цели и реални бројеви" href="9_1_cas.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'Информатика и рачунарство за шести разред';
  eBookConfig.logLevel = 0;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = false;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'Информатика и рачунарство за шести разред';
  eBookConfig.runestone_version = '3.1.1';
  eBookConfig.imagesDir = '../_images/';
  eBookConfig.staticDir = '../_static/';
  if(typeof(Sk) != "undefined")
      Sk.imgPath = eBookConfig.imagesDir;
</script>

<div id="fb-root"></div>



  </head><body>


<!-- Begin navbar -->

<nav id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <button type="button" class="navbar-toggle collapsed" data-toggle="collapse" data-target="#bs-example-navbar-collapse-1">
        <span class="sr-only">Toggle navigation</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>
     <img src="../_static/img/logo.png" style="width:80px; height: 80px;"> 
    </div>

    <div class="collapse navbar-collapse" id="bs-example-navbar-collapse-1" style="margin-top: 10px; margin-left: 25px;">
      <ul class="nav navbar-nav">
        <li class="active"><a href="../">Petlja - Информатика и рачунарство за шести разред</a></li>
      </ul>
    </div>
  </div>
</nav>

<div style="margin-left: 50%;margin-top: 5%;font-size: 20px;">
  
      <div class="sphinxsidebar" role="navigation" aria-label="main navigation">
        <div class="sphinxsidebarwrapper">
    <ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../1_lekcija/index.html">1. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_1_cas.html">1.1. Робот Карел и линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_2_cas.html">1.2. Робот Карел и бројачка петља <cite>for</cite></a></li>
<li class="toctree-l2"><a class="reference internal" href="../1_lekcija/1_3_cas.html">1.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../2_lekcija/index.html">2. лекција - робот Карел</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_1_cas.html">2.1. Робот Карел - условне петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_2_cas.html">2.2. Робот Карел - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../2_lekcija/2_3_cas.html">2.3. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../3_lekcija/index.html">3. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../3_lekcija/3_1_cas.html">3.1. Вежбање - робот Карел</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../4_lekcija/index.html">4. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_1_test.html">4.1. Квиз - робот Карел (линијска структура и понављање)</a></li>
<li class="toctree-l2"><a class="reference internal" href="../4_lekcija/4_2_test.html">4.2. Квиз - робот Карел (условна петља и гранање)</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../5_lekcija/index.html">5. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_1_cas.html">5.1. Корњача графика - линијски програми</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_2_cas.html">5.2. Корњача графика - понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../5_lekcija/5_3_cas.html">5.3. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../6_lekcija/index.html">6. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_1_cas.html">6.1. Корњача графика - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_2_cas.html">6.2. Корњача графика - угнежђене петље</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_3_cas.html">6.3. Корњача графика - процедуре</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_4_cas.html">6.4. Корњача графика - торке/листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../6_lekcija/6_5_cas.html">6.5. Вежбање - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../7_lekcija/index.html">7. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../7_lekcija/7_1.html">Тест - корњача графика</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../8_lekcija/index.html">8. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_1_cas.html">8.1. Основне аритметичке операције и примена</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_2_cas.html">8.2. Променљиве - имена међурезултата</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_3_cas.html">8.3. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../8_lekcija/8_4_cas.html">8.4. Вежбање - изрази</a></li>
</ul>
</li>
<li class="toctree-l1 current"><a class="reference internal" href="index.html">9. лекција</a><ul class="current">
<li class="toctree-l2"><a class="reference internal" href="9_1_cas.html">9.1. Цели и реални бројеви</a></li>
<li class="toctree-l2 current"><a class="current reference internal" href="#">9.2. Реално и целобројно дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="9_3_cas.html">9.3. Вежбање - цели и реални бројеви</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../10_lekcija/index.html">10. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_1_cas.html">10.1. Коришћење и дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_2_cas.html">10.2. Дефинисање функција</a></li>
<li class="toctree-l2"><a class="reference internal" href="../10_lekcija/10_3_cas.html">10.3. Вежбање - функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../11_lekcija/index.html">11. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_1.html">11.1. Квиз - Аритметичке операције</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_2.html">11.2.Квиз - целобројно и реално дељење</a></li>
<li class="toctree-l2"><a class="reference internal" href="../11_lekcija/11_3.html">11.3. Квиз - Функције</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../12_lekcija/index.html">12. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_1_cas.html">12.1. Гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../12_lekcija/12_2_cas.html">12.2. Комбинација више логичких услова</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../13_lekcija/index.html">13. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_1_cas.html">13.1 Понављање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_2_cas.html">13.2. Условна петља</a></li>
<li class="toctree-l2"><a class="reference internal" href="../13_lekcija/13_3_cas.html">13.3. Угнежђене петље</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../14_lekcija/index.html">14. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_1_test.html">Тест - гранање</a></li>
<li class="toctree-l2"><a class="reference internal" href="../14_lekcija/14_2_test.html">Тест - понављање</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../15_lekcija/index.html">15. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_1_cas.html">15.1. Структуре података: Торке, листе</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_2_cas.html">15.2. Ниске</a></li>
<li class="toctree-l2"><a class="reference internal" href="../15_lekcija/15_3_cas.html">15.3. Скупови</a></li>
</ul>
</li>
<li class="toctree-l1"><a class="reference internal" href="../16_lekcija/index.html">16. лекција</a><ul>
<li class="toctree-l2"><a class="reference internal" href="../16_lekcija/16_1.html">16.1. Квиз - структуре података</a></li>
</ul>
</li>
</ul>
</ul>

        </div>
      </div>
  </div>
<div class="container col-md-12" id="continue-reading"></div>

<div class="container col-md-10 col-md-offset-2" id="main-content" style="margin-top: 70px;">
  
  <div class="section" id="id1">
<h1>9.2. Реално и целобројно дељење<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>Хајде да заједно погледамо шта је то реално, а шта целобројно дељење, погледај наредни видео:</p>

    <div  id="wl2wvHeB0pQ" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/wl2wvHeB0pQ/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-wl2wvHeB0pQ" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-wl2wvHeB0pQ">
        </div>
    </div>
    
    
    <div class="section" id="id2">
<h2>Реално дељење<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>У наставку ћемо се посветити операцији дељења. У већини програмских
језика разликују се два облика дељења: <strong>реално</strong> и
<strong>целобројно</strong>. Резултат (количник) реалног дељења је реалан број (на
пример, када се 13 реално дели са 4, добија се количник 3,25), док је
код целобројног дељења количник увек цео број, при чему је могуће да
постоји и целобројни остатак (на пример, када се 13 целобројно дели са
4, добија се количник 3 и остатак 1).</p>
<p>Реалним дељењем се могу делити и реални и цели бројеви, а резултат је
реалан број. На пример, ако се користи реално дељење тада је <em>7,5 :
2,5</em> једнако <em>3</em> док је <em>5 : 2</em> једнако <em>2,5</em>. Реално дељење се у
језику Python3 обележава знаком <code class="docutils literal notranslate"><span class="pre">/</span></code>. Тако је вредност израза <code class="docutils literal notranslate"><span class="pre">7.5</span> <span class="pre">/</span>
<span class="pre">2.5</span></code> једнака <code class="docutils literal notranslate"><span class="pre">3.0</span></code>.</p>
<p>Провери своје знање наредним питањем.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill_проценат">
            <p>Вредност израза <code class="docutils literal notranslate"><span class="pre">4.5</span> <span class="pre">/</span> <span class="pre">5</span></code> је <input type="text"></p>

        <script type="application/json">
            [[{"number": [0.9, 0.9], "feedback": "<p>T\u0430\u0447\u043d\u043e</p>\n"}, {"regex": "^ *0,9 *$", "regexFlags": "", "feedback": "<p>\u0422\u0430\u0447\u043d\u043e, \u0430\u043b\u0438 \u0443\u043c\u0435\u0441\u0442\u043e \u0434\u0435\u0446\u0438\u043c\u0430\u043b\u043d\u043e\u0433 \u0437\u0430\u0440\u0435\u0437\u0430 \u0442\u0440\u0435\u0431\u0430 \u0441\u0442\u0430\u0432\u0438\u0442\u0438 \u0434\u0435\u0446\u0438\u043c\u0430\u043b\u043d\u0443 \u0442\u0430\u0447\u043a\u0443</p>\n"}, {"regex": "^ *.* *$", "regexFlags": "", "feedback": "<p>\u041d\u0435\u0442\u0430\u0447\u043d\u043e</p>\n"}]]
        </script>

        </div>
        </div>
            <p>Посматрајмо наредни једноставан задатак у којем ћемо употребити реално
дељење.</p>
<div class="section" id="id3">
<h3>Просек скокова у даљ<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Скакач у даљ је у квалификацијама у првој серији скочио 8,12, у
другој 8,23, а у трећој 8,17 метара. Колико је износио његов
просечни скок?</p>

    </p></div></div>
<p>Просек (каже се и аритметичка средина) неколико бројева једнак је
количнику њиховог збира и њиховог броја. Са просеком сте се сигурно
већ срели када сте рачунали просек својих оцена. Дакле, да бисмо решили
овај задатак потребно је сабрати дужине сва три скока и поделити са
три.</p>

<div data-childcomponent="Просек_скокова" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Просек_скокова data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
skok1 = 8.12
skok2 = 8.23
skok3 = 8.17
prosek = (skok1 + skok2 + skok3) / 3
print(prosek)
</textarea>
</div>
</div>
<div class="section" id="id5">
<h3>Сложени израз из збирке из математике<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У једној збирици из математике за шести разред јавља се задатак у
коме се тражи да се израчуна вредност израза 1 + (3 - (-4)) : 2 +
0,7. Израчунај ту вредност у Python-у.</p>

    </p></div></div>

<div data-childcomponent="Сложени_израз_са_дељењем" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Сложени_израз_са_дељењем data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(1 + (3 - (-4)) / 2 + 0.7)
</textarea>
</div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Израчунај вредност израза <span class="math notranslate nohighlight">\(7 + \frac{4 - (-5)}{(-3) \cdot 2 -
7}\)</span> у Python-у.</p>

    </p></div></div>

<div data-childcomponent="Сложени_израз_са_дељењем_1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Сложени_израз_са_дељењем_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print()  # popravi ovaj red
</textarea>
</div>
<p>Ако урадиш све како треба, добићеш резултат <code class="docutils literal notranslate"><span class="pre">6.3076923076923075</span></code>.</p>
</div>
<div class="section" id="id8">
<h3>Полице са књигама<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>На првој полици има 150 књига. На другој има дупло мање него на
првој, а на трећој три пута мање него на другој. Колико је укупно
књига на полицама.</p>

    </p></div></div>

<div data-childcomponent="Полице_са_књигама" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Полице_са_књигама data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
polica1 = 150
polica2 = polica1 / 2
polica3 = polica2 / 3
ukupno = polica1 + polica2 + polica3
print(ukupno)
</textarea>
</div>
<p>Приметимо да се као резултат добија број <code class="docutils literal notranslate"><span class="pre">250.0</span></code> уместо <code class="docutils literal notranslate"><span class="pre">250</span></code>. То
је због тога што се након примене операције реалног дељења (операције
<code class="docutils literal notranslate"><span class="pre">/</span></code>) добије увек резултат у облику реалног броја. Пошто је број
књига цео број и пошто је број 150 дељив и са 2, а број 75 са 3 (иначе
задатак не би имао смисла) на овом месту је било могуће употребити и
операцију целобројног дељења о којој ће више речи бити у наставку.</p>
</div>
</div>
<div class="section" id="id10">
<h2>Целобројно дељење<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>
<p>Целобројно дељење обично подразумева дељење целих бројева и као
резултат се одређују целобројни количник и остатак при дељењу. На
пример, ако се целобројно деле бројеви <em>14</em> и <em>3</em> тада се добија
целобројни количник <em>4</em> и остатак <em>2</em>.</p>

    <div class="rst-level rst-level-2">
<p>У општем случају, целобројни количник и остатак при дељењу бројева
<span class="math notranslate nohighlight">\(a\)</span> и <span class="math notranslate nohighlight">\(b\)</span> су бројеви <span class="math notranslate nohighlight">\(q\)</span> и <span class="math notranslate nohighlight">\(r\)</span> такви да је
<span class="math notranslate nohighlight">\(a = q \cdot b + r\)</span> и <span class="math notranslate nohighlight">\(0 \leq r &lt; b\)</span>. Приметимо да ова
веза важи у примеру дељења <span class="math notranslate nohighlight">\(14\)</span> и <span class="math notranslate nohighlight">\(3\)</span> важи управо ова
веза тј. важи да је <span class="math notranslate nohighlight">\(14 = 4 \cdot 3 + 2\)</span>, при чему је
<span class="math notranslate nohighlight">\(0 \leq 2 &lt; 3\)</span>. Други услов каже да остатак мора бити мањи од
делиоца тј. да количник мора бити што је могуће већи. Тај услов је
веома важан (на пример, важи да је <span class="math notranslate nohighlight">\(14 = 3 \cdot 3 + 5\)</span>,
међутим, нећемо рећи да је целобројни количник <span class="math notranslate nohighlight">\(3\)</span> а остатак
<span class="math notranslate nohighlight">\(5\)</span> јер број <span class="math notranslate nohighlight">\(5\)</span> није мањи од делиоца).</p>

    </div>
<p>Кроз наредно питање провери колико разумеш операције целобројног
дељења и остатка при дељењу.</p>

        <div class="course-box course-box-question course-content">
        <div data-component="fillintheblank" id="fill1412">
            <p>При дељењу бројева 13 и 5 целобројни количник је <input type="text"> а остатак је <input type="text"></p>

        <script type="application/json">
            [[{"number": [2, 2], "feedback": "<p>\u0422\u0430\u0447\u043d\u043e</p>\n"}, {"regex": "^ *x *$", "regexFlags": "", "feedback": "<p>\u0412\u0430\u0436\u0438 \u0434\u0430 \u0458\u0435 13 = 2 \u00b7 5 + 3</p>\n"}], [{"number": [3, 3], "feedback": "<p>\u0422\u0430\u0447\u043d\u043e</p>\n"}, {"regex": "^ *x *$", "regexFlags": "", "feedback": "<p>\u0412\u0430\u0436\u0438 \u0434\u0430 \u0458\u0435 13 = 2 \u00b7 5 + 3</p>\n"}]]
        </script>

        </div>
        </div>
            <p>У језику Python3 операција <strong>целобројног дељења</strong> се означава са
<code class="docutils literal notranslate"><span class="pre">//</span></code>, а операција израчунавања <strong>остатка при дељењу</strong> се означава са
<code class="docutils literal notranslate"><span class="pre">%</span></code>.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>У математици се знак % користи да означи проценат (стоти део
нечега). Коришћење истог знака за остатак при дељењу је заправо
несрећна околност и треба бити обазрив да се та два заправо
неповезана појма случајно не помешају.</p>

    </p></div></div>
<p>Дакле, оператором <code class="docutils literal notranslate"><span class="pre">/</span></code> се израчунава реални, оператором <code class="docutils literal notranslate"><span class="pre">//</span></code>
целобројни количник, а оператором <code class="docutils literal notranslate"><span class="pre">%</span></code> остатак при дељењу. Провери колико
ово разумеш.</p>

<div class="course-box course-box-problem course-content">
<ul data-component="dragndrop" id="дељење">
    <span data-component="question">Q-86: Превлачењем упари изразе са њиховим вредностима.</span>
	<span data-component=feedback>Покушај поново</span>

    <li data-component="draggable" id="дељење_drag1">27 / 10</li>
    <li data-component="dropzone" for="дељење_drag1">2.7</li>

    <li data-component="draggable" id="дељење_drag2">27 // 10</li>
    <li data-component="dropzone" for="дељење_drag2">2</li>

    <li data-component="draggable" id="дељење_drag3">27 % 10</li>
    <li data-component="dropzone" for="дељење_drag3">7</li>
</ul></div>
<div class="course-box course-box-problem course-content">
<ul data-component="dragndrop" id="дељење1">
    <span data-component="question">Q-87: Упари изразе са њиховим вредностима.</span>
	<span data-component=feedback>Покушај поново</span>

    <li data-component="draggable" id="дељење1_drag1">43 / 8</li>
    <li data-component="dropzone" for="дељење1_drag1">5.375</li>

    <li data-component="draggable" id="дељење1_drag2">43 // 8</li>
    <li data-component="dropzone" for="дељење1_drag2">5</li>

    <li data-component="draggable" id="дељење1_drag3">43 % 8</li>
    <li data-component="dropzone" for="дељење1_drag3">3</li>
</ul></div><p>Покажимо једноставну примену израчунавања целобројног количника и остатка
на следећем задатку.</p>
<div class="section" id="id11">
<h3>Подела чоколадних бананица<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>У школи се организује новогодишња приредба за децу. Од пара које су
зарадили тако што су организовали сајам својих рукотворина купили
су неколико крем бананица које желе да равномерно поделе свој деци
(тако да свако дете добије исти број бананица). Ако се зна колико
ће деце доћи на приредбу, колико ће свако дете добити бананица, a
колико ће бананица остати нерасподељено?</p>

    </p></div></div>

<div data-childcomponent="чоколадне_бананице" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=чоколадне_бананице data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
broj_dece = int(input("Koliko će dece doći na priredbu: "))
ukupno_bananica = int(input("Koliko ukupno ima bananica: "))
bananica_po_detetu = ukupno_bananica // broj_dece
ostalo_bananica = ukupno_bananica % broj_dece
print("Svako će dete dobiti", bananica_po_detetu, "bananica.")
print("Ostaće", ostalo_bananica, "bananica.")
</textarea>
</div>
<p>Приметимо и да смо број преосталих бананица могли израчунати и тако
што од укупног броја бананица одузмемо број бананица које су подељене
деци (а то је производ броја деце и броја бананица које је свако дете
добило) тј. помоћу израза <code class="docutils literal notranslate"><span class="pre">ukupno_bananica</span> <span class="pre">-</span> <span class="pre">broj_dece</span> <span class="pre">*</span>
<span class="pre">bananica_po_detetu</span></code>. Ипак, коришћење оператора <code class="docutils literal notranslate"><span class="pre">%</span></code> којим се
израчунава остатак је једноставније решење.</p>
<p>Ево једног сличног задатка, за вежбу.</p>
</div>
<div class="section" id="id13">
<h3>Разломак у мешовити број<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Бројилац разломка је 37, а именилац је 12. Преведи овај разломак у
мешовит број.</p>

    </p></div></div>
<p>Важи да је <span class="math notranslate nohighlight">\(37 = 3 \cdot 12 + 1\)</span>, па је <span class="math notranslate nohighlight">\(\frac{37}{12} =
\frac{3 \cdot 12 + 1}{12} = 3 \frac{1}{12}\)</span>. У општем случају када
разломак <span class="math notranslate nohighlight">\(\frac{a}{b}\)</span> преводимо у мешовит број потребно је да
бројилац напишемо у облику <span class="math notranslate nohighlight">\(a = q \cdot b + r\)</span>, при чему мора да
важи да је <span class="math notranslate nohighlight">\(0 \leq r &lt; b\)</span> и тада се добија межовити број
<span class="math notranslate nohighlight">\(q \frac{r}{b}\)</span>. Број <span class="math notranslate nohighlight">\(q\)</span> је целобројни количник бројева
<span class="math notranslate nohighlight">\(a\)</span> и <span class="math notranslate nohighlight">\(b\)</span>, док је <span class="math notranslate nohighlight">\(r\)</span> остатак при њиховом дељењу.</p>

<div data-childcomponent="Мешовит_број" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Мешовит_број data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
brojilac = 37
imenilac = 12
mesoviti_ceo_deo = 0  # ispravi ovaj red
mesoviti_brojilac = 0 # ispravi ovaj red
mesoviti_imenilac = 0 # ispravi ovaj red
print(mesoviti_ceo_deo, "celih i", mesoviti_brojilac, "/", mesoviti_imenilac)
</textarea>
</div>
<p>Наравно, резултат треба да буде <code class="docutils literal notranslate"><span class="pre">3</span> <span class="pre">celih</span> <span class="pre">i</span> <span class="pre">1</span> <span class="pre">/</span> <span class="pre">12</span></code>.</p>

    <div data-component="reveal" id="пресек_решење31"   data-showtitle="Прикажи решење" data-hidetitle="Сакриј решење">
    
<div data-childcomponent="Мешовит_број_решење" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Мешовит_број_решење data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
brojilac = 37
imenilac = 12
mesoviti_ceo_deo = brojilac // imenilac
mesoviti_brojilac = brojilac % imenilac
mesoviti_imenilac = imenilac
print(mesoviti_ceo_deo, "celih i", mesoviti_brojilac, "/", mesoviti_imenilac)
</textarea>
</div>

    </div>
    </div>
</div>
<div class="section" id="id16">
<h2>Целобројно дељење - конверзија јединица<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h2>
<p>Целобројни количник и остатак често користимо када желимо да
прерачунавамо јединице. Размотримо следећих неколико задатака.</p>
<div class="section" id="id17">
<h3>Конверзија центиметара у метре и центиметре<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Напиши програм који на основу дате дужине у центиметрима израчунава
исту дужину у метрима и центиметрима. На пример, ако је дужина 178
центиметара, програм израчунава да је то 1 метар и 78 центиметара.</p>

    </p></div></div>
<p>Пошто у једном метру има 100 центиметара, задатак се своди на
израчунавање целобројног количника и остатка при дељењу
са 100. Заиста, ако имамо <span class="math notranslate nohighlight">\(m\)</span> метара и <span class="math notranslate nohighlight">\(c\)</span> центиметара,
тада је укупан број центиметара једнак <span class="math notranslate nohighlight">\(m\cdot 100 + c\)</span>, при
чему је <span class="math notranslate nohighlight">\(0 \leq c &lt; 100\)</span>.</p>

<div data-childcomponent="центиметри_у_метре_и_центиметре" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=центиметри_у_метре_и_центиметре data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
ukupno_centimetara = int(input("Unesi dužinu u centimetrima: "))
metara = ukupno_centimetara // 100
centimetara = ukupno_centimetara % 100
print("Dužina je", metara, "m", centimetara, "cm")
</textarea>
</div>
<p>Рецимо поново да је често решење до којег ученици самостално долазе и
оно у којем се преостали број центиметара рачуна као <code class="docutils literal notranslate"><span class="pre">centimetara</span> <span class="pre">=</span>
<span class="pre">ukupno_centimetara</span> <span class="pre">-</span> <span class="pre">metara</span> <span class="pre">*</span> <span class="pre">100</span></code>. Иако је ово решење исправно, на
располагању нам је оператор израчунавања остатака <code class="docutils literal notranslate"><span class="pre">%</span></code> и требало би
да се навикнемо да га користимо.</p>
</div>
</div>
<div class="section" id="id19">
<h2>Целобројно дељење - позициони запис броја<a class="headerlink" href="#id19" title="Permalink to this headline">¶</a></h2>
<p>Ако је укупан број центиметара био 123, тада је број метара 1, број
дециметара 2 и број центиметара 3. Решавањем претходног задатка смо
заправо одређивали појединачне цифре коришћене у запису тог
троцифреног броја. Приказана техника може бити уопштена тако да се
одређују све цифре и у запису дужих бројева.</p>
<div class="section" id="id20">
<h3>Цифре броја<a class="headerlink" href="#id20" title="Permalink to this headline">¶</a></h3>

    <div data-level="2" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Чест начин откривања грешака при слању података је да се уз податке
које треба послати, пошаљу и одређени контролни подаци, израчунати
на основу самих података. Када прималац прими податке, он на основу
примљених података поново израчунава контролне податке и упоређује
их са контролним подацима које је примио. Ако се приликом преноса
података, услед неких сметњи, подаци случајно измене, прималац ће
то приметити тако што ће видети да се контролни подаци које је он
израчунао неће поклопити са онима које је примио. Сви подаци се у
рачунарима представљају помоћу бројева, а још од најранијег доба
рачунарства као метода контроле коришћен је збир цифара. На пример,
ако би податак који се шаље био број 12345, онда би се уз њега слао
и контролни број 15 (збир 1 + 2 + 3 + 4 + 5). Ако би приликом слања
нека цифра била случајно промењена (на пример, ако би прималац
грешком примио број 12335) тада би се и контролни број
највероватније разликовао (контролни број би у нашем примеру био
14). Напиши програм који за дати петоцифрени број одређује његов
контролни број (збир његових цифара).</p>

    </p></div></div>

<div data-childcomponent="цифре_броја" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=цифре_броја data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="broj kontrolni_broj"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
broj = 12345
# -*- acsection: main -*-
c0 = (broj // 1) % 10
c1 = (broj // 10) % 10
c2 = 0 # ispravi ovaj red
c3 = (broj // 1000) % 10
c4 = 0 # ispravi ovaj red
kontrolni_broj = c0 + c1 + c2 + c3 + c4
# -*- acsection: after-main -*-
print(kontrolni_broj)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for broj, kontrolni_broj in [(71425, 19), (33214, 13), (62040, 12)]:
          self.assertEqual(acMainSection(broj = broj)["kontrolni_broj"],kontrolni_broj,"За број %s контролни број је %s." % (broj, kontrolni_broj))
myTests().main()
</textarea>
</div>
<p>Приметимо да смо цифре броја одређивали веома слично као у претходном
задатку.  Цифре одређујемо сдесна налево, тако што делимо број са
тежином цифре (за цифру јединица број делимо са 1, десетица са 10,
стотина са 100 итд.) и проналазимо остатак при дељењу са 10.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Људи су контролу података примењивали и ручно, а сада рачунари
обављају такве контроле у овиру прецизно задатих поступака –
<em>протокола за размену података</em>. Иако је контролни збир
најједноставнији начин контроле, он не може да открије грешке до
којих може доћи услед случајне размене редоследа цифара (на пример,
ако се уместо броја 12345 грешком пошаље број 12435, контролни збир
оба броја ће бити исти и грешка неће бити примећена). Зато се
уместо збира понекад користе изрази облика <span class="math notranslate nohighlight">\(c_0 + 2c_1 +
3c_2 + 4c_3 + 5c_4\)</span>. Покушај да измениш претходни програм тако да
израчунава контролни број на овај начин.</p>

    </p></div></div>
</div>
</div>
<div class="section" id="id22">
<h2>Целобројно дељење - време и углови<a class="headerlink" href="#id22" title="Permalink to this headline">¶</a></h2>
<p>За разлику од бројева и јединица мере које записујемо у систему чија
је основа 10, при раду са временом и угловима користимо систем чија је
основа број 60. Тако један сат има 60 минута, а један минут 60
секунди. Слично један степен има 60 угаоних минута, а један угаони
минут има 60 угаоних секунди. Прикажимо сада кроз неколико задатака
како можемо у програмима вршити израчунавања у којима учествују време
и углови.</p>
<div class="section" id="id23">
<h3>Конверзија сати и минута у минуте и обратно<a class="headerlink" href="#id23" title="Permalink to this headline">¶</a></h3>

    <div data-level="1" style="display:none">

    </div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Ако се зна колико је тренутно сати и минута, израчунај колико је
минута протекло од претходне поноћи.</p>

    </p></div></div>
<p>Пошто у једном сату има 60 минута, довољно је да помоножиш број сати
са 60 и на то да додаш број минута.</p>

<div data-childcomponent="сати_и_минути_у_минуте" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=сати_и_минути_у_минуте data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="sati minuta minuta_od_ponoci"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
sati = 2
minuta = 60
# -*- acsection: main -*-
minuta_od_ponoci = 0 # ispravi ovaj red
# -*- acsection: after-main -*-
print(minuta_od_ponoci)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for sati, minuta, minuta_od_ponoci in [(14, 19, 859), (11, 13, 673), (23, 59, 1439)]:
          self.assertEqual(acMainSection(sati = sati, minuta = minuta)["minuta_od_ponoci"],minuta_od_ponoci,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
myTests().main()
</textarea>
</div>

    <div class="course-box course-box-special">
        <div class="course-content">
            <h4 class="carbox-title">
                <img class="corner-image float-right" src="../_static/img/question-mark.png" />
            </h4>
            <p>
<p>Ако се зна колико је минута протекло од претходне поноћи, израчунај
колико је тренутно сати и минута.</p>

    </p></div></div>
<p>Ако са <span class="math notranslate nohighlight">\(s\)</span> обележимо тренутни број сати, са <span class="math notranslate nohighlight">\(m\)</span> тренутни
број минута, а са <span class="math notranslate nohighlight">\(M\)</span> број минута протеклих од поноћи, тада важи
да је <span class="math notranslate nohighlight">\(M = s \cdot 60 + m\)</span>, при чему за <span class="math notranslate nohighlight">\(m\)</span> важи да је
број између <span class="math notranslate nohighlight">\(0\)</span> и <span class="math notranslate nohighlight">\(59\)</span>, што јасно указује на то да се
тражене вредности могу израчунати применом целобројног дељења.</p>

<div data-childcomponent="минути_у_сате_и_минуте" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=минути_у_сате_и_минуте data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          data-runortest="minuta_od_ponoci sati minuta"     data-hidehistory=true
       data-enablecopy="true">
# -*- acsection: general-init -*-
# -*- acsection: var-init -*-
minuta_od_ponoci = 125
# -*- acsection: main -*-
sati = 0     # ispravi ovaj red
minuta = 0   # ispravi ovaj red
# -*- acsection: after-main -*-
print(sati, minuta)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
       for sati, minuta, minuta_od_ponoci in [(14, 19, 859), (11, 13, 673), (23, 59, 1439)]:
          self.assertEqual(acMainSection(minuta_od_ponoci = minuta_od_ponoci)["sati"],sati,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
          self.assertEqual(acMainSection(minuta_od_ponoci = minuta_od_ponoci)["minuta"],minuta,"У %s:%s протекло је %s минута од поноћи." % (sati, minuta, minuta_od_ponoci))
myTests().main()
</textarea>
</div>
</div>
</div>
</div>


  
      <div class="col-md-12">
<ul class="pager">
    <li id="relations-prev" title='Previous chapter - 9.1. Цели и реални бројеви' data-toggle="tooltip"><a href="9_1_cas.html"><i class="fas fa-arrow-left"></i></a></li>
    
        <li id="relations-next" title='Next chapter - 9.3. Вежбање - цели и реални бројеви' data-toggle="tooltip"><a href="9_3_cas.html"><i class="fas fa-arrow-right"></i></a></li>
</ul>

<!-- <ul class="pager"> -->
    <!-- -->
        <!-- <li id="relations-prev" title='Претходно поглавље - 9.1. Цели и реални бројеви' data-toggle="tooltip"><a href="9_1_cas.html">Претходно поглавље</a></li> -->
    <!--  -->
    <!-- -->
        <!-- <li id="relations-next" title='Следеће поглавље - 9.3. Вежбање - цели и реални бројеви' data-toggle="tooltip"><a href="9_3_cas.html">Следеће поглавље</a></li> -->
    <!-- -->
<!-- </ul> -->

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});

</script>
</div>
  
</div>
<footer class="footer col-md-12">
    <div class="container">
        <div class="text-center">
            <hr>
            <p class="text-muted">
                <span class="pull-left">&copy; 2019 Petlja (Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>)</span>
            </p>
        </div>
    </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>