Структуре података: Торке, листе, ниске, скупови, речници
##############################################

У свим програмима које смо до сада писали податке смо чували у
појединачним променљивама. То су били обично појединачни бројеви или
евентуално краћи делови текста (нпр. име или презиме).  Елементарни
подаци које смо до сада сретали били су цели бројеви (типа ``int``),
реални бројеви (типа ``float``), логичке вредности (типа ``bool``) и
ниске (типа ``str``). У многим реалним применама имамо потребу да у
програму запамтимо, а касније и обрадимо веће количине података. У тим
ситуацијама коришћење појединачних променљивих за памћење сваког
појединачног података није практично или, чак, није могуће и потребно је
користити **структуре података** које нам омогућавају да помоћу једне
променљиве представимо читаве **колекције** података, које се састоје
од више вредности неког једноставнијег типа. У зависности од података
које желимо да сачувамо и операција које са њима желимо да спроведемо,
бирамо коју од неколико различитих врста структура података које имамо
на располагању треба да употребимо.

Торке и листе
-------------

Запитајмо се како бисмо у програму могли представити списак имена
кошаркаша у једном тиму? Нека је то наша кошаркашка репрезентација
која се 2019. такмичила на светском првенству у
кошарци. Претпоставићемо да сваки играч има свој редни број и то од 1
до 12.

Један могући начин би био да уведемо пуно засебних променљивих:

.. activecode:: кошаркаши_без_листе
   :passivecode: true

   igrac1 = "Никола Јокић"
   igrac2 = "Богдан Богдановић"
   igrac3 = "Немања Бјелица"
   igrac4 = "Василије Мицић"
   ...

Међутим, оваква репрезентација података није згодна јер не можемо
једноставно да одговоримо на разна интересантна питања о кошаркашима
(на пример, колико укупно има играча, ко носи број који се уноси са
улаза, који број носи неки дати играч и слично). Бољи начин да
представимо списак играча је да употребимо **торку** или **листу**.
На пример, торку имена кошаркаша можемо дефинисати на следећи начин.

.. activecode:: кошаркаши_торка
   :passivecode: true
                 
   igraci = ("Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов")

.. infonote::
   
   Запис облика (35, 128) се у математици назива *уређени пар*, запис
   облика (250, 120, 310) *уређена тројка*, запис облика (135, 22,
   83, 57) *уређена четворка* итд. Када дужина није позната, каже се
   *уређена n-торка* (*уређена енторка*), али у програмирању је
   одомаћено краће име - *торка*.
             
Дакле, торку дефинишемо тако што између обичних малих заграда (заграда
``(`` и ``)``) наводимо елементе раздвојене зарезима. Приметимо да је
листа имена кошаркаша веома дугачка, тако да смо је прегледности ради
разбили у више редова. Иако у општем случају није допуштено делити
наредбе на више линија, то је могуће ако је садржај тих линија окружен
неким заградама (пошто је у првој линији отворена заграда ``(``, након
читања прве, аутоматски се читају и друга и трећа и четврта линија, све
док се не наиђе на затворену заграду ``)``).

Дефинисање листе би се разликовало од дефинисања торке само по томе
што би се уместо обичних малих заграда ``(...)`` употребиле угласте
заграде употребиле обичне заграде ``[...]``.


.. activecode:: кошаркаши_листа
   :passivecode: true
                 
   igraci = ["Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов"]


Суштинска разлика између торке и листе је то што је торка
**имутабилна** што значи да јој не можемо мењати вредности када је
једном дефинишемо. Са друге стране, листе допуштају да им мењамо
појединачне елементе, уклањамо елементе, додајемо нове и слично. Све
операције у којима се врши анализа садржаја биће идентичне, без обзира
да ли се ради о листи или о торки и описаћемо их у наставку (иако ћемо
у примерима користити торке, све оне се на потпуно исти начин могу
применити и на листе). Операције у којима се врши промена су
специфичне за листе и није их могуће спровести над торкама па ћемо њих
описати засебно.
             
Појединачни чланови торке или листе називају се и
**елементи**. Елементи могу бити ниске (као у овом примеру), али и
бројеви, друге торке или листе и било шта друго. Чак је могуће у исту
торку тј. листу поставити елементе различитог типа. Као пример торке
бројева, формирајмо торку висина ових наших кошаркаша.

.. activecode:: висине_кошаркаша
   :passivecode: true

   visine = (198, 197, 198, 198, 203, 203, 210, 208, 213, 221, 213, 212)

Двочлане торке (уређени парови) и вишечлане торке ( уређене тројке,
четворке итд.) се најчешће користе за груписање тек неколико
појединачних података. На пример, позицију фигуре на шаховској табли
можемо представити помоћу уређеног пара који чини ознака врсте (слова
од ``a`` до ``h``) и ознака колоне (броја од 1 до 8). На пример,
``("b", 6)``. Слично, позиције на географској карти се описују помоћу
географске ширине и дужине тј. помоћу пара реалних бројева.  Тако се
град Париз налази на позицији која се може описати помоћу пара
``(48.8566, 2.3522)``. Време можемо представити помоћу уређене тројке
коју чине сат, минут и секунд (на пример, ``(7, 25, 37)``).

За груписање већег броја података чешће се користе листе, поготово ако
није унапред познато колико таквих података може бити тј. ако се број
података може изменити током рада програма. На пример, ако желимо да
представимо оцене ученика, чешће бисмо употребили листу јер ученик
може да добије још неку оцену, па је листу његових оцена потребно
проширити, што није могуће ако се оцене представе торком.

.. infonote::

   Већина програмских језика не разликује имутабилне и мутабилне
   колекције. Зато програмери често користе искључиво листе, чак и у
   ситуацијама када је коришћење торки примереније. То се не сматра
   великим пропустом.

   
Издвајање елемената
&&&&&&&&&&&&&&&&&&&

Елементима торке (или листе) приступа се на основу њихове **позиције**
тј. **индекса**. Бројање креће од нуле. Први елемент листе се налази
на позицији 0, други на позицији 1 и тако даље.  На пример, елементи
у листи ``a = [3, 7, 4, 2, 5]`` се броје на следећи начин.

+------------+-+-+-+-+-+
|индекс      |0|1|2|3|4|
+------------+-+-+-+-+-+
|вредност    |3|7|4|2|5|
+------------+-+-+-+-+-+

Зато је првом елементу могуће приступити са ``a[0]`` и вредност тога
израза је ``3`` (јер се на позицији 0 у листи налази број ``3``),
другом елементу са ``a[1]`` и вредност тог израза је ``7`` и тако
даље.

Негативни индекси упућују на бојање од краја (с десна на лево), па -1
означава последњи елемент, -2 претпоследњи и тако даље.

+--------+---+---+---+---+--+
|индекс  | -5| -4| -3| -2|-1|
+--------+---+---+---+---+--+
|вредност| 3 | 7 | 4 | 2 | 5|
+--------+---+---+---+---+--+

Зато је последњем елементу могуће приступити помоћу ``a[-1]``,
претпоследњем помоћу ``a[-2]`` итд.

Још једна операција која је често корисна је издвајање дела торке
тј. листе.  Елементима између позиција ``a`` и ``b`` укључујући
позицију ``a``, али искључујући позицију ``b`` могуће је приступити
помоћу ``a[a:b]``.

На пример, ``ime[1:4]`` издваја елементе имена на позицијама 1, 2 и 3
(распон ``[1:4]`` је полуотворен тј. позиција 1 је урачуната, а
позиција 4 није). Ако је ``a = [3, 7, 4, 2, 5]`` тада је ``a[1:4]``
листа чија је вредност ``[7, 4, 2]``. Ако се изостави број иза
двотачке, издвајају се елементи до краја. На пример ``a[2:]`` је листа
која има вредност ``[4, 2, 5]``.

Употребимо ове операције да решимо неколико задатака о нашим играчима.

Играч са датим бројем дреса
'''''''''''''''''''''''''''
.. level:: 1

.. questionnote::

  Познат је списак играча у тиму. Они носе дресове са бројевима од 1
  па на даље. Напиши програм који за дати број дреса одређује играча
  који игра под тим редним бројем.

Пошто се уносе бројеви од 1 до дужине торке (или листе), а индекси се
крећу од 0 па до претходника дужине торке, након учитавања броја дреса
приступиће се елементу чији је индекс за 1 мањи од броја дреса.
   
.. activecode :: висина_играча_са_датим_бројем

   igraci = ("Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов")
   dres = int(input("Број дреса: "))
   print(igraci[dres - 1])

Начин да се избегне ово померање индекса за 1 био би да се на почетно
место убаци неки вештачки елемент (на пример, празна ниска).

.. activecode :: играч_са_датим_бројем_1

   igraci = ("", "Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов")
             
   dres = int(input("Број дреса: "))
   print()   # dopuni ovaj red

   
Играчи са датим бројевима дреса
'''''''''''''''''''''''''''''''
.. level:: 1

   
.. questionnote::

   Испиши имена играча са бројевима дреса 2, 3 и 4, као и име играча
   који има највећи број дреса у тиму.

Када би се дресови бројали од 0, као што се броје индекси, потребно би
било издвојити елементе у распону ``[2:5]`` (то обухвата тачно индексе
2, 3 и 4). Пошто се места броје од 1, потребно је вредности наведене у
распону умањити за 1 тј. употребити распон ``[1:4]``. Последњи играч
тј. последњи елемент листе се увек налази на позицији ``-1``.
   
.. activecode:: играчи_са_датим_дресовима
                
   igraci = ("Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов")
             
   print(igraci[0])   # ispravi ovaj red
   print(igraci[0])   # ispravi ovaj red

Ако исправиш програм како треба, требало би да добијеш одговоре
``['Василије Мицић', 'Богдан Богдановић', 'Марко Гудурић']`` и
``Никола Милутинов``.
   
.. questionnote::

   Ако је на прво место у листи играча уписана празна листа, напиши
   програм који одређује играче са бројевима 7, 8 и 9, као и
   претпоследњег играча у тиму.
	     
.. activecode :: играчи_са_датим_дресовима_2

   igraci = ("Стефан Јовић", "Василије Мицић", "Богдан Богдановић",
             "Марко Гудурић", "Марко Симоновић", "Владимир Лучић",
	     "Стефан Бирчевић", "Немања Бјелица", "Никола Јокић",
	     "Бобан Марјановић", "Мирослав Радуљица", "Никола Милутинов")
             
   print(igraci[0])  # ispravi ovaj red
   print(igraci[0])  # ispravi ovaj red

Ако урадиш све како треба, требало би да добијеш одговоре ``["Стефан
Бирчевић", "Немања Бјелица", "Никола Јокић"]`` и ``Мирослав
Радуљица``.

Распакивање
&&&&&&&&&&&

Ако знамо дужину торке или листе, могуће је на лак начин именовати
сваки њен елемент тј. сместити сваки елемент у посебну променљиву.

На пример, претпоставимо да торка ``matematicari`` садржи пуна имена
четири велика математичара: Ренеа Декарта, Жозефа Луја Лагранжа, Карла
Фридриха Гауса и Леонарда Ојлера. Направимо четири посебне променљиве
``dekart``, ``lagranz``, ``gaus`` и ``ojler`` које ће садржати пуна
имена одговарајућих математичара.

Један начин да се то уради је да се помоћу индекса приступи појединачним
елементима. 

.. activecode:: математичари

   matematicari = ["Рене Декарт", "Жозеф Луј Лагранж", "Карл Фридрих Гаус", "Леонард Ојлер"]
   dekart = matematicari[0]
   lagranz = matematicari[1]
   gaus = matematicari[2]
   ojler = matematicari[3]
		
Међутим, постоји и једноставнији начин да се постигне исти ефекат.
   
.. activecode:: распакивање_листе

   matematicari = ["Рене Декарт", "Жозеф Луј Лагранж", "Карл Фридрих Гаус", "Леонард Ојлер"]
   dekart, lagranz, gaus, ojler = matematicari
   print(gaus)

Првој променљивој са леве стране се додељује први елемент листе са
десне стране, другој променљивој други елемент листе и тако даље. Овај
поступак називамо **распакивање**. Исти ефекат би се постигао и да су
четири променљиве са леве стране биле наведене у склопу торке
тј. листе.

.. activecode:: распакивање_листе_1

   matematicari = ["Рене Декарт", "Жозеф Луј Лагранж", "Карл Фридрих Гаус", "Леонард Ојлер"]
   [dekart, lagranz, gaus, ojler] = matematicari
   print(gaus)
   
Распакивање се може применити и на уређене парове и торке. Примени то
у решавању наредног једноставног задатка.

Географске координате
'''''''''''''''''''''
.. level:: 1

.. questionnote::

   Уређени пар садржи географске координате града Париза. Напиши програм
   који одређује и посебно исписује његову географску ширину и географску
   дужину.

Задатак можемо решити очитавањем елемената на основу индекса.
           
.. activecode:: географске_координате_града

   Pariz = (48.8566, 2.3522)
   sirina = ???
   duzina = ???
   print("Географска ширина:", sirina)
   print("Географска дужина:", duzina)

Реши сада задатак применом распакивања.

.. activecode:: географске_координате_града_1

   Pariz = (48.8566, 2.3522)
   ???
   print("Географска ширина:", sirina)
   print("Географска дужина:", duzina)


Функције за рад са торкама и листама
&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&

Језик Python нам на располагање ставља неколико функција којима можемо
одредити често коришћене статистике елемената листе. Функцијом ``len``
израчунавамо дужину листе, функцијом ``sum`` израчунавамо збир
елемената листе, функцијом ``min`` најмањи елемент у листи, а
функцијом ``max`` највећи. Ове функције можемо употребити да
једноставно решимо разне практичне задатке.

Просечна оцена
''''''''''''''
.. level:: 1

.. questionnote::

   Дате су оцене из неколико предмета. Израчунај просечну оцену.

Просек тј. аритметичка средина се једноставно добија дељењем збира
елемената њиховим бројем.
   
.. activecode:: просек_оцена
		
   ocene = [5, 4, 5, 3, 5]
   prosek = sum(ocene) / len(ocene)
   print(prosek)

Победнички скок
'''''''''''''''
.. level:: 1

.. questionnote::

   На Олимпијским играма у Рију наша атлетичарка Ивана Шпановић је
   скакала редом 6,95m, затим у наредне две серије преступила, а затим
   скакала, 6,91m, 7,08m и 7,05m. Одреди дужину скока (у метрима) који
   јој је донео бронзану медаљу.
   

Да би се ордредио најдужи скок потребно је одредити највећи елемент
који садржи дата торка.
           
.. activecode:: победнички_скок
		
   skokovi = (6.95, 0.0, 0.0, 6.91, 7.08, 7.05)
   # ispravi naredni red
   print()

Наравно, програм ако је исправан треба да испише ``7.08``.
   
Најмања оцена
'''''''''''''
.. level:: 1
   
.. questionnote::

   Ако се зна да су оцене из природних наука последње три у листи
   оцена, израчунај Горанову најмању оцену из тих предмета.

У овом задатку комбиновање идвајање сегмента елемената торке и
функцију за проналажење минимума.
   
.. activecode:: последње_оцене

   ocene = (5, 4, 5, 3, 5, 4, 4, 5)
   ocene_iz_prirodnih_nauka = ocene[-3:] 
   print(min(ocene_iz_prirodnih_nauka))

Распон температура
''''''''''''''''''
.. level:: 1

.. questionnote::

   Време се често мења и дешава се да се у једној недељи смењују и
   хладни и топли дани. Ако су дате су температуре у подне у данима
   током једне недеље, одреди колики је распон температура тј. разлика
   између највише и најниже подневне температуре.


.. activecode:: распон_температура
   
   temperature = [17, 23, 12, 15, 19, 21, 25]
   print(????)  # ispravi ovaj red


Претрага
&&&&&&&&

Проверу да ли елемент припада торки тј. листи могуће је извршити
коришћењем оператора ``in`` (на пример, ``"Бобан Марјановић" in
igraci``). У наредном примеру проверавамо да ли је кошаркаш чије се
име и презиме учитавају при покретању програма играо за репрезентацију
на СП 2019. године.

.. activecode:: Калинић_не_игра
   
   igraci = ["", "Стефан Јовић", "Марко Симоновић", "Богдан Богдановић",
             "Никола Калинић", "Милан Мачван", "Стефан Марковић",
	     "Немања Недовић", "Мирослав Радуљица", "Милош Теодосић",
	     "Никола Јокић", "Владимир Штимац", "Стефан Бирчевић"]
   igrac = input("Унеси име и презиме кошаркаша:")
   if igrac in igraci:
       print(igrac, "је играо за репрезентацију")
   else:
       print(igrac, "није играо за репрезентацију")


Проналажење најмање позиције (индекса) на којој се налази неки тражени
елемент може се урадити коришћењем методе ``index``, којој се као
једини параметар прослеђује елемент који се тражи.

.. infonote::

   Методе су посебан облик функција које се позивају у облику
   ``struktura.funkcija(parametri)``, уместо у облику
   ``funkcija(struktura, parametri)``. Дакле, ако у листи ``igraci``
   тражимо Николу Јокића, уместо да наведемо ``index(igraci, "Никола
   Јокић")``, навешћемо ``igraci.index("Никола Јокић")``. Приметимо да
   су све наредбе корњачи заправо биле методе (користили смо
   ``turtle.forward(100)``, а не ``forward(turtle, 100)``.

Број дреса датог играча
'''''''''''''''''''''''
.. level:: 1

.. questionnote::

   Који број дреса носи Никола Јокић?

Кроз овај задатак можемо илустровати претрагу тј. проналажење индекса
(позиције) у торци тј. листи на којој се налази тражени елемент. Један
начин да се то уради је помоћу методе ``index`` која враћа индекс
(подсетимо се, индекси се броје од 0) првог појављивања траженог
елемента тј. проузрокује грешку ако се тај елемент не налази у листи.

Ако торку не допунимо празним елементом на позицији 0, пошто се
дресови броје од 1, а индекси од 0, пронађени индекс је потребно
увећати за 1. 

.. activecode:: Николин_број

   igraci = ["Стефан Јовић", "Марко Симоновић", "Богдан Богдановић", \
             "Никола Калинић", "Милан Мачван", "Стефан Марковић", \
	     "Немања Недовић", "Мирослав Радуљица", "Милош Теодосић",\
	     "Никола Јокић", "Владимир Штимац", "Стефан Бирчевић"]
   print(igraci.index("Никола Јокић") + 1)


Надовезивање
&&&&&&&&&&&&

Две торке тј. листе можемо надовезати (спојити у једну) коришћењем
операције ``+``. Тиме се формира нова торка тј. листа, док се
оригинали не мењају. Није могуће надовезати торку и листу, нити листу
и торку.

Висине девојчица и дечака у одељењу
'''''''''''''''''''''''''''''''''''
.. level:: 1

.. questionnote::

 Познате су висине девојчица и висине дечака у једном одељењу. Направи
 торку свих висина и израчунај затим број и просечну висину свих ђака.

.. activecode:: спајање_листа
 
   visine_devojcica = [165, 153, 155, 155, 157]
   visine_decaka = [170, 168, 173, 156, 172]
   visine = visine_devojcica + visine_decaka
   print(len(visine))
   print(sum(visine) / len(visine))

Слично као што код бројева множење представља узастопно сабирање, исти
је случај и овде и може се наслутити да множење ниске природним бројем
заправо представља њено понављање одређени број пута. На пример, ако
желимо да направимо листу која садржи 100 нула, то можео урадити
изразом ``[0] * 100``, док израз ``(0, 1, 2) * 3`` гради торку ``(0,
1, 2, 0, 1, 2, 0, 1, 2)``.
   
Сортирање
&&&&&&&&&

Елементе торке тј. листе је веома једноставно уредити од најмањег до
највећег (па и обратно, од највећег до најмањег). Функција ``sorted``
се може применити и на торке и на листе и она на основу датих података
гради нову торку тј. листу која садржи исте податке, у сортираном
редоследу. Метода ``sort`` применљива је само на листе и она мења
листу тако што јој елементе rеђа у сортираном редоследу. Ово је
ефикасније него да се прави нова листа, али се оригинални редослед
елемената губи.


Три најјефтинија и најскупља производа
''''''''''''''''''''''''''''''''''''''
.. level:: 1

.. questionnote::

   Дата је листа цена производа. Колико коштају три најјефтинија, а
   колико три најскупља производа?

.. activecode:: најјефтинији_и_најскупљи_производи

   cene = (58.00, 104.95, 117.50, 11.95, 10.4, 37.95, 85.5)
   sortirane_cene = sorted(cene)
   print(sum(sortirane_cene[0:3]))
   print(sum(sortirane_cene[-3:]))

   
Сортирање имена ученика
'''''''''''''''''''''''
.. level:: 1
   
.. questionnote::

   Наставница треба да у дневник унесе имена ученика, међутим, од
   педагога је добила списак ученика који није сортиран. Напиши
   програм који помаже наставници да добије ученике сортиране по
   абецедном реду.

Приликом сортирања, ниске се пореде по абецедном поретку (више о овом
поретку биће речено у делу о нискама). Напоменимо да овакво сортирање
може имати проблем са нашим карактерима ćđščž (потребно је навести
посебне опције функцији ``sort`` ако желимо да наш језик буде обрађен
потпуно исправно, али се тиме у овом приручнику нећемо бавити).

   
.. activecode:: лексикографско_сортирање_ниски

		
   ucenici = ["Ljubić Milenko", "Stojković Milica", "Vilimonović Aleksandar",
              "Jokić Đurađ", "Filipović Kalina", "Zlatković Jasmina"]
   print(sorted(ucenici))

   
Враћање торки из функције
&&&&&&&&&&&&&&&&&&&&&&&&&

Као што смо већ раније поменули, торке могу бити веома корисне да би
се вратило више вредности из функције. Подсетимо сеове корисне
технике.

Конверзија угла у степене и минуте
''''''''''''''''''''''''''''''''''
.. level:: 2
	   
.. questionnote::

  Напиши функцију која за угао дат у облику децималног броја степени
  одређује њему најближи угао дат у степенима и минутима. Употреби га
  да одредиш колико степени и минута има угао :math:`36,2^\circ`.

.. activecode:: функција_враћа_торку
  
   def ugao(alfa):
       # prevodimo ugao u minute i zaokruzujemo na najblizi ceo broj
       alfa_min = int(round(alfa * 60))
       # izdvajamo stepene i minute
       stepeni = alfa_min // 60
       minuta = alfa_min % 60
       # vracamo rezultat
       return (stepeni, minuta)

   (stepeni, minuta) = ugao(36.2);
   print(stepeni, ":", minuta)

Приметимо да се приликом прихватања резултата врши распакивање торке
тј. да се променљивима ``stepeni`` и ``minuta`` додељују редом први и
други елемент уређеног пара који је функција вратила.

   
Сложене торке и листе
&&&&&&&&&&&&&&&&&&&&&

Торке и листе могу да буду елементи других торки тј. листа. Направимо,
на пример, торку која садржи податке о месецима током једне године
која није преступна. За сваки месец знамо назив и број дана и те
податке ћемо организовати помоћу уређених парова (на пример,
``("април", 30)``).

Дани у месецу
'''''''''''''
.. level:: 2

.. activecode:: листа_торки

   meseci = (("јануар", 31), ("фебруар", 28), ("март", 31), \
             ("април", 30), ("мај", 31), ("јун", 30), \
             ("јул", 31), ("август", 31), ("септембар", 30), \
	     ("октобар", 31), ("новембар", 30), ("децембар", 31))
   broj = int(input("Унеси редни број месеца:"))
   mesec = meseci[broj - 1]
   print("Назив:", mesec[0], "Број дана:", mesec[1])

   
Помоћну променљиву ``mesec`` није било неопходно користити. Поправи
индексе у наредном програму тако да ради исто као и претходни.

.. activecode:: листа_торки_1

   meseci = (("јануар", 31), ("фебруар", 28), ("март", 31), \
             ("април", 30), ("мај", 31), ("јун", 30), \
             ("јул", 31), ("август", 31), ("септембар", 30), \
	     ("октобар", 31), ("новембар", 30), ("децембар", 31))
   broj = int(input("Унеси редни број месеца:"))
   # popravi indekse u narednom redu
   print("Назив:", meseci[0][0], "Број дана:", meseci[0][0])

   
Ниске
-----

Поред бројева, рачунари су веома добри и у раду са текстом и, као што
смо видели, за рад са текстом се користе ниске и променљиве типа
``str``. Иако се могу посматрати и као атомичке вредности (појединачни
подаци), ниске се могу посматрати и као торке карактера (слова,
интерпункцијских знакова и слично). Самим тим, на ниске се примењују
скоро све технике које се користе у раду са торкама. Нагласимо да су,
као и торке, и ниске имутабилне и да једном формирану ниску није
могуће променити. 

		   
Надовезивање ниски
&&&&&&&&&&&&&&&&&&

Слично као и над листама, над нискама се могу вршити одређене операције.
Једна од основних операција је спајање две ниске. Ова операција
донекле подсећа на сабирање и обележава се знаком ``+``.  На пример,
вредност израза ``"abraka" + "dabra"`` је ``"abrakadabra"``.  Још
једна интересантна операција је да се ниска помножи природним
бројем. И множење ниске бројем функционише исто као код торки
тј. листа. На пример, вредност израза ``"ba"*2`` је ``"baba"``.

.. mchoice:: ниске_надовезивање_множење
   :multiple_answers:
   :answer_a: "ba" + "na" * 2
   :answer_b: "a" + "na" * 2 + "s"
   :answer_c: "ba" * 2 + "na"
   :answer_d: "an" * 2 + "nas"
   :correct: a,b
   :feedback_a: Тачно! "banana"
   :feedback_b: Тачно! "ananas"
   :feedback_c: Нетачно! "babana"
   :feedback_d: Нетачно! "anannas"

   Који од наредних израза крију у себи исправно записан назив воћа?

Допуни наредни програм тако да исправно испише име и презиме (у овом
примеру треба да се испише ``Петар Петровић``).

.. activecode:: име_и_презиме

   ime = "Петар"
   prezime = "Петровић"
   ime_i_prezime = ime prezime         # ispravi ovaj red
   print(ime_i_prezime)

А шта је са цифрама у ниски?
&&&&&&&&&&&&&&&&&&&&&&&&&&&&

Ниске у себи могу да садрже цифре, па чак могу да буду састављене
искључиво од цифара. Међутим, када се на две ниске примени оператор
``+`` ниске се надовезују. Тако је резултат операције ``"12" + "34"``
једнак ``"1234"``, а не ``46`` како би неки очекивали. Помоћу
``input`` уноси се увек текст, тј. резултат ове операције је увек
ниска, чак иако тај текст садржи само цифре. Имајући ово у виду,
покушај да предвидиш шта ће израчунати наредни програм када, након
његовог покретања, корисник унесе 3, а затим и 5.

.. activecode:: погрешно_сабирање

   prvi_sabirak  = input("Unesi prvi sabirak: ")
   drugi_sabirak = input("Unesi drugi sabirak: ")
   zbir = prvi_sabirak + drugi_sabirak
   print(zbir)

.. mchoice:: ниске_надовезивање_множење_2
   :answer_a: 8
   :answer_b: 15
   :answer_c: "35"
   :answer_d: "8"
   :correct: c
   :feedback_a: Нетачно! Иако корисник куца цифре, ``prvi_sabirak`` и
                ``drugi_sabirak`` нису бројеви него ниске (текст) и
                зато се оператором ``+`` оне надовезују.
   :feedback_b: Нетачно! Покушај поново.
   :feedback_c: Тачно!
   :feedback_d: Нетачно! Иако корисник куца цифре, ``prvi_sabirak`` и
                ``drugi_sabirak`` нису бројеви него ниске (текст) и
                зато се оператором ``+`` оне надовезују.

   Који је резултат извршавања претходног програма ако корисник унесе
   прво ``3``, а затим ``5``.

Ако текст садржи само цифре, онда се број представљен тим цифрама може
добити помоћу ``int``. На пример, ``int("123")`` је број ``123``. Тако
је ``int("12") + int("34")`` једнако ``12 + 34`` тј.  ``46``. Стога се
учитавање броја може постићи помоћу ``int(input("Unesi broj:
"))``. Тако се претходни програм који сабира два учитана броја може
поправити на следећи начин:

.. activecode:: сабирање

   prvi_sabirak  = int(input("Unesi prvi sabirak: "))
   drugi_sabirak = int(input("Unesi drugi sabirak: "))
   zbir = prvi_sabirak + drugi_sabirak
   print(zbir)

Ниска може да садржи и децимални запис неког броја и тада се број
представљен том ниском може добити помоћу ``float``. На пример,
``float("123.45")`` je број ``123.45``. Претварање ниске у број је и у
овом случају веома важно урадити, јер се у супротном оператор ``+``
односи на надовезивање ниски, а не на сабирање бројева. Провери да ли
ово добро разумеш.

.. dragndrop:: ниске_и_бројеви
    :feedback: Покушај поново
    :match_1: float("3.5") + float("3.5")|||7.0
    :match_2: "3.5" + "0.5"|||"3.50.5"
    :match_3: 3.5 + "3.5"|||greška
    :match_4: float("2.5") + 1.5|||4.0
	      
Тако се децимални број може унети са ``broj = float(input("Unesi
decimalan broj:"))``.
   

Дужина ниске, издвајање делова ниске
&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&

Дужину ниске тј. број њених карактера можемо добити помоћу функције
``len``. Тако је ``len("Zdravo")`` једнако 6, јер ниска ``"Zdravo"``
има тачно 6 карактера.

.. fillintheblank:: fill_len_1

      Вредност ``len("Popokatepetl")`` је |blank|

      - :12: Тачно!
        :x: Изброј карактере у речи

.. fillintheblank:: fill_len_2
		    
      Вредност ``len("Супер Марио 3!")`` је |blank|

      - :14: Тачно!
        :x: Изброј карктере у речи рачунајући посебно и празнине и интерпункцијске знаке
         
Као и у другим торкама и листама и карактери у ниски имају своје редне
бројеве тј. позиције. Први карактер се налази на позицији 0, други на
позицији 1 и тако даље. Могуће је издвојити појединачни карактер из
ниске. На пример, ако је ``ime = "Zorana"`` тада се карактер ``Z``
може добити изразом ``ime[0]``, а карактер ``r`` изразом ``ime[2]``.

Подржани су и негативни индекси тако што -1 означава последњи карактер,
-2 претпоследњи и тако даље. На пример, ако је ``ime = "Zorana"`` тада
је ``ime[-1]`` карактер ``a`` док је ``ime[-4]`` карактер ``r``.

И издвајање дела ниске (подниске) функционише на исти начин као и код
других торки тј. листа. Подсети се овога.

.. fillintheblank:: fill_индекси_ниске
		    
      Дата је ниска ``s = "Programiranje je mnogo zabavno"``.
      Вредност израза ``s[0:4]`` је |blank|
    
      Израз којим је из ниске ``s`` могуће издвојити реч ``mnogo`` је |blank|


      - :Prog|"Prog": Тачно!
        :Progr: Карактер са последње наведене позиције се не узима
        :x: Изброј карактере у речи
      - :s\[(17\:22|-13\:-8)\]: Тачно!
        :x: Пажљиво преброј позиције. Можеш да користиш било позитивне индексе (слева на десно), било негативне индексе (сдесна на лево).

И наредни задатак се могу урадити коришћењем издвајања делова ниске.
   
ЈМБГ
''''
.. level:: 1
   
.. questionnote::

   Сваки грађанин Републике Србије има свој јединствени матични број
   (ЈМБГ).  У њему прве две цифре одређују дан рођења, друге две
   месец, а наредне три цифре одређују годину рођења. Наредне две
   цифре одређују општину рођења, наредне три цифре су јединствене за
   ту особу при чему се из њих може одредити пол (комбинације од 000
   до 499 се додељују дечацима, а од 500 до 999 девојчицама). Последња
   цифра је контролна. Она се израчунава применом једне посебне
   формуле на претходне цифре. За дати ЈМБГ одредите ког дана и месеца
   се особа родила.

Иако делује да је ЈМБГ број, њега је боље посматрати као ниску
карактера тј.  стринг. Наиме, обично смо заинтересовани само за
издвајање одређених делова ЈМБГ, док на ЈМБГ никада не примењујемо
аритметичке операције (нпр. нема смисла сабирати два ЈМБГ нити ЈМБГ
множити са 2). Издвајање делова је веома једноставно ако ЈМБГ
представимо у облику ниске. 

.. activecode:: jmbg

   jmbg  = "1702992850011"
   dan   = jmbg[0:2]     # izdvajamo tekst sa pozicija 0 i 1
   mesec = ""            # izdvajamo tekst sa pozicija 2 i 3 - popravi ovaj red
   print("Rodjen si " + dan + "." + mesec + ".")
            
Претрага ниске
&&&&&&&&&&&&&&

Често је потребно да проверимо да ли једна ниска садржи неки карактер
или садржи неку другу ниску. То можемо урадити коришћењем ``find``.
На пример, ако је ``ime_i_prezime = "Љубица Љубичић"`` тада је
вредност ``ime_i_prezime.find(" ")`` једнака ``6``, јер се размак може
наћи на позицији 6. Ако потражимо карактер који ниска не садржи (на
пример, ``ime_i_prezime.find(",")``), добићемо вредност ``-1``.
Уместо ``find`` можемо употребити и ``index`` (као и код торки и
листа), међутим, када се оно што се тражи не налази у ниски, ``find``
враћа вредност ``-1``, док ``index`` даје грешку при извршавању
програма.

Претрагу можемо употребити да бисмо, на пример, издвојили име и
презиме из датог имена и презимена (што је управо супротно од задатка
у којем смо на основу посебног имена и презимена и добили спојено име
и презиме).

.. activecode:: издвајање_имена_и_презимена

   ime_i_prezime = "Љубица Љубичић"
   razmak = ime_i_prezime.find(" ")
   ime = ime_i_prezime[0:razmak]
   prezime = ime_i_prezime[razmak+1:]
   print("Име: ", ime)
   print("Презиме: ", prezime)

Овај задатак има и лепше решење које користи функцију ``split``, која
формира листу речи дате реченице (врши поделу ниске на подниске на
основу размака који се јављају унутар ниске).

.. activecode:: издвајање_имена_и_презимена_split

   ime_i_prezime = "Љубица Љубичић"
   (ime, prezime) = ime_i_prezime.split()
   print("Име: ", ime)
   print("Презиме: ", prezime)

   
Осим што можемо пронаћи позицију на којој се неко слово или нека
подниска јавља унутар ниске, можемо и избројати колико се пута оно
појављује. За то можемо користити ``count``. На пример, број
појављивања слова ``a`` у неком имену можемо израчунати помоћу
``ime.count("a")``. Наредни задатак се веома једноставно може урадити
применом методе ``count``.

Врсте реченица
''''''''''''''
.. level:: 1

   
.. questionnote::

   Дат је текст (ниска) који садржи неколико једноставних реченица.
   Одреди колико међу њима има обавештајних, колико има упитних и
   колико има узвичних реченица.

Обавештајне реченице се завршавају тачком, упитне знаком питања, а
узвичне узвичником. Број реченица сваке врсте можемо одредити тако што
пребројимо колико се пута јавља сваки од ова три карактера (карактер
тачка ``.``, карактер упитник ``?`` и карактер узвичник ``!``).

.. activecode:: врсте_реченица   

   tekst = "Zdravo! Ja sam Marko! Kako se ti zoveš? Koji si ti razred? Ja sam šesti."
   print("Obaveštajnih:", tekst.count("."))
   print("Upitnih:", 0)         # ispravi ovaj red
   print("Uzvičnih:", 0)        # ispravi ovaj red

Измене листе
------------

Торке и ниске су имутабилни типови података, што значи да се, када се
једном формирају, они не могу даље мењати. Са друге стране, листе је
могуће мењати. Појединачни елементи листе се могу мењати на следећи
начин.

.. activecode:: измена_елемената_листе

   boje = ["red", "green", "blue"]
   boje[0] = "purple"
   boje[2] = "orange"
   print(boje)

На првом месту у листи црвену боју смо заменили љубичастом, а на
последњем плаву са нарандџастом, тако да на крају програма листа
``boje`` чува вредности ``["purple", "green", "orange"]``.

Елемент на крај листе можемо додати методом ``append``. На пример,
уместо да мењамо, листу боја можемо да проширимо за две вредности.

.. activecode:: додавање_елемената_у_листу

   boje = ["red", "green", "blue"]
   boje.append("purple")
   boje.append("orange")
   print(boje)

На крају овог програма, листа садржи вредности ``["red", "green",
"blue", "purple", "orange"]``.

Елементе из листе можемо брисати помоћу ``del``. На пример, ако из
листе дана желимо да обришемо дане викенда, то можемо урадити на
наредни начин.

.. activecode:: брисање_елемената_из_листе

   dani = ["nedelja", "ponedeljak", "utorak", "sreda", "četvrtak", "petak", "subota"]
   del dani[6]
   del dani[0]
   print(dani)


Скупови
-------

Ако знамо имена свих фудбалера који су давали голове током једне
ревијалне утакмице и то редом, како су давали голове можемо
једноставно одредити скуп свих фудбалера који су се уписали у листу
стрелаца. Наиме, ако формирамо скуп на основу имена стрелаца
аутоматски ће бити уклоњени дупликати.

.. activecode:: скуп
		
   strelci = {"Mesi", "Ronaldo", "Mesi", "Ibrahimović", "Ibrahimović", "Nejmar", "Nejmar"}
   print(strelci)

Скуп дефинишемо тако што између витичастих заграда (заграда ``{`` и
``}``) наводимо елементе раздвојене зарезима. Елементи скупа у овом
примеру су биле ниске, а можемо разматрати и скупове елемената другог
типа (на пример, скупове бројева).

Скуп од торке/листе/ниске
&&&&&&&&&&&&&&&&&&&&&&&&&

Ако су елементи скупа чланови неке торке или листе, можемо лако добити
скуп свих елемената (коришћењем функције ``set`` као у следећем
програму).

.. activecode:: скуп_од_листе
		
   golovi = ["Mesi", "Ronaldo", "Mesi", "Ibrahimović", "Ibrahimović", "Nejmar", "Nejmar"]
   strelci = set(golovi)
   print(strelci)

И од ниске можемо једноставно добити скуп карактера које она садржи
(опет помоћу функције ``set``).

Провера припадности скупу
&&&&&&&&&&&&&&&&&&&&&&&&&

Слично као и код осталих структура података, провера припадности скупу
може се извршити оператором ``in``. Покушај да то примениш у наредном
задатку.

Да ли је Миодраг награђен?
''''''''''''''''''''''''''

.. questionnote:: 

   Познат је скуп награђених ученика. Напиши програм који проверава да
   ли је Миодраг освојио награду.

.. activecode:: члан_скупа_награђених

   nagradjeni = {"Lenka", "Sofija", "Dimitrije"} 
   if "Miodrag" ???:
       print("Bravo!")
   else:                       
       print("Nažalost, Miodrag nije među nagrađenima.")


Скуповне операције
&&&&&&&&&&&&&&&&&&

У језику Python можемо веома једноставно израчунати **унију**,
**пресек** и **разлику** скупова. Ако су ``A`` и ``B`` скупови, онда
је ``A | B`` њихова унија, ``A & B`` њихов пресек, а ``A - B`` њихова
разлика. Провери да ли се сећаш ових операција из математике тако што
ћеш решити наредни тест.


.. dragndrop:: identifier
    :match_1: A & B|||{3}
    :match_2: A | B|||{3, 4, 5, 6, 7}
    :match_3: A - B|||{6, 7}
    :match_4: B - A|||{4, 5}
    :match_5: (A - B) | (B - A)|||{4, 5, 6, 7}

    Ако је A = {3, 6, 7}, а B = {3, 4, 5}, повежи скуповне
    операције са њиховим резултатима.


Прикажимо како се скуповне операције могу употребити у решењу наредног
задатка.

Девојчице које тренирају два спорта
'''''''''''''''''''''''''''''''''''
.. level:: 1
   
.. questionnote::

  Један скуп садржи девојчице из одељења које тренирају ритмичку
  гимнастику, а други оне које тренирају одбојку. Одреди скуп
  девојчица које тренирају оба спорта, скуп девојчица које тренирају
  бар један од њих и скуп девојчица које тренирају само одбојку.

.. activecode:: скупови_спортисткиња  

   ritmicka = {"Ana", "Milica", "Jovana", "Gordana"}
   odbojka = {"Tara", "Nađa", "Milica", "Jovana", "Aleksandra"}
   dva_sporta = ritmicka ??? odbojka       # upiši odgovarajući operator
   bar_jedan_sport = ritmicka ??? odbojka  # upiši odgovarajući operator
   samo_odbojka = odbojka ??? ritmicka     # upiši odgovarajući operator
   print(dva_sporta)
   print(bar_jedan_sport)
   print(samo_odbojka)

Ако све урадиш како треба, добићеш решење

::

   {'Milica', 'Jovana'}
   {'Ana', 'Milica', 'Jovana', 'Gordana', 'Tara', 'Nađa', 'Aleksandra'}
   {'Tara', 'Nađa', 'Aleksandra'}   

Речници
-------

Листе нам дају могућност да организујемо податке за које је познат
неки редослед и у којима сваки податак има свој редни број (на пример,
листу смо употребили да организујемо имена путника у авиону или
податке о месецима у години) и приступ елементима листе врши се на
основу њиховог индекса (редног броја, тј. позиције). Међутим, често
подаци представљају неко пресликавање у којем на основу датог *кључа*
желимо да приступимо *вредности* која је придружена том кључу. На
пример, на основу шифре производа приступамо његовој цени, на основу
ЈМБГ грађана приступамо његовом имену, на основу броја банковног
рачуна приступамо салду (износу новца) на том рачуну и слично. У тим
ситуацијама уместо листе много је боље користити **речнике** (каже се
и мапе или асоцијативне низове).

Дефинисање и коришћење речника
&&&&&&&&&&&&&&&&&&&&&&&&&&&&&&

Цене аутомобила
'''''''''''''''
.. level:: 1

.. questionnote::

   Аутомобили у каталогу имају придружене цене и ми желимо да у нашем
   програму можемо да одредимо цену аутомобила на основу његовог
   модела. Напиши програм који на основу унетог модела аутомобила
   (ниска) одређује његову цену (цео број).

.. image:: ../../_images/sajam_automobila.jpg
   :width: 300px   
   :align: center

Кључна одлука за ефикасно решење овог задатка је да каталог цена
аутомобила представимо помоћу речника.

.. activecode:: цене_аутомобила
		
   cene_automobila = {"fiat 500l": 11990, 
                      "renault clio": 9650, 
                      "toyota corolla": 13990}
   automobil = input("Unesi model automobila:")
   print(cene_automobila[automobil])

У претходном примеру модели аутомобила су кључеви који се пресликавају
у цене аутомобила које су вредности. Кључеви су биле ниске, а
вредности цели бројеви, што је чест случај.


Географске координате градова
'''''''''''''''''''''''''''''
.. level:: 2

У наредном примеру вредности у речнику су уређени парови.
	   
.. questionnote::

   Познате су географске координате неколико главних европских
   градова. За дато име града одреди њене географске
   координате. Одреди посебно географску дужину и посебно географску
   ширину.

.. activecode:: географске_координате
		
   gradovi = {"Beograd": (44.7866, 20.4489), 
              "Budimpešta": (47.4979, 19.0402),
              "Beč": (48.2082, 16.3738), 
              "Bratislava": (48.1486, 17.1077)}
   grad = input("Unesi ime grada: ")
   # ispravi naredni red tako da se iz rečnika pročitaju koordinate grada
   koordinate = ???
   print(koordinate)
   # dopuni naredni kod tako da se ispišu geografska širina i dužina
   print("Geografska širina: ")
   print("Geografska dužina: ")


Цена производа
''''''''''''''

Слично као и код осталих структура података и провера да ли је неком
кључу у речнику придружена вредност се може проверити оператором
``in``.

.. level:: 1
      
.. questionnote::

   У речнику се чувају цене различитих производа у продавници. Пронађи
   и испиши цену датог производа или пријави да се тај производ не
   продаје.

.. activecode:: припада_речнику
		
   cene = {"hleb": 39.90, "mleko": 89, "jaje": 14.90}
   proizvod = "hleb"
   if proizvod in cene:
       print(proizvod, "košta", cene[proizvod], "dinara")
   else:
       print(proizvod, "se ne prodaje u ovoj prodavnici")
   
Домаћи задатак
--------------

Ова лекција је била прилично обимна, па је могуће да нисте током часа
успели да је целу пређете. За домаћи пређи све задатке које су на часу
остали неурађени, а затим уради и наредних неколико задатака.

Задаци
&&&&&&


Да ли си награђен?
''''''''''''''''''
.. level:: 1

.. questionnote:: 

   Позната је листа награђених ученика. Напиши програм који проверава
   да ли је корисник чије се име учитава на почетку програма освојио
   награду.

.. activecode:: члан_листе_награђених

   nagradjeni = ["Lenka", "Sofija", "Dimitrije"]
   ime = ""    # ispravi ovaj red tako da se unosi ime ucenika
   if ???:     # dodaj uslov koji proverava da li se uneto ime nalazi u listi
       print("Bravo, nagrada je tvoja!")
   else:
       print("Nažalost, nisi među nagrađenima")

Путници на датим седиштима у авиону
'''''''''''''''''''''''''''''''''''
.. level:: 1

.. questionnote::

  Познат је списак имена путника у авиону. Седишта су нумерисана од 1
  па надаље. Ако стјуардеса унесе број седишта, напиши програм који
  одређује име путника на том седишту. Након тога испиши имена путника
  која седе на прва четири, као и на последња два седишта, као и број
  седишта на којем седи Мика Микић.

Да бисмо олакшали бројање, на првом месту у листи је постављен
вештачки елемент (празна ниска).
  
.. activecode:: путници_у_авиону

  putnici = ["", "Петар Петровић", "Зорана Зорановић", "Мика Микић", \
             "Јелена Јеленковић", "Ана Анић", "Милан Милановић"]
  # ispravi naredne redove tako da se ispise ime putnika na unetom broju sedista
  sediste = int(input("Unesi broj sedista:"))
  print("Na sedistu", sediste, "sedi", putnici[0])
  # ispravi naredni red tako da se ispisu imena prva 4 putnika
  print(putnici[0])
  # ispravi naredni red tako da se ispisu imena poslednja 2 putnika
  print(putnici[0])
  # ispravi naredni red tako da se ispise broj sedista na kojem sedi Мика Микић
  print(putnici[0])

Шифра
'''''
.. level:: 1
   
.. questionnote::

   Лидијина шифра за gmail се састоји од 6 слова. Одлучила је да је
   закомпликује, али на правилан начин, тако да касније може да је се
   сети. Шифру ће закомпликовати тако што ће заменити редослед прва
   два и последња два слова, а средња два слова поновити два пута. На
   пример, ако јој је шифра била lidija закомпликована шифра јој је
   jadidili. Напиши програм који исписује закомпликовану шифру.

.. activecode:: закомпликована_шифра   

   sifra = "lidija"
   zakomplikovana_sifra = sifra[4:6] + ""  # popravi ovaj red
   print(zakomplikovana_sifra)

Учитавање разломка
''''''''''''''''''
.. level:: 2
		
.. questionnote::

   Корисник уноси текст који описује разломак (на пример,
   `3/4`). Израчунај децималну вредност тог разломка.


.. activecode:: разломак

   razlomak = input("Unesi razlomak: ")
   pozicija_crte = razlomak.find("/")
   brojilac = int(0)    # popravi ovaj red
   imenilac = int(razlomak[pozicija_crte+1:])
   print(0)   # popravi ovaj red

Сугласници
''''''''''
.. level:: 1

.. questionnote::
   
   Сугласници се по месту изговора могу бити предњонепчани (Ј, Љ, Њ, Ћ,
   Ђ, Ш, Ж, Ч, Џ) и задњонепчани К, Г, Х. По звучности, сугласници се
   деле на звучне (Б, Г, Д, Ђ, Ж, З, Џ), безвучне (П, К, Т, Ћ, Ш, С, Ч,
   Ф, Х, Ц), и сонанте(М, В, Р, Л, Н, Љ, Њ, Ј).

   - Који сугласници су истовремено и звучни и предњонепчани?
   - Који безвучни сугласници нису задњонепчани?
   - Који сугласници су сонанти или звучни?

Поправи наредни код тако да исправно решава овај задатак. Програм
након покретања треба да испише наредне резултате.
     
::      

  {'Ђ', 'Џ', 'Ж'}
  {'Ч', 'Ћ', 'Ш', 'Т', 'П', 'Ф', 'С', 'Ц'}
  {'Џ', 'З', 'Д', 'Ђ', 'Ј', 'Л', 'Н', 'М', 'Г', 'Р', 'Б', 'Љ', 'Њ', 'Ж', 'В'}

.. activecode:: скупови_сугласника
		
   zadnjonepcani  = {"К", "Г", "Х"}
   prednjonepcani = {"Ј", "Љ", "Њ", "Ћ", "Ђ", "Ш", "Ж", "Ч", "Џ"}
   zvucni   = {"Б", "Г", "Д", "Ђ", "Ж", "З", "Џ"}
   bezvucni = {"П", "К", "Т", "Ћ", "Ш", "С", "Ч", "Ф", "Х", "Ц"}
   sonanti  = {"М", "В", "Р", "Л", "Н", "Љ", "Њ", "Ј"}

   zvucni_prednjonepcani = {}           # popravi ovaj red
   bezvucni_ne_zadnjonepcani = {}       # popravi ovaj red
   zvucni_ili_sonanti = {}              # popravi ovaj red
   print(zvucni_prednjonepcani)
   print(bezvucni_ne_zadnjonepcani)
   print(zvucni_ili_sonanti)
   
Звучни и безвучни сугласници
''''''''''''''''''''''''''''
.. level:: 1
   
.. questionnote::

   Звучни и безвучни сугласници јављају се у паровима.  Звучни
   сугласници су БГДЂЖЗЏ (можеш их лакше запамтити помоћу првих слова
   реченице Баба грди деду: "Ђаволе живи зашто џандрљаш"?) Њихови
   безвучни парови су редом ПКТЋШСЧ. Напиши програм који за дати
   звучни сугласник одређује његов безвучни пар.

Поправи наредни програм тако што ћеш попунити речник одговарајућим
паровима сугласника. Ако све урадиш како треба, требало би да добијеш
7 зелених поља (сваки сугласник мора да прође тест). Користи велика
слова ћирилице.
   
.. activecode:: безвучни_звучни_парови
   :nocodelens:
      
   def bezvucni(zvucni):
       bezvucni_par = {}    # поправи овај ред
       return bezvucni_par[zvucni]
       
   ====
   from unittest.gui import TestCaseGui

   class myTests(TestCaseGui):

       def testOne(self):
          self.assertEqual(bezvucni("Б"), "П", "Безвучни пар гласа Б је глас П.")
          self.assertEqual(bezvucni("Г"), "К", "Безвучни пар гласа Г је глас К.")
          self.assertEqual(bezvucni("Д"), "Т", "Безвучни пар гласа Д је глас Т.")
          self.assertEqual(bezvucni("Ђ"), "Ћ", "Безвучни пар гласа Ђ је глас Ћ.")
          self.assertEqual(bezvucni("Ж"), "Ш", "Безвучни пар гласа Ж је глас Ш.")
          self.assertEqual(bezvucni("З"), "С", "Безвучни пар гласа З је глас С.")
          self.assertEqual(bezvucni("Џ"), "Ч", "Безвучни пар гласа Џ је глас Ч.")

   myTests().main()

Главни градови
''''''''''''''
.. level:: 1

.. questionnote::

   У речнику се чувају главни градови неколико држава. Напиши програм
   који учитава име државе и исписује главни град или каже да одговор
   на то питање не зна.

.. activecode:: главни_град

   glavni_grad = {"Srbija": "Beograd",
                  "Mađarska": "Budimpešta",
                  "Rumunija": "Bukurešt",
                  "Bugarska": "Sofija",
                  "Makedonija": "Skoplje",
                  "Albanija": "Tirana",
                  "Crna Gora": "Podgorica",
                  "Bosna i Hercegovina": "Sarajevo",
                  "Hrvatska": "Zagreb"}
   # završi ovaj program
   
